﻿---
id: 66ad8294a0ad902f1b31b612
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将通过构建一个问候机器人来学习如何使用JavaScript基础知识。

在第一步中，你需要从问候机器人向控制台输出一条消息。

请记住，你在之前的课程中学习了 `console.log()` 和字符串。

以下是使用 `console.log()` 与字符串的提醒：

```js
console.log("Hello, World!");
```

添加一个 `console.log()` 语句，将字符串 `"Hi there!"` 输出到控制台。不要忘记消息前后的引号。


# --hints--

你的代码中应该有一个 `console.log()` 语句。

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

你的 `console` 语句应该有消息 `"Hi there!"`。不要忘记引号。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```