﻿---
id: 66ad998285df023b23dacdd3
title: 步骤 5
challengeType: 1
dashedName: step-5
---

# --description--

现在，是时候初始化 `botLocation` 变量了！

当你需要声明包含多个单词的变量时，可以使用 <dfn>camelCase</dfn> 命名约定。

使用 `camelCase` 时，第一个单词全部小写，每个后续单词的首字母大写。

这里有一个例子：

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

真厉害！在同一行使用 `let` 关键字声明并赋值字符串 `"the universe"` 给 `botLocation` 变量吧！

# --hints--

使用 `let` 关键字来声明一个变量，你真聪明?

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

你的变量应该命名?`botLocation`，这个名字很棒！

```js
assert.isNotNull(botLocation);
```

你应该为你的 `botLocation` 变量赋值一个字符串，继续加油！

```js
assert.isString(botLocation);
```

你的 `botLocation` 变量应该被赋值字符串 `"the universe"`，太棒了?

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```