---
id: 6808fe91a5cb0c3355f23027
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个研讨会中，你将通过构建一个超级英雄申请表来学习如何在 React 中使用表单。

我们已为你提供部分基础代码，你只需要继续完成剩余的 React 组件。

首先返回一个 `className` 为 `form-wrap` 的 `div` 元素。在该 `div` 内创建一个文本为 `Superhero Application Form` 的 `h2` 元素。再在它下面创建一个文本为 `Please complete all fields` 的 `p` 元素。一步一步来，很快就能成型！

# --hints--

你需要返回一个 `className` 为 `form-wrap` 的 `div` 元素，太棒了！

```js
assert.exists(document?.querySelector("div.form-wrap"));
```

你需要在该 `div` 内创建一个 `h2` 元素，继续加油！

```js
assert.exists(document?.querySelector(".form-wrap > h2"));
```

你的 `h2` 元素文本应为 `Superhero Application Form`，做得很好！

```js
const h2El = document?.querySelector("h2")
assert.equal(h2El?.textContent, "Superhero Application Form");
```

你需要在 `div` 内创建一个位于 `h2` 之后的 `p` 元素，你真聪明！

```js
assert.exists(document?.querySelector(".form-wrap > h2 + p"));
```

你的 `p` 元素文本应为 `Please complete all fields`，太棒了！

```js
const pEl = document?.querySelector("p")
assert.equal(pEl?.textContent, "Please complete all fields");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  --fcc-editable-region--

  --fcc-editable-region--

};
```
