---
id: 63c21839f56eaf3ef4e027c4
title: 第50步
challengeType: 0
dashedName: step-50
---

# --description--

在你的[HTMLString](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L247-L250)末尾创建另一个[label](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L248-L248)元素（在新行上）。这个[label](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L248-L248)应该有文本"条目 # 卡路里"，使用你的模板字面量语法将`#`替换为[entryNumber](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L246-L246)的值，并将[for](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225)属性设置为`X-#-calories`，其中`X`是[entryDropdown](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L227-L227)的[value](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/670ea2c4f42def724c77e3c7.md#L225-L225)，`#`是[entryNumber](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L246-L246)的值。

# --hints--

你的[HTMLString](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L247-L250)中应该有两个[label](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L248-L248)元素。

```js
const HTMLstring = code.split(/HTMLString\s*=/)[1];
assert.equal(HTMLstring.match(/<label/g).length, 2);
```

你的新[label](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L248-L248)元素应该在新行上。

```js
const HTMLstring = code.split(/HTMLString\s*=/)[1];
assert.equal(HTMLstring.match(/\n\s*<label/g).length, 2);
```

你的新[label](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L248-L248)元素应该在你的[input](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c2187f55eb0f400269568f.md#L214-L218)元素之后。

```js
const HTMLstring = code.split(/HTMLString\s*=/)[1];
const inputIndex = HTMLstring.indexOf("<input");
const labelIndex = HTMLstring.lastIndexOf("<label");
assert.isBelow(inputIndex, labelIndex);
```

你的新[label](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L248-L248)元素应该有一个[for](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225)属性设置为`${entryDropdown.value}-${entryNumber}-calories`。

```js
const HTMLstring = code.split(/HTMLString\s*=/)[1];
const label = HTMLstring.match(/<label.*>.*<\/label>/g)[1];
assert.match(label, /<label\s+for="\$\{entryDropdown.value\}-\$\{entryNumber\}-calories"\s*>/)
```

你的新[label](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L248-L248)元素应该有文本`Entry ${entryNumber} Calories`。

```js
const HTMLstring = code.split(/HTMLString\s*=/)[1];
const label = HTMLstring.match(/<label.*>.*<\/label>/g)[1];
assert.match(label, /<label\s+for="\$\{entryDropdown.value\}-\$\{entryNumber\}-calories"\s*>Entry\s\$\{entryNumber\}\sCalories<\/label>/);
```

你不应该修改现有的元素。

```js
const HTMLstring = code.split(/HTMLString\s*=/)[1];
assert.match(HTMLstring, /`\n\s*<label\s+for="\$\{entryDropdown.value\}-\$\{entryNumber\}-name"\s*>Entry\s\$\{entryNumber\}\sName<\/label>\n\s*<input\stype="text"\sid="\$\{entryDropdown.value\}-\$\{entryNumber\}-name"\splaceholder="Name"\s\/>\n\s*<label\s+for="\$\{entryDropdown.value\}-\$\{entryNumber\}-calories"\s*>Entry\s\$\{entryNumber\}\sCalories<\/label>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>卡路里计算器</title>
  </head>
  <body>
    <main>
      <h1>卡路里计算器</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">预算</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="每日卡路里预算"
            required
          />
          <fieldset id="breakfast">
            <legend>早餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>午餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>晚餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>零食</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>运动</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">添加食物或运动：</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>早餐</option>
                <option value="lunch">午餐</option>
                <option value="dinner">晚餐</option>
                <option value="snacks">零食</option>
                <option value="exercise">运动</option>
              </select>
              <button type="button" id="add-entry">添加条目</button>
            </span>
          </div>
          <div>
            <button type="submit">
              计算剩余卡路里
            </button>
            <button type="button" id="clear">清除</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

function cleanInputString(str) {
  const regex = /[+-\s]/g;
  return str.replace(regex, '');
}

function isInvalidInput(str) {
  const regex = /\d+e\d+/i;
  return str.match(regex);
}

--fcc-editable-region--
function addEntry() {
  const targetInputContainer = document.querySelector(`#${entryDropdown.value} .input-container`);
  const entryNumber = targetInputContainer.querySelectorAll('input[type="text"]').length;
  const HTMLString = `
  <label for="${entryDropdown.value}-${entryNumber}-name">条目 ${entryNumber} 名称</label>
  <input type="text" id="${entryDropdown.value}-${entryNumber}-name" placeholder="名称" />
  <label for="${entryDropdown.value}-${entryNumber}-calories">条目 ${entryNumber} 卡路里</label>`;
}
--fcc-editable-region--
```