---
id: 63c2194dce265f429300c8b1
title: 第54步
challengeType: 0
dashedName: step-54
---

# --description--

试着在"早餐"分类中添加几个条目，你可能会注意到一些小问题呢！首先我们需要修复的是条目计数——第一个条目应该是`1`，而不是`0`。

出现这个小bug是因为你在将新条目添加到页面*之前*就查询了`input[type="text"]`元素。要解决这个问题，你需要更新你的[entryNumber](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L208-L208)变量，将其设置为查询结果的[length](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225)值加`1`。请在你的声明行上添加这个操作，而不是在模板字符串中添加哦！

# --hints--

你应该在[querySelectorAll()](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L208-L208)方法的[length](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225)后面加上`1`。

```js
assert.match(addEntry.toString(), /entryNumber\s*=\s*targetInputContainer\.querySelectorAll\(\s*'\s*input\s*\[\s*type\s*=\s*"text"\s*]\s*'\s*\)\.length\s*\+\s*1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>卡路里计算器</title>
  </head>
  <body>
    <main>
      <h1>卡路里计算器</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">预算</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="每日卡路里预算"
            required
          />
          <fieldset id="breakfast">
            <legend>早餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>午餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>晚餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>零食</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>运动</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">添加食物或运动：</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>早餐</option>
                <option value="lunch">午餐</option>
                <option value="dinner">晚餐</option>
                <option value="snacks">零食</option>
                <option value="exercise">运动</option>
              </select>
              <button type="button" id="add-entry">添加条目</button>
            </span>
          </div>
          <div>
            <button type="submit">
              计算剩余卡路里
            </button>
            <button type="button" id="clear">清除</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

function cleanInputString(str) {
  const regex = /[+-\s]/g;
  return str.replace(regex, '');
}

function isInvalidInput(str) {
  const regex = /\d+e\d+/i;
  return str.match(regex);
}

--fcc-editable-region--
function addEntry() {
  const targetInputContainer = document.querySelector(`#${entryDropdown.value} .input-container`);
  const entryNumber = targetInputContainer.querySelectorAll('input[type="text"]').length + 1;
  const HTMLString = `
  <label for="${entryDropdown.value}-${entryNumber}-name">条目 ${entryNumber} 名称</label>
  <input type="text" id="${entryDropdown.value}-${entryNumber}-name" placeholder="名称" />
  <label for="${entryDropdown.value}-${entryNumber}-calories">条目 ${entryNumber} 卡路里</label>
  <input
    type="number"
    min="0"
    id="${entryDropdown.value}-${entryNumber}-calories"
    placeholder="卡路里"
  />`;
  targetInputContainer.innerHTML += HTMLString;
}
--fcc-editable-region--

addEntryButton.addEventListener("click", addEntry);
```