---
id: 63c9bce376ca4f09c15a3768
title: 第66步
challengeType: 0
dashedName: step-66
---

# --description--

记住，[return](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L259-L259)会结束函数的执行。在你的[if](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L256-L260)代码块之后，你需要处理输入有效时的逻辑。因为你的[if](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L256-L260)语句会返回一个值，所以你不需要[else](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L256-L260)语句。

使用加法赋值运算符将[currVal](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L253-L253)添加到你的[calories](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L250-L250)总数中。你需要使用[Number](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L261-L261)构造函数将[currVal](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L253-L253)转换为数字。

[Number](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L261-L261)构造函数是一个将值转换为数字的函数。如果值无法转换，它会返回[NaN](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L261-L261)，表示"不是数字"。

这是一个示例：

```js
Number('10'); // 返回数字 10
Number('abc'); // 返回 NaN
```

# --hints--

你不应该添加[else](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L256-L260)语句。

```js
assert.notMatch(getCaloriesFromInputs.toString(), /else/);
```

在你的[if](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L256-L260)语句之后，你应该对[calories](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L250-L250)使用加法赋值运算符。

```js
assert.match(getCaloriesFromInputs.toString(), /if.*}\s*calories\s*\+=/s);
```

你应该将[currVal](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L253-L253)传递给[Number()](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L261-L261)构造函数。

```js
assert.match(getCaloriesFromInputs.toString(), /Number\(\s*currVal\s*\)/);
```

你应该将[Number(currVal)](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L261-L261)的值添加到[calories](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9bce376ca4f09c15a3768.md#L250-L250)中。

```js
assert.match(getCaloriesFromInputs.toString(), /calories\s*\+=\s*Number\(\s*currVal\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>卡路里计算器</title>
  </head>
  <body>
    <main>
      <h1>卡路里计算器</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">预算</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="每日卡路里预算"
            required
          />
          <fieldset id="breakfast">
            <legend>早餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>午餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>晚餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>零食</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>运动</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">添加食物或运动：</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>早餐</option>
                <option value="lunch">午餐</option>
                <option value="dinner">晚餐</option>
                <option value="snacks">零食</option>
                <option value="exercise">运动</option>
              </select>
              <button type="button" id="add-entry">添加条目</button>
            </span>
          </div>
          <div>
            <button type="submit">
              计算剩余卡路里
            </button>
            <button type="button" id="clear">清除</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

function cleanInputString(str) {
  const regex = /[+-\s]/g;
  return str.replace(regex, '');
}

function isInvalidInput(str) {
  const regex = /\d+e\d+/i;
  return str.match(regex);
}

function addEntry() {
  const targetInputContainer = document.querySelector(`#${entryDropdown.value} .input-container`);
  const entryNumber = targetInputContainer.querySelectorAll('input[type="text"]').length + 1;
  const HTMLString = `
  <label for="${entryDropdown.value}-${entryNumber}-name">条目 ${entryNumber} 名称</label>
  <input type="text" id="${entryDropdown.value}-${entryNumber}-name" placeholder="名称" />
  <label for="${entryDropdown.value}-${entryNumber}-calories">条目 ${entryNumber} 卡路里</label>
  <input
    type="number"
    min="0"
    id="${entryDropdown.value}-${entryNumber}-calories"
    placeholder="卡路里"
  />`;
  targetInputContainer.insertAdjacentHTML('beforeend', HTMLString);
}

function getCaloriesFromInputs(list) {
  let calories = 0;

  for (const item of list) {
    const currVal = cleanInputString(item.value);
    const invalidInputMatch = isInvalidInput(currVal);

    if (invalidInputMatch) {
      alert(`无效输入: ${invalidInputMatch[0]}`);
      isError = true;
      return null;
    }
    calories += Number(currVal);
  }
  return calories;
}

addEntryButton.addEventListener("click", addEntry);