---
id: afcc8d540bea9ea2669306b6
title: 构建字符串重复函数
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

在这个实验中，你将创建一个函数，该函数将给定的字符串重复特定次数。为了完成这个实验，_不要_使用内置的 [.repeat()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L25-L25) 方法。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 [repeatStringNumTimes](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 的函数，该函数接受两个参数：一个字符串和一个数字。
2. 该函数应返回指定次数重复的字符串。
3. 如果数字小于或等于零，函数应返回空字符串。

# --hints--

你应该创建一个名为 [repeatStringNumTimes](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 的函数。

```js
assert.isFunction(repeatStringNumTimes);
```

[repeatStringNumTimes](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应该接受两个参数。

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

函数 [repeatStringNumTimes](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应始终返回一个字符串。

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

[repeatStringNumTimes("*", 3)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应返回字符串 [***](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102)。

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

[repeatStringNumTimes("abc", 3)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应返回字符串 [abcabcabc](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102)。

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

[repeatStringNumTimes("abc", 4)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应返回字符串 [abcabcabcabc](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102)。

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

[repeatStringNumTimes("abc", 1)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应返回字符串 [abc](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102)。

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

[repeatStringNumTimes("*", 8)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应返回字符串 [********](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102)。

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

[repeatStringNumTimes("abc", -2)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应返回空字符串 ("")。

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

[repeatStringNumTimes("abc", 0)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-repeat-a-string/afcc8d540bea9ea2669306b6.md#L99-L102) 应返回 ""。

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

不应使用内置的 [repeat()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L25-L25) 方法。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
