---
id: 587d78af367417b2b2512b03
title: æž„å»ºä¸€ä¸ªè°ƒæŸ¥è¡¨å?
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
---

# --description--

**ç›®æ ‡ï¼?*å®žçŽ°ä»¥ä¸‹ç”¨æˆ·éœ€æ±‚å¹¶é€šè¿‡æ‰€æœ‰æµ‹è¯•ä»¥å®Œæˆå®žéªŒã€?

**ç”¨æˆ·éœ€æ±‚ï¼š**

1. ä½ åº”è¯¥æŠŠé¡µé¢æ ‡é¢˜æ”¾åœ¨ `h1` å…ƒç´ ä¸­ï¼Œå…?`id` ä¸?`title`ã€?
1. ä½ åº”è¯¥æœ‰ä¸€ä¸?`id` ä¸?`description` çš?`p` å…ƒç´ ï¼Œç”¨äºŽç®€è¦è¯´æ˜Žã€?
1. ä½ åº”è¯¥æœ‰ä¸€ä¸?`form` å…ƒç´ ï¼Œ`id` ä¸?`survey-form`ã€?
1. åœ¨è¡¨å•å…ƒç´ ä¸­ï¼Œä½ åº”è¯¥æœ‰ä¸€ä¸ªè¾“å…¥å§“åçš„å¿…å¡« `input` å­—æ®µï¼Œå…¶ `id` ä¸?`name`ï¼Œ`type` ä¸?`text`ã€?
1. åœ¨è¡¨å•å…ƒç´ ä¸­ï¼Œä½ åº”è¯¥æœ‰ä¸€ä¸ªè¾“å…¥ç”µå­é‚®ä»¶çš„å¿…å¡« `input` å­—æ®µï¼Œå…¶ `id` ä¸?`email`ã€?
1. å¦‚æžœä½ è¾“å…¥äº†æ ¼å¼ä¸æ­£ç¡®çš„é‚®ç®±ï¼Œä½ å°†ä¼šçœ‹è§ HTML5 éªŒè¯é”™è¯¯ä¿¡æ¯ã€?
1. åœ¨è¡¨å•å…ƒç´ ä¸­ï¼Œä½ åº”è¯¥æœ‰ä¸€ä¸?`input` å­—æ®µï¼Œç”¨äºŽè¾“å…¥æ•°å­—ï¼Œå…?`id` ä¸?`number`ã€?
1. æ•°å­—è¾“å…¥ä¸åº”æŽ¥å—éžæ•°å­—ï¼Œæˆ–æ˜¯é˜»æ­¢ä½ è¾“å…¥å®ƒä»¬ï¼Œæˆ–æ˜¯æ˜¾ç¤ºä¸€ä¸?HTML5 éªŒè¯é”™è¯¯ï¼ˆå–å†³äºŽä½ çš„æµè§ˆå™¨ï¼‰ã€?
1. å¦‚æžœä½ è¾“å…¥çš„æ•°å­—è¶…å‡ºäº†èŒƒå›´ï¼ˆä½¿ç”¨ `min` å’?`max` å±žæ€§å®šä¹‰ï¼‰ï¼Œä½ å°†ä¼šçœ‹è§ HTML5 éªŒè¯é”™è¯¯ä¿¡æ¯ã€?
1. è¡¨å•ä¸­çš„åå­—ã€é‚®ç®±å’Œæ•°å­—è¾“å…¥æ¡†åº”æœ‰å¯¹åº”çš„åŒ…å«æè¿°è¾“å…¥æ¡†ç”¨é€”çš„ `label` å…ƒç´ ï¼Œid åº”åˆ†åˆ«ä¸º `name-label`ã€`email-label` å’?`number-label`ã€?
1. åœ¨è¡¨å•ä¸­çš„åå­—ã€é‚®ç®±å’Œæ•°å­—è¾“å…¥æ¡†ä¸­ï¼Œåº”çœ‹åˆ°å„è‡ªçš„æè¿°æ–‡å­—ä½œä¸ºå ä½ç¬¦ã€?
1. åœ¨è¡¨å•å…ƒç´ å†…ï¼?ä½ åº”è¯¥æœ‰ä¸€ä¸?`select` ä¸‹æ‹‰å…ƒç´ ï¼?`id` ä¸?`dropdown`ï¼Œå®ƒè‡³å°‘æœ‰ä¸¤ä¸ªé€‰é¡¹ã€?
1. åœ¨è¡¨å•å…ƒç´ å†…éƒ¨ï¼Œä½ å¯ä»¥ä»Žä¸€ç»„è‡³å°‘åŒ…å«ä¸¤ä¸ªå•é€‰æŒ‰é’®çš„é€‰é¡¹ä¸­è¿›è¡Œé€‰æ‹©ï¼Œè¿™äº›æŒ‰é’®é€šè¿‡ `name` å±žæ€§è¿›è¡Œåˆ†ç»„ã€?
1. åœ¨è¡¨å•å…ƒç´ å†…éƒ¨ï¼Œä½ å¯ä»¥ä»Žä¸€ç³»åˆ—å¤é€‰æ¡†ä¸­é€‰æ‹©å¤šä¸ªå­—æ®µï¼Œæ¯ä¸ªå¤é€‰æ¡†éƒ½å¿…é¡»å…·æœ?`value` å±žæ€§ã€?
1. åœ¨è¡¨å•å…ƒç´ å†…ï¼Œä½ åº”è¯¥æœ‰ä¸€ä¸?`textarea` ç”¨äºŽé™„åŠ æ³¨é‡Šã€?
1. åœ¨è¡¨å•å…ƒç´ ä¸­ï¼Œåº”è¯¥æœ‰ä¸€ä¸ªæŒ‰é’®ï¼Œå…?`id` ä¸?`submit`ï¼Œç”¨äºŽæäº¤æ‰€æœ‰è¾“å…¥ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`id` ä¸?`title` çš?`h1` å…ƒç´ ã€?

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

ä½ çš„ id ä¸?`#title` çš„å…ƒç´ ä¸åº”ä¸ºç©ºã€?

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`p` å…ƒç´ ï¼Œå…¶ `id` ä¸?`description`ã€?

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

ä½ çš„ `#description` ä¸åº”ä¸ºç©ºã€?

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`id` ä¸?`survey-form` çš?`form` å…ƒç´ ã€?

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`input` å…ƒç´ ï¼Œå…¶ `id` ä¸?`name`ã€?

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

ä½ çš„ `#name` åº”è¯¥æœ‰ä¸€ä¸?`type`ï¼Œå€¼ä¸º `text`ã€?

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

ä½ çš„ `#name` åº”è¯¥æ˜¯å¿…å¡«çš„è¾“å…¥ã€?

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

ä½ çš„ `#name` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`input` å…ƒç´ ï¼Œå…¶ `id` ä¸?`email`ã€?

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

ä½ çš„ `#email` åº”è¯¥æœ‰ä¸€ä¸?`type`ï¼Œå€¼ä¸º `email`ã€?

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

ä½ çš„ `#email` åº”è¯¥æ˜¯å¿…å¡«çš„è¾“å…¥ã€?

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

ä½ çš„ `#email` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`input` å…ƒç´ ï¼Œå…¶ `id` ä¸?`number`ã€?

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

ä½ çš„ `#number` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

ä½ çš„ `#number` åº”è¯¥æœ‰ä¸€ä¸?`type`ï¼Œå€¼ä¸º `number`ã€?

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

ä½ çš„ `#number` åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸ºæ•°å­—çš?`min` å±žæ€§ã€?

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

ä½ çš„ `#number` åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸ºæ•°å­—çš?`max` å±žæ€§ã€?

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ï¼Œå…¶ `id` ä¸?`name-label`ã€?

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ï¼Œå…¶ `id` ä¸?`email-label`ã€?

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ï¼Œå…¶ `id` ä¸?`number-label`ã€?

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

ä½ çš„ `#name-label` åº”åŒ…å«æè¿°è¾“å…¥çš„æ–‡æœ¬ã€?

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

ä½ çš„ `#email-label` åº”åŒ…å«æè¿°è¾“å…¥çš„æ–‡æœ¬ã€?

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

ä½ çš„ `#number-label` åº”è¯¥åŒ…å«æè¿°è¾“å…¥çš„æ–‡æœ¬ã€?

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

ä½ çš„ `#name-label` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

ä½ çš„ `#email-label` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

ä½ çš„ `#number-label` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

ä½ çš„ `#name` åº”è¯¥æœ‰ä¸€ä¸?`placeholder` å±žæ€§ä¸”æœ‰å€¼ã€?

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

ä½ çš„ `#email` åº”è¯¥æœ‰ä¸€ä¸?`placeholder` å±žæ€§ä¸”æœ‰å€¼ã€?

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

ä½ çš„ `#number` åº”è¯¥æœ?`placeholder` å±žæ€§ä¸”æœ‰å€¼ã€?

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`select` å­—æ®µï¼Œå…¶ `id` ä¸?`dropdown`ã€?

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

ä½ çš„ `#dropdown` åº”è¯¥è‡³å°‘æœ‰ä¸¤ä¸ªå¯é€‰æ‹©çš„ï¼ˆæœªç¦ç”¨ï¼‰`option` å…ƒç´ ã€?

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

ä½ çš„ `#dropdown` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

ä½ åº”è¯¥æœ‰è‡³å°‘ä¸¤ä¸ª `input` å…ƒç´ ï¼Œ`type` ä¸?`radio`ï¼ˆå•é€‰æŒ‰é’®ï¼‰ã€?

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

ä½ åº”è¯¥æœ‰è‡³å°‘ä¸¤ä¸ªå•é€‰æŒ‰é’®ï¼Œæ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

ä½ æ‰€æœ‰çš„å•é€‰æŒ‰é’®éƒ½åº”è¯¥æœ‰ä¸€ä¸?`value` å±žæ€§å’Œå€¼ã€?

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

ä½ æ‰€æœ‰çš„å•é€‰æŒ‰é’®éƒ½åº”è¯¥æœ‰ä¸€ä¸?`name` å±žæ€§å’Œå€¼ã€?

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

æ¯ä¸ªå•é€‰æŒ‰é’®ç»„åº”è‡³å°‘æœ‰ 2 ä¸ªå•é€‰æŒ‰é’®ã€?

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

ä½ åº”è¯¥è‡³å°‘æœ‰ä¸¤ä¸ª `input` å…ƒç´ ï¼Œ`type` ä¸?`checkbox`ï¼ˆå¤é€‰æ¡†ï¼‰ï¼Œå®ƒä»¬æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

ä½ åœ¨ `#survey-form` ä¸­çš„æ‰€æœ‰å¤é€‰æ¡†éƒ½åº”è¯¥æœ‰ `value` å±žæ€§å’Œå€¼ã€?

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

ä½ è‡³å°‘åº”è¯¥æœ‰ä¸€ä¸?`textarea` å…ƒç´ ï¼Œå®ƒæ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`input` æˆ?`button` å…ƒç´ ï¼Œå…¶ `id` ä¸?`submit`ã€?

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

ä½ çš„ `#submit` åº”è¯¥æœ‰ä¸€ä¸?`type`ï¼Œå€¼ä¸º `submit`ã€?

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

ä½ çš„ `#submit` åº”è¯¥æ˜?`#survey-form` çš„å­å…ƒç´ ã€?

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">MasterPuti Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend MasterPuti to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of MasterPuti?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```

