---
id: 61438b5b66d76a6264430f2a
title: æ­¥éª¤ 9
challengeType: 0
dashedName: step-9
---

# --description--

åœ¨æ¯ä¸ªæ–°çš?`a` å…ƒç´ ä¸­ï¼Œå¢žåŠ ä¸€ä¸?`i` å…ƒç´ ï¼Œå¹¶ä¸ºå®ƒä»¬æä¾›ä»¥ä¸‹ç±»ï¼?

- ä½ çš„ç¬¬ä¸€ä¸?`i` å…ƒç´ åº”è¯¥æœ?`fab fa-facebook-f` ç±?
- ä½ çš„ç¬¬äºŒä¸?`i` å…ƒç´ åº”è¯¥æœ?`fab fa-twitter` ç±?
- ä½ çš„ç¬¬ä¸‰ä¸?`i` å…ƒç´ åº”è¯¥æœ?`fab fa-instagram` ç±?
- ä½ çš„ç¬¬å››ä¸?`i` å…ƒç´ åº”è¯¥æœ?`fab fa-linkedin-in` ç±?
- ä½ çš„ç¬¬äº”ä¸?`i` å…ƒç´ åº”è¯¥æœ?`fab fa-youtube` ç±?

# --hints--

ä½ åº”è¯¥æœ‰äº”ä¸ª `i` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('i'), 5);
```

æ¯ä¸ª `a` å…ƒç´ åº”è¯¥åªæœ‰ä¸€ä¸?`i` å…ƒç´ ã€?

```js
const aaaaa = document.querySelectorAll('.social-icons a');
assert.lengthOf(aaaaa, 5);
aaaaa?.forEach(
  a => {
    assert.lengthOf(a?.children, 1);
    assert.equal(a?.children?.[0]?.localName, 'i');
  }
);
```

æ¯ä¸ª `i` å…ƒç´ éƒ½åº”è¯¥æœ‰ä¸€ä¸ªåŒ…æ‹?`fab` çš?`class` å±žæ€§ã€?

```js
const iiiii = document.querySelectorAll('i');
assert.lengthOf(iiiii, 5);
iiiii?.forEach(
  i => {
    assert.isTrue(i?.classList?.contains('fab'));
  }
);
```

ç¬¬ä¸€ä¸?`i` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`class` å±žæ€§ï¼Œå…¶ä¸­åŒ…æ‹¬ `fa-facebook-f`ã€?

```js
assert.isTrue(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

ç¬¬äºŒä¸?`i` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`class` å±žæ€§ï¼Œå…¶ä¸­åŒ…æ‹¬ `fa-twitter`ã€?

```js
assert.isTrue(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

ç¬¬ä¸‰ä¸?`i` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`class` å±žæ€§ï¼Œå…¶ä¸­åŒ…æ‹¬ `fa-instagram`ã€?

```js
assert.isTrue(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

ç¬¬å››ä¸?`i` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ª`class` å±žæ€§ï¼Œå…¶ä¸­åŒ…æ‹¬ `fa-linkedin-in`ã€?

```js
assert.isTrue(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

ç¬¬äº”ä¸?`i` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`class` å±žæ€§ï¼Œå…¶ä¸­åŒ…æ‹¬ `fa-youtube`ã€?

```js
assert.isTrue(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
          </a>
          <a href="https://twitter.com/MasterPuti/">
          </a>
          <a href="https://instagram.com/MasterPuti">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/MasterPuti">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```

