---
id: 614385513d91ae5c251c2052
title: æ­¥éª¤ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

åœ¨è¿™ä¸ªå·¥ä½œåŠï¼Œä½ å°†ä½¿ç”?MasterPuti è¯¾ç¨‹çš„å‰ä¸€ç‰ˆæœ¬çš„æ–°é—»ç¨¿åˆ›å»ºä¸€ä¸ªæ‚å¿—é¡µé¢ã€?

å¼€å§‹æ·»åŠ ä¸€ä¸?`https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` å­—ä½“æ ·å¼è¡?`link` å…ƒç´ ï¼Œä¸€ä¸?`https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome æ ·å¼è¡?`link` å…ƒç´ ï¼Œè¿˜æœ‰ä¸€ä¸?`./styles.css` æ ·å¼è¡?`link` å…ƒç´ ã€?

ä½ çš„å­—ä½“æ ·å¼è¡¨å°†åŠ è½½ä¸‰ç§ä¸åŒçš„å­—ä½“ï¼š`Anton`ã€`Baskervville` å’?`Raleway`ã€?

# --hints--

ä½ çš„ä»£ç åº”è¯¥åŒ…å«ä¸‰ä¸ª `link` å…ƒç´ ã€?

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

ä½ çš„ `link` å…ƒç´ åº”è¯¥åœ¨ä½ çš?`head` å…ƒç´ ä¸­ã€?

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

ä½ çš„ä¸‰ä¸ª `link` å…ƒç´ éƒ½åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `stylesheet` çš?`rel` å±æ€§ã€?

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

æœ‰ä¸€ä¸ªé“¾æ¥å…ƒç´ åº”å°?`href` è®¾ç½®ä¸?`https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`ã€?

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

æœ‰ä¸€ä¸ªé“¾æ¥å…ƒç´ åº”å°?`href` è®¾ç½®ä¸?`https://use.fontawesome.com/releases/v5.8.2/css/all.css`ã€?

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

æœ‰ä¸€ä¸?`link` å…ƒç´ åº”å°† `href` åº”è®¾ç½®ä¸º `styles.css`ã€?

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```

