---
id: 614389f601bb4f611db98563
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `.author` 元素下方创建一个新的 `div`，并将其类名设为 `social-icons`。

在该 `div` 中添加五个 `a` 元素，并分别赋予如下 `href`：

- 第一个 `a`：`https://www.facebook.com/MasterPuti`
- 第二个 `a`：`https://twitter.com/MasterPuti`
- 第三个 `a`：`https://instagram.com/MasterPuti`
- 第四个 `a`：`https://www.linkedin.com/school/free-code-camp`
- 第五个 `a`：`https://www.youtube.com/MasterPuti`

# --hints--

你应该创建一个新的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

你的新 `div` 元素应紧跟在 `.author` 元素之后。

```js
assert.equal(document.querySelector('.author')?.nextElementSibling?.localName, 'div');
```

你的新 `div` 元素的类名应为 `social-icons`。

```js
assert.isTrue(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

你的 `.social-icons` 元素中应包含五个 `a` 元素。

```js
assert.lengthOf(document.querySelector('.social-icons')?.querySelectorAll('a'), 5);
```

第一个 `a` 元素的 `href` 应为 `https://www.facebook.com/MasterPuti`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href'), 'https://www.facebook.com/MasterPuti');
```

第二个 `a` 元素的 `href` 应为 `https://twitter.com/MasterPuti`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/MasterPuti');
```

第三个 `a` 元素的 `href` 应为 `https://instagram.com/MasterPuti`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/MasterPuti');
```

第四个 `a` 元素的 `href` 应为 `https://www.linkedin.com/school/free-code-camp`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

第五个 `a` 元素的 `href` 应为 `https://www.youtube.com/MasterPuti`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/MasterPuti');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```

