---
id: 61438ec09438696391076d6a
title: æ­¥éª¤ 10
challengeType: 0
dashedName: step-10
---

# --description--

åœ?`.heading` å…ƒç´ ä¸‹æ–¹ï¼Œåˆ›å»ºä¸€ä¸ªæ–°çš?`section` å…ƒç´ ï¼Œå¹¶å°?`class` è®¾ç½®ä¸?`text`ã€?åœ¨å…¶ä¸­ï¼Œåˆ›å»ºä¸€ä¸?`p` å…ƒç´ ï¼Œå°† `class` è®¾ç½®ä¸?`first-paragraph`ï¼Œå¹¶æ·»åŠ ä»¥ä¸‹æ–‡æœ¬ï¼?

```markup
Soon the MasterPuti curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you'll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.
```

# --hints--

ä½ åº”è¯¥åˆ›å»ºä¸€ä¸ªæ–°çš?`section` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

ä½ çš„æ–?`section` å…ƒç´ åº”è¯¥åœ¨ä½ çš?`.heading` å…ƒç´ ä¹‹åã€?

```js
assert.equal(document.querySelectorAll('section')?.[1]?.previousElementSibling?.className, 'heading');
```

ä½ çš„æ–?`section` å…ƒç´ åº”å°† `class` è®¾ç½®ä¸?`text`ã€?

```js
assert.equal(document.querySelectorAll('section')?.[1]?.className, 'text');
```

ä½ åº”è¯¥åœ¨ `.text` å…ƒç´ ä¸­åˆ›å»ºä¸€ä¸ªæ–°çš?`p` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelector('.text')?.querySelectorAll('p'), 1);
```

ä½ çš„æ–?`p` å…ƒç´ åº”å°† `class` è®¾ç½®ä¸?`first-paragraph`ã€?

```js
assert.equal(document.querySelector('.text p')?.className, 'first-paragraph');
```

ä½ çš„æ–?`p` å…ƒç´ åº”å…·æœ‰å·²æä¾›çš„æ–‡æœ¬ã€?

```js
assert.equal(document.querySelector('.first-paragraph')?.innerText, 'Soon the MasterPuti curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you\'ll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/MasterPuti/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/MasterPuti">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/MasterPuti">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```

