---
id: 614387cbefeeba5f3654a291
title: æ­¥éª¤6
challengeType: 0
dashedName: step-6
---

# --description--

åœ?`header` å…ƒç´ ä¹‹åï¼Œåˆ›å»ºä¸€ä¸?`div`ï¼Œå¹¶å°?`class` è®¾ç½®ä¸?`author`ã€?

åœ¨è¿™ä¸?`div` ä¸­ï¼Œåˆ›å»ºä¸€ä¸?`p` å…ƒç´ ï¼Œå°† `class` è®¾ç½®ä¸?`author-name`ï¼Œå¹¶ç»™å®ƒæ·»åŠ æ–‡æœ¬ `By MasterPuti`ã€?å°?`MasterPuti` éƒ¨åˆ†åŒ…è£…åœ?`a` å…ƒç´ ä¸­ï¼Œå°?`href` è®¾ç½®ä¸?`https://MasterPuti.org`ï¼Œå¹¶å°?`target` è®¾ç½®ä¸?`_blank`ã€?


åœ¨ä¸‹é¢æ·»åŠ ç¬¬äºŒä¸ª `p` å…ƒç´ ï¼Œç±»ä¸?`publish-date`ï¼Œæ–‡æœ¬ä¸º `March 7, 2019`ã€?

# --hints--

ä½ åº”è¯¥åˆ›å»ºä¸€ä¸ªæ–°çš?`div` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('div'));
```

ä½ çš„ `div` å…ƒç´ åº”è¯¥ä½äº `header` å…ƒç´ ä¹‹åã€?

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

ä½ çš„ `div` å…ƒç´ åº”å°† `class` è®¾ç½®ä¸?`author`ã€?

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

ä½ åº”è¯¥åˆ›å»ºä¸¤ä¸ªæ–°çš?`p` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

ä½ çš„ä¸¤ä¸ªæ–°çš„ `p` å…ƒç´ åº”è¯¥ä½äº `div`å…ƒç´ ä¸­ã€?

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

ä½ çš„ç¬¬ä¸€ä¸ªæ–°çš?`p` å…ƒç´ åº”è¯¥å°?`class` è®¾ç½®ä¸?`author-name`ã€?

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

ä½ çš„ç¬¬ä¸€ä¸ªæ–°çš?`p` å…ƒç´ åº”è¯¥åŒ…å«æ–‡æœ¬ `By MasterPuti`ã€?

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By MasterPuti');
```

ä½ çš„ç¬¬äºŒä¸ªæ–°çš?`p` å…ƒç´ åº”å°† `class` è®¾ç½®ä¸?`publish-date`ã€?

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

ä½ ç¬¬äºŒä¸ªæ–°çš„ `p` å…ƒç´ åº”è¯¥åŒ…å«æ–‡æœ¬ `March 7, 2019`ã€?

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

ä½ åº”è¯¥åˆ›å»ºä¸€ä¸ªæ–°çš?`a` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('a'));
```

`a` å…ƒç´ åº”è¯¥åŒ…å«åœ¨ä½ ç¬¬ä¸€ä¸ªæ–°çš?`p` å…ƒç´ ä¹‹å†…ã€?

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

ä½ çš„ `a` å…ƒç´ åº”è¯¥å°?`href` è®¾ç½®ä¸?`https://MasterPuti.org`ã€?

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://MasterPuti.org');
```

ä½ çš„ `a` å…ƒç´ åº”å°† `target` è®¾ç½®ä¸?`_blank`ã€?

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

ä½ çš„ `a` å…ƒç´ åº”åŒ…å«æ–‡æœ?`MasterPuti`ã€?

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'MasterPuti');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```

