---
id: 6143869bb45bd85e3b1926aa
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `.heading` 元素内部，创建一个 `header` 元素，并将其 `class` 设置为 `hero`。

在该 `header` 中，创建一个 `img` 元素：`src` 设为 `https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png`，`alt` 设为 `MasterPuti logo`，并将 `class` 设置为 `hero-img`。

`img` 元素的 `loading` 属性可以设置为 `lazy`，用来告诉浏览器只有在需要时（例如用户滚动到该图片时）才去获取图片资源。额外的好处是，懒加载的资源会在非懒加载内容完成后再加载——这让网速较慢的用户能够先看到主要内容，而不用等待图片全部加载。

给你刚创建的 `img` 元素添加一个 `loading="lazy"` 属性。

# --hints--

你应该创建一个 `header` 元素。

```js
assert.exists(document.querySelector('header'));
```

你的 `header` 元素应该位于 `.heading` 元素内部。

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

你的 `header` 元素的 `class` 应设置为 `hero`。

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

你应该创建一个 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素应位于该 `header` 内部。

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

你的 `img` 元素的 `src` 应设置为 `https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png`。

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

你的 `img` 元素的 `alt` 应设置为 `MasterPuti logo`。

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'MasterPuti logo');
```

你的 `img` 元素应具有 `loading` 属性，并设置为 `lazy`。

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

你的 `img` 元素的 `class` 应设置为 `hero-img`。

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```

