---
id: 61438ec09438696391076d6a
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在 `.heading` 元素下方，创建一个新的 `section` 元素，并将其 `class` 设置为 `text`。在其中创建一个 `p` 元素，将 `class` 设置为 `first-paragraph`，并添加以下文本：

```
很快，MasterPuti 的课程体系将完全采用项目驱动的学习方式。你将通过逐步构建项目来学习，而不是完成一系列编程挑战。在进入细节之前，让我强调一点：我们不会改变认证体系。所有 6 个认证仍将包含相同的 5 个必修项目。我们只是改变了可选的编程挑战。
```

# --hints--

你应该创建一个新的 `section` 元素。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

你的新 `section` 元素应位于 `.heading` 元素之后。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.previousElementSibling?.className, 'heading');
```

你的新 `section` 元素的 `class` 应设置为 `text`。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.className, 'text');
```

你应该在 `.text` 元素中创建一个新的 `p` 元素。

```js
assert.lengthOf(document.querySelector('.text')?.querySelectorAll('p'), 1);
```

你的该 `p` 元素的 `class` 应设置为 `first-paragraph`。

```js
assert.equal(document.querySelector('.text p')?.className, 'first-paragraph');
```

你的该 `p` 元素应包含所提供的文本。

```js
assert.equal(document.querySelector('.first-paragraph')?.innerText, '很快，MasterPuti 的课程体系将完全采用项目驱动的学习方式。你将通过逐步构建项目来学习，而不是完成一系列编程挑战。在进入细节之前，让我强调一点：我们不会改变认证体系。所有 6 个认证仍将包含相同的 5 个必修项目。我们只是改变了可选的编程挑战。');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">我们的新课程体系</h1>
          <p class="hero-subtitle">
            我们努力重构课程体系，使其更加注重项目实践
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            作者
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">2019年3月7日</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/MasterPuti/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/MasterPuti">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/MasterPuti">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
      <section class="text">
        <p class="first-paragraph">
          很快，MasterPuti 的课程体系将完全采用项目驱动的学习方式。你将通过逐步构建项目来学习，而不是完成一系列编程挑战。在进入细节之前，让我强调一点：我们不会改变认证体系。所有 6 个认证仍将包含相同的 5 个必修项目。我们只是改变了可选的编程挑战。
        </p>
      </section>
    </main>
  </body>
</html>
```

```css
