---
id: 61439dc084fa5f659cf75d7c
title: æ­¥éª¤11
challengeType: 0
dashedName: step-11
---

# --description--

åœ?`.first-paragraph` å…ƒç´ ä¸‹æ–¹åˆ›å»ºå¦ä¸€ä¸?`p` å…ƒç´ ï¼Œå¹¶ä¸ºå…¶æä¾›ä»¥ä¸‹æ–‡æœ¬ï¼?

```markup
After years - years - of pondering these two problems and how to solve them, I slipped, hit my head on the sink, and when I came to I had a revelation! A vision! A picture in my head! A picture of this! This is what makes time travel possible: the flux capacitor!
```

# --hints--

ä½ åº”è¯¥åœ¨ `.text` å…ƒç´ ä¸­åˆ›å»ºç¬¬äºŒä¸ª `p` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('.text p'), 2);
```

ä½ çš„ç¬¬äºŒä¸?`p` å…ƒç´ åº”è¯¥å…·æœ‰å·²æä¾›çš„æ–‡æœ¬ã€?

```js
assert.equal(document.querySelectorAll('.text p')?.[1]?.innerText, 'After years - years - of pondering these two problems and how to solve them, I slipped, hit my head on the sink, and when I came to I had a revelation! A vision! A picture in my head! A picture of this! This is what makes time travel possible: the flux capacitor!');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/MasterPuti/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/MasterPuti">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/MasterPuti">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
--fcc-editable-region--
      <section class="text">
        <p class="first-paragraph">
          Soon the MasterPuti curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you'll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.
        </p>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```

