---
id: 614e0e588f0e8a772a8a81a6
title: æ­¥éª¤5
challengeType: 0
dashedName: step-5
---

# --description--

ä½ çš„å›¾åƒç›®å‰å ç”¨äº†å¤§é‡ç©ºé—´ã€?ä¸ºäº†æ›´å¥½åœ°æŸ¥çœ‹ä½ æ­£åœ¨å¤„ç†çš„å†…å®¹ï¼Œè¯·ç»™ `img` å…ƒç´ æ·»åŠ  `width` å±æ€§ï¼Œå…¶å€¼ä¸º `400`ã€?

ç¨åå½“ä½ å¤„ç† CSS æ—¶ï¼Œä½ å°†åˆ é™¤å®ƒã€?

# --hints--

ä½ çš„ `img` å…ƒç´ åº”å°† `width` å±æ€§è®¾ç½®ä¸º `400`ã€?

```js
assert.equal(document.querySelector('img')?.getAttribute('width'), '400');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```

