---
id: 6165d3b702a5d92ad970b30c
title: æ­¥éª¤4
challengeType: 0
dashedName: step-4
---

# --description--

åœ?`img` å…ƒç´ ä¹‹åï¼Œæ·»åŠ ä¸€ä¸?`h1` å…ƒç´ ï¼Œå…¶ä¸?`class` è®¾ç½®ä¸?`hero-title`ï¼Œæ–‡æœ¬è®¾ç½®ä¸º `OUR NEW CURRICULUM`ï¼Œç„¶åæ·»åŠ ä¸€ä¸?`p` å…ƒç´ ï¼Œå…¶ä¸?`class` è®¾ç½®ä¸?`hero-subtitle`ï¼Œæ–‡æœ¬è®¾ç½®ä¸º `Our efforts to restructure our curriculum with a more project-based focus`ã€?

# --hints--

ä½ åº”è¯¥åˆ›å»ºä¸€ä¸?`h1` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('h1'));
```

ä½ çš„ `h1` å…ƒç´ åº”è¯¥åœ¨ä½ çš?`img` å…ƒç´ ä¹‹åã€?

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

ä½ çš„ `h1` å…ƒç´ åº”å°† `class` è®¾ç½®ä¸?`hero-title`ã€?

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

ä½ çš„ `h1` å…ƒç´ çš„æ–‡æœ¬åº”è¯¥ä¸º `OUR NEW CURRICULUM`ã€?

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

ä½ åº”è¯¥åˆ›å»ºä¸€ä¸ªæ–°çš?`p` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('p'));
```

ä½ çš„ `p` å…ƒç´ åº”è¯¥åœ?`h1` å…ƒç´ ä¹‹åã€?

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

ä½ çš„ `p` å…ƒç´ åº”å°† `class` è®¾ç½®ä¸?`hero-subtitle`ã€?

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

ä½ çš„ `p` å…ƒç´ çš„æ–‡æœ¬åº”è¯¥ä¸º `Our efforts to restructure our curriculum with a more project-based focus`ã€?

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```

