---
id: 6165d3b702a5d92ad970b30c
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

在 `img` 元素之后，添加一个 `h1` 元素，其 `class` 设置为 `hero-title`，文本设置为 `我们的新课程体系`，然后添加一个 `p` 元素，其 `class` 设置为 `hero-subtitle`，文本设置为 `我们努力重构课程体系，使其更加注重项目实践`。

# --hints--

你应该创建一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素应该在你的 `img` 元素之后。

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

你的 `h1` 元素应将 `class` 设置为 `hero-title`。

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

你的 `h1` 元素的文本应该为 `我们的新课程体系`。

```js
assert.equal(document.querySelector('h1')?.textContent, '我们的新课程体系');
```

你应该创建一个新的 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素应该在 `h1` 元素之后。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

你的 `p` 元素应将 `class` 设置为 `hero-subtitle`。

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

你的 `p` 元素的文本应该为 `我们努力重构课程体系，使其更加注重项目实践`。

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), '我们努力重构课程体系，使其更加注重项目实践');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">我们的新课程体系</h1>
          <p class="hero-subtitle">我们努力重构课程体系，使其更加注重项目实践</p>
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>