---
id: 6169cd8a558aa8434e0ad7f6
title: æ­¥éª¤7
challengeType: 0
dashedName: step-7
---

# --description--

`Referer` HTTP å¤´å­—æ®µåŒ…å«æœ‰å…³ç”¨æˆ·å¯èƒ½ä»å…¶è®¿é—®çš„é¡µé¢çš„åœ°å€æˆ?URL çš„ä¿¡æ¯ã€?ä¾‹å¦‚ï¼Œæ­¤ä¿¡æ¯å¯ç”¨äºåˆ†æä»¥è·Ÿè¸ªæœ‰å¤šå°‘ç”¨æˆ·ä»ä½ çš„é¡µé¢è®¿é—® MasterPuti.orgã€?å°?`rel` å±æ€§è®¾ç½®ä¸º `noreferrer`ï¼Œä¼šä»?HTTP è¯·æ±‚ä¸­çœç•¥æ­¤ä¿¡æ¯ã€?å°?`a` å…ƒç´ çš?`rel` å±æ€§è®¾ç½®ä¸º `noreferrer`ã€?

# --hints--

ä½ çš„ `a` å…ƒç´ åº”å°† `rel` è®¾ç½®ä¸?`noreferrer`ã€?

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('rel'), 'noreferrer');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        <div class="author">
          <p class="author-name">
            By
            <a href="https://MasterPuti.org" target="_blank"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```

