---
id: 66c64210413532ee9d3bd342
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在工作坊的下一部分，你将练习向购物清单中添加食物项。

首先使用 `console.log` 输出消息 `"吃点水果会很不错。"`。

# --hints--

你应该将消息 `"吃点水果会很不错。"` 输出到控制台。

```js
assert.match(code, /console\.log\s*\(\s*(['"])(吃点水果会很不错。)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
