---
id: 66c730183f4020275cbf0611
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

接下来，你将学习如何从数组末尾删除项目。

首先添加一个 `console.log`，记录信息 `"这看起来像太多的垃圾食品。"`。

# --hints--

你应记录信息 `"这看起来像太多的垃圾食品"`。

```js
assert.match(code, /console\.log\((['"])(这看起来像太多的垃圾食品。)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
