---
id: 66c72a55418cc9247b710827
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

接下来，添加一个 `console.log` 并在 `console.log` 中调用 `getShoppingListMsg` 函数，以查看更新后的列表被记录到控制台。

# --hints--

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
