---
id: 66c72b0ffbc5522525768558
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

在前面的步骤中，你回顾了如何使用 `push` 方法在数组末尾添加一个项目。

但是 `push` 方法可接受多个参数，因此你可以像这样在数组末尾添加多个项目：

```js
array.push(item1, item2, item3);
```

在此步骤中，使用 `push` 方法在 `shoppingList` 数组中添加字符串 `"Popcorn", "Beef Jerky", "Potato Chips"`。

顺序很重要，因此请确保按照列出的顺序添加项目。

# --hints--

你的 `shoppingList` 数组总共应有 `6` 个字符串。

```js
assert.lengthOf(shoppingList, 6);
```

你的 `shoppingList` 应包含字符串 `"Popcorn"`。

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

你的 `shoppingList` 应包含字符串 `"Beef Jerky"`。

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

你的 `shoppingList` 应包含字符串 `"Potato Chips"`。

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
