---
id: 66c742d045c9fc2e09fa64b1
title: 步骤 19
challengeType: 1
dashedName: step-19
---

# --description--

对购物清单的最后一个修改是更新列表中的第一个项目。

在之前的课程中，你已经学习了如何使用方括号表示法和想要更新的项目的索引来更新项目。

以下是更新数组中项目的方法提醒：

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// 结果: [10, 2, 3, 4, 5]
console.log(array); 
```

将 `shoppingList` 数组中的第一个项目更新为 `"Canola Oil"`。

# --hints--

你应该将 `shoppingList` 数组中的第一个项目更新为 `"Canola Oil"`。

```js
assert.equal(shoppingList[0], "Canola Oil");
assert.notEqual(shoppingList[1], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("这看起来像太多的垃圾食品。");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("吃点甜点会很不错。");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("再想想，也许我们应该更注重健康。");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
