---
id: 66cbe2319d3845545a293a0b
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

现在是时候查看记录到控制台的消息了。

添加一个 `console.log` 并在 `console.log` 中调用 `getShoppingListMsg` 函数，以查看记录到控制台的消息。

# --hints--

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)` 函数。

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
