---
id: a103376db3ba46b2d50db289
title: Implement a Spinal Case Converter
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

脊柱式（spinal case）是一种字符串格式，其中所有单词都使用小写字母并用连字符分隔。`"this-is-spinal-tap"` 是脊柱式字符串的一个示例。

在这个实验中，你将创建一个函数，将给定的字符串转换为脊柱式格式。

**目标**：完成下面的用户故事并让所有测试通过以完成实验。

**用户故事：**

1. 你应该创建一个名为 `spinalCase` 的函数。
2. `spinalCase` 函数应该接受一个参数，即一个字符串。
3. `spinalCase` 函数应该返回脊柱式格式的字符串。例如，如果参数是 `JavaScript is awesome`，函数应该返回 `javascript-is-awesome`。

# --hints--

你应该创建一个名为 `spinalCase` 的函数。

```js
assert.isFunction(spinalCase);
```

`spinalCase` 应该接受一个参数。

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` 应该返回字符串 `this-is-spinal-tap`。

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` 应该返回字符串 `this-is-spinal-tap`。

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` 应该返回字符串 `the-andy-griffith-show`。

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` 应该返回字符串 `teletubbies-say-eh-oh`。

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` 应该返回字符串 `all-the-small-things`。

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
