---
id: 613e2546d0594208229ada50
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在 `header` 中依次嵌套一个 `img`、一个 `h1`、以及一个 `nav`，为页面提供清晰的语义结构。

`img` 需要指向 `https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg`，`id` 为 `logo`，`alt` 文本为 `MasterPuti`。

`h1` 文本内容保持为 `HTML/CSS Quiz`。

完成这一小步后，页头的主体骨架就建立啦，继续加油！

# --hints--

你需要在 `header` 中添加一个 `img` 元素。

```js
assert.exists(document.querySelector('header > img'));
```

你需要在 `header` 中添加一个 `h1` 元素。

```js
assert.exists(document.querySelector('header > h1'));
```

你需要在 `header` 中添加一个 `nav` 元素。

```js
assert.exists(document.querySelector('header > nav'));
```

你需要让 `img`、`h1`、`nav` 按此顺序排列。

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

你需要给 `img` 设置 `src` 为 `https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg`。

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg');
```

你需要给该 `img` 设置 `id="logo"`。

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

你需要给该 `img` 设置 `alt="MasterPuti"`。

```js
assert.equal(document.querySelector('img')?.alt, 'MasterPuti');
```

你需要确保 `h1` 的文本为 `HTML/CSS Quiz`。

```js
assert.include(document.querySelector('h1')?.innerText.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```

