---
id: 6140827cff96e906bd38fc2b
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

根据 <a href="https://design-style-guide.MasterPuti.org/" target="_blank" rel="noopener noreferrer nofollow">MasterPuti 风格指南</a>，logo 需要保持 `35 / 4` 的纵横比，并在文字四周留出统一内边距。

请先把 `background-color` 改为 `#0a0a23` 方便观察；然后使用 `aspect-ratio` 属性设为 `35 / 4`；最后为它添加 `padding: 0.4rem`。继续，你的细节掌控力越来越棒！

# --hints--

你需要让 `#logo` 的 `background-color` 为 `#0a0a23`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

你需要使用 `aspect-ratio` 属性。

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

你不应该使用 `height` 属性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

你不应该修改现有的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

你需要给该图片设置 `aspect-ratio: 35 / 4`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

你需要给该图片设置 `padding: 0.4rem`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(10rem, 18vw);

}
--fcc-editable-region--

```

