---
id: 6140883f74224508174794c0
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

让 `header` 占满父容器宽度，设置 `height: 50px`，背景色为 `#1b1b32`。然后将 `display` 改为 Flex，后续布局更灵活。继续，页面框架越来越清晰！

# --hints--

你应该使用 `header` 元素选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

你需要让 `header` 的 `width` 为 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

你需要让 `header` 的 `height` 为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

你需要让 `header` 的 `background-color` 为 `#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

你需要让 `header` 的 `display` 为 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```

