---
id: 61408e4ae3e35d08feb260eb
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

把 `h1` 的字体颜色设置为 `#f1be32`，字号设为 `min(5vw, 1.2em)`，在不同屏幕下都保持舒适显示。继续，主标题更醒目啦！

# --hints--

你应该使用 `h1` 元素选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

你需要将 `h1` 的 `color` 设置为 `#f1be32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.color, 'rgb(241, 190, 50)');
```

你需要将 `h1` 的 `font-size` 设置为 `min(5vw, 1.2em)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.fontSize, 'min(5vw, 1.2em)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

--fcc-editable-region--

--fcc-editable-region--
```

