---
id: 61408e4ae3e35d08feb260eb
title: æ­¥éª¤11
challengeType: 0
dashedName: step-11
---

# --description--

å°?`h1` å­—ä½“é¢œè‰²æ›´æ”¹ä¸?`#f1be32`ï¼Œå¹¶å°†å­—ä½“å¤§å°è®¾ç½®ä¸º `min(5vw, 1.2em)`ã€?

# --hints--

ä½ åº”è¯¥ä½¿ç”?`h1` å…ƒç´ é€‰æ‹©å™¨ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

ä½ åº”è¯¥å°† `h1` çš?`color` è®¾ç½®ä¸?`#f1be32`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.color, 'rgb(241, 190, 50)');
```

ä½ åº”å°?`h1` çš?`font-size` è®¾ç½®ä¸?`min(5vw, 1.2em)`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.fontSize, 'min(5vw, 1.2em)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

--fcc-editable-region--

--fcc-editable-region--
```

