---
id: 61408f155e798909b6908712
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

为了启用页面导航，请在 `nav` 中添加一个无序列表，包含三个列表项：

- `INFO`
- `HTML`
- `CSS`

每个列表项文字要包裹在 `a` 标签中。继续，导航结构马上完成！

# --hints--

你应该在 `nav` 内嵌套一个 `ul`。

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

你应该在该 `ul` 中嵌套三个 `li`。

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

你应该在每个 `li` 中嵌套一个 `a`。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

第一个 `a` 文本需要是 `INFO`。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[0]?.textContent.trim(), 'INFO');
```

第二个 `a` 文本需要是 `HTML`。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[1]?.textContent.trim(), 'HTML');
```

第三个 `a` 文本需要是 `CSS`。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')[2]?.textContent.trim(), 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```

