---
id: 6141fed65b61320743da5894
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

为了增强页面可访问性，可以使用 `role` 属性向辅助技术说明元素的语义用途。`role` 是 WAI 规范的一部分。

现在给三个 `section` 都添加 `role="region"`，让它们对屏幕阅读器更友好。继续，可访问语义再升级！

# --hints--

你需要给第一个 `section` 添加 `role="region"`。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('role'), 'region');
```

你需要给第二个 `section` 添加 `role="region"`。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('role'), 'region');
```

你需要给第三个 `section` 添加 `role="region"`。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('role'), 'region');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
        <section></section>
        <section></section>
        <section></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```

