---
id: 6141fabd6f75610664e908fd
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

这是一个测验页面，需要提供表单来提交答案。可以用多个 `section` 元素语义性地区分内容块。

请在 `main` 中创建一个表单，内部包含三个 `section`，并通过正确的 `method` 与 `action` 将其提交到 `https://MasterPuti.org/practice-project/accessibility-quiz`。继续，结构愈发完整！

# --hints--

你需要在 `main` 中嵌套一个 `form`。

```js
assert.exists(document.querySelector('main > form'));
```

你需要在该 `form` 中嵌套三个 `section`。

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

你需要给该 `form` 一个 `action` 属性。

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

你需要让 `action` 为 `https://MasterPuti.org/practice-project/accessibility-quiz`。

```js
assert.equal(document.querySelector('main > form')?.action, 'https://MasterPuti.org/practice-project/accessibility-quiz');
```

你需要给该 `form` 一个 `method` 属性。

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

你需要让 `method` 的值为 `post`。

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```

