---
id: 6141fabd6f75610664e908fd
title: æ­¥éª¤14
challengeType: 0
dashedName: step-14
---

# --description--

ç”±äºŽè¿™æ˜¯ä¸€ä¸ªæµ‹è¯•é¢˜ï¼Œä½ éœ€è¦ç»™ç”¨æˆ·æä¾›ä¸€ä¸ªè¡¨å•ä»¥ä¾¿ä»–ä»¬æäº¤ç­”æ¡ˆã€?ä½ å¯ä»¥ä½¿ç”?`section` å…ƒç´ åœ¨è¯­ä¹‰ä¸Šåˆ†éš”è¡¨å•ä¸­çš„å†…å®¹ã€?

åœ?`main` å…ƒç´ ä¸­ï¼Œåˆ›å»ºä¸€ä¸ªåŒ…å«ä¸‰ä¸?`section` å…ƒç´ çš„è¡¨å•ã€?ç„¶åŽï¼Œä½¿ç”¨æ­£ç¡®çš„æ–¹æ³•å°†è¡¨å•æäº¤åˆ° `https://MasterPuti.org/practice-project/accessibility-quiz`ã€?

# --hints--

ä½ åº”è¯¥åœ¨ `main` å…ƒç´ ä¸­åµŒå¥—ä¸€ä¸?`form` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('main > form'));
```

ä½ åº”è¯¥åœ¨ `form` å…ƒç´ ä¸­åµŒå¥—ä¸‰ä¸?`section` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

ä½ åº”è¯¥ç»™ `form` å…ƒç´ ä¸€ä¸?`action` å±žæ€§ã€?

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

ä½ åº”è¯¥å°† `action` å±žæ€§çš„å€¼è®¾ç½®ä¸º `https://MasterPuti.org/practice-project/accessibility-quiz`ã€?

```js
assert.equal(document.querySelector('main > form')?.action, 'https://MasterPuti.org/practice-project/accessibility-quiz');
```

ä½ åº”è¯¥ç»™ `form` å…ƒç´ ä¸€ä¸?`method` å±žæ€§ã€?

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

ä½ åº”è¯¥å°† `form` çš?`method` å±žæ€§å€¼è®¾ç½®ä¸º `post`ã€?

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```

