---
id: 614202874ca576084fca625f
title: æ­¥éª¤ 16
challengeType: 0
dashedName: step-16
---

# --description--

æ¯ä¸ª `region` è§’è‰²éƒ½éœ€è¦ä¸€ä¸ªæ ‡ç­¾ï¼Œä»¥å¸®åŠ©å±å¹•é˜…è¯»å™¨ç”¨æˆ·ç†è§£è¯¥åŒºåŸŸçš„ç”¨é€”ã€?æ·»åŠ æ ‡ç­¾çš„ä¸€ç§æ–¹æ³•æ˜¯ï¼Œåœ¨åŒºåŸŸå†…æ·»åŠ ä¸€ä¸ªæ ‡é¢˜å…ƒç´ ï¼Œç„¶åç”?`aria-labelledby` å±æ€§å¼•ç”¨å®ƒã€?

ç»?`section` å…ƒç´ æ·»åŠ ä»¥ä¸‹ `aria-labelledby` å±æ€§ï¼š

- `student-info`
- `html-questions`
- `css-questions`

ç„¶åï¼Œåœ¨æ¯ä¸ª `section` å…ƒç´ ä¸­ï¼ŒåµŒå¥—ä¸€ä¸?`h2` å…ƒç´ ï¼Œå…¶ `id` ä¸ç›¸åº”çš„ `aria-labelledby` å±æ€§åŒ¹é…ã€?ç»™æ¯ä¸?`h2` æ·»åŠ é€‚å½“çš„æ–‡æœ¬å†…å®¹ã€?

# --hints--

ä½ åº”è¯¥ç»™ç¬¬ä¸€ä¸?`section` å…ƒç´ æ·»åŠ å€¼ä¸º `student-info` çš?`aria-labelledby` å±æ€§ã€?

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

ä½ åº”è¯¥ç»™ç¬¬äºŒä¸?`section` å…ƒç´ æ·»åŠ å€¼ä¸º `html-questions` çš?`aria-labelledby` å±æ€§ã€?

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

ä½ åº”è¯¥ç»™ç¬¬ä¸‰ä¸?`section` å…ƒç´ æ·»åŠ å€¼ä¸º `css-questions` çš?`aria-labelledby` å±æ€§ã€?

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

ä½ åº”è¯¥åœ¨ç¬¬ä¸€ä¸?`section` å…ƒç´ ä¸­åµŒå¥—ä¸€ä¸?`h2` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

ä½ åº”è¯¥åœ¨ç¬¬äºŒä¸?`section` å…ƒç´ ä¸­åµŒå¥—ä¸€ä¸?`h2` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

ä½ åº”è¯¥åœ¨ç¬¬ä¸‰ä¸?`section` å…ƒç´ ä¸­åµŒå¥—ä¸€ä¸?`h2` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

ä½ åº”è¯¥ç»™ç¬¬ä¸€ä¸?`h2` å…ƒç´ ä¸€ä¸?`id` å±æ€§ï¼Œå€¼ä¸º `student-info`ã€?

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

ä½ åº”è¯¥ç»™ç¬¬äºŒä¸?`h2` å…ƒç´ ä¸€ä¸?`id` å±æ€§ï¼Œå€¼ä¸º `html-questions`ã€?

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

ä½ åº”è¯¥ç»™ç¬¬ä¸‰ä¸?`h2` å…ƒç´ ä¸€ä¸?`id` å±æ€§ï¼Œå€¼ä¸º `css-questions`ã€?

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

ä½ åº”è¯¥ç»™ç¬¬ä¸€ä¸?`h2` å…ƒç´ æ·»åŠ é€‚å½“çš„æ–‡æœ¬å†…å®¹ã€?_æç¤ºï¼šæˆ‘ä¼šä½¿ç”?`Student Info`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

ä½ åº”è¯¥ç»™ç¬¬äºŒä¸?`h2` å…ƒç´ æ·»åŠ é€‚å½“çš„æ–‡æœ¬å†…å®¹ã€?_æç¤ºï¼šæˆ‘ä¼šä½¿ç”?`HTML`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

ä½ åº”è¯¥ç»™ç¬¬ä¸‰ä¸?`h2` å…ƒç´ æ·»åŠ é€‚å½“çš„æ–‡æœ¬å†…å®¹ã€?_æç¤ºï¼šæˆ‘ä¼šä½¿ç”?`CSS`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```

