---
id: 6143639d5eddc7094161648c
title: æ­¥éª¤ 20
challengeType: 0
dashedName: step-20
---

# --description--

å°†æ¯ä¸?`input` é“¾æŽ¥åˆ°ç›¸åº”çš„ `label` å…ƒç´ å¾ˆé‡è¦ã€?è¿™ä¸ºè¾…åŠ©æŠ€æœ¯ç”¨æˆ·æä¾›äº†å¯¹è¾“å…¥çš„è§†è§‰å‚è€ƒã€?

å®žçŽ°æ–¹æ³•æ˜¯ä¸º `label` å…ƒç´ æ·»åŠ  `for` å±žæ€§ï¼Œå…¶å€¼ä¸Žå¯¹åº”çš?`input` çš?`id` ç›¸åŒã€?

è¿™ä¸ªéƒ¨åˆ†å°†å¡«å†™å­¦ç”Ÿçš„åå­—ã€ç”µå­é‚®ä»¶åœ°å€å’Œå‡ºç”Ÿæ—¥æœŸã€?ä¸?`label` å…ƒç´ æä¾›é€‚å½“çš?`for` å±žæ€§ä»¥åŠæ–‡æœ¬å†…å®¹ã€?ç„¶åŽï¼Œå°† `input` å…ƒç´ é“¾æŽ¥åˆ°ç›¸åº”çš„ `label` å…ƒç´ ã€?

# --hints--

ä½ åº”è¯¥ç»™ç¬¬ä¸€ä¸?`label` å…ƒç´ ä¸€ä¸ªé€‚å½“çš?`for` å±žæ€§ã€?

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.htmlFor?.length, 1);
```

ä½ åº”è¯¥ç»™ç¬¬äºŒä¸?`label` å…ƒç´ ä¸€ä¸ªé€‚å½“çš?`for` å±žæ€§ã€?

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.htmlFor?.length, 1);
```

ä½ åº”è¯¥ç»™ç¬¬ä¸‰ä¸?`label` å…ƒç´ ä¸€ä¸ªé€‚å½“çš?`for` å±žæ€§ã€?

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.htmlFor?.length, 1);
```

ä½ åº”è¯¥ç»™ç¬¬ä¸€ä¸?`label` å…ƒç´ æ·»åŠ é€‚å½“çš„æ–‡æœ¬å†…å®¹ã€?

```js
assert.isNotEmpty(document.querySelectorAll('label')[0]?.textContent.trim());
```

ä½ åº”è¯¥ç»™ç¬¬äºŒä¸?`label` å…ƒç´ æ·»åŠ é€‚å½“çš„æ–‡æœ¬å†…å®¹ã€?

```js
assert.isNotEmpty(document.querySelectorAll('label')[1]?.textContent.trim());
```

ä½ åº”è¯¥ç»™ç¬¬ä¸‰ä¸?`label` å…ƒç´ æ·»åŠ é€‚å½“çš„æ–‡æœ¬å†…å®¹ã€?

```js
assert.isNotEmpty(document.querySelectorAll('label')[2]?.textContent.trim());
```

ä½ åº”è¯¥ç»™ç¬¬ä¸€ä¸?`input` å…ƒç´ ä¸€ä¸ªä¸Žç¬¬ä¸€ä¸?`label` çš?`for` å±žæ€§ç›¸åŒ¹é…çš?`id` å±žæ€§ã€?

```js
assert.equal(document.querySelectorAll('input')?.[0]?.id, document.querySelectorAll('label')?.[0]?.htmlFor);
```

ä½ åº”è¯¥ç»™ç¬¬äºŒä¸?`input` å…ƒç´ ä¸€ä¸ªä¸Žç¬¬äºŒä¸?`label` çš?`for` å±žæ€§ç›¸åŒ¹é…çš?`id` å±žæ€§ã€?

```js
assert.equal(document.querySelectorAll('input')?.[1]?.id, document.querySelectorAll('label')?.[1]?.htmlFor);
```

ä½ åº”è¯¥ç»™ç¬¬ä¸‰ä¸?`input` å…ƒç´ ä¸€ä¸ªä¸Žç¬¬ä¸‰ä¸?`label` çš?`for` å±žæ€§ç›¸åŒ¹é…çš?`id` å±žæ€§ã€?

```js
assert.equal(document.querySelectorAll('input')?.[2]?.id, document.querySelectorAll('label')?.[2]?.htmlFor);
```

ä½ ä¸åº”è¯¥ä¸ºå¤šä¸?`input` å…ƒç´ ä½¿ç”¨ç›¸åŒçš?`id` å±žæ€§ã€?

```js
const id = (n) => document.querySelectorAll('input')?.[n]?.id;
assert.notEqual(id(0), id(1));
assert.notEqual(id(0), id(2));
assert.notEqual(id(1), id(2));
```

ä½ ä¸åº”è¯¥ä¸ºå¤šä¸?`label` å…ƒç´ ä½¿ç”¨ç›¸åŒçš?`for` å±žæ€§ã€?

```js
const htmlFor = (n) => document.querySelectorAll('label')?.[n]?.htmlFor;
assert.notEqual(htmlFor(0), htmlFor(1));
assert.notEqual(htmlFor(0), htmlFor(2));
assert.notEqual(htmlFor(1), htmlFor(2));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```

