---
id: 614394fb41985e0d2012a93e
title: 步骤 24
challengeType: 0
dashedName: step-24
---

# --description--

仅有题号并不够描述性，视力受限的用户可能难以理解其语义。一个常用技巧是添加只供屏幕阅读器读取的隐藏文本。

请在每个题号 `h3` 中的数字前嵌入一个 `span`，类名为 `sr-only`，用于提供更详细的朗读信息（稍后可填充）。继续，你开始为无障碍做细致加分项！

# --hints--

你需要在第一个 `h3` 中添加一个 `span`。

```js
assert.exists(document.querySelector('h3 > span'));
```

你需要让第一个 `span` 的类名为 `sr-only`。

```js
assert.equal(document.querySelector('span')?.className, 'sr-only');
```

你需要在第二个 `h3` 中添加一个 `span`。

```js
assert.exists(document.querySelectorAll('h3 > span')[1]);
```

你需要让第二个 `span` 的类名为 `sr-only`。

```js
assert.equal(document.querySelectorAll('span')[1]?.className, 'sr-only');
```

你需要确保第一个 `h3` 中的数字 `1` 在 `span` 后出现。

```js
assert.equal(document.querySelector('span')?.nextSibling?.textContent.trim(), '1');
```

你需要确保第二个 `h3` 中的数字 `2` 在 `span` 后出现。

```js
assert.equal(document.querySelectorAll('span')[1]?.nextSibling?.textContent.trim(), '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          <div class="question-block">
--fcc-editable-region--
            <h3>1</h3>
            <fieldset class="question"></fieldset>
          </div>
          <div class="question-block">
            <h3>2</h3>
--fcc-editable-region--
            <fieldset class="question"></fieldset>
          </div>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}


```

