---
id: 68af273960e009597dd861e2
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

你希望有一个可读的显示用于调试和打印。创建一个 `__str__` 方法，返回格式为 `title (year) - duration min, director` 的字符串（其中 `title`、`year`、`duration` 和 `director` 应该被替换为相应属性的值），这样打印电影时会给出一个很好的摘要。

# --hints--

你的 `Movie` 类应该有一个 `__str__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

你的 `__str__` 方法应该返回格式为 `title (year) - duration min, director` 的字符串，其中 `title`、`year`、`duration` 和 `director` 应该被替换为相应属性的值。

```js
({
    test: () => runPython(`
    import io
    import sys

    captured_output = io.StringIO()
    sys.stdout = captured_output

    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration


--fcc-editable-region--
```
