---
id: 6747237371f13173c9f6c27e
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `#music-player` 元素内部，添加一个 `header` 元素，其中包含一个 `h1` 元素，文本为 `音乐播放器`。

# --hints--

你应该有一个 `header` 元素。

```js
assert.exists(document.querySelector('header'));
```

你的 `header` 元素应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('header h1'));
```

你的 `h1` 元素应该包含文本 `音乐播放器`。

```js
assert.equal(document.querySelector('header h1')?.textContent, '音乐播放器');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">

    </div>
  );
}
```

