---
id: 674728eda5829d7b4c360643
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `.album-art-container` 元素之后，添加一个 `div` 元素，其 `className` 为 `track-info`。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应该有一个 `className` 值为 `track-info`。

```js
assert.equal(document.querySelector('#player-container div.track-info')?.className, 'track-info');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>

      </div>
    </div>
  );
}
```

