---
id: 67472ab9f570797dbb01568a
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `#track-title` 元素之后，添加一个 `p` 元素，其 `id` 为 `track-artist`。

# --hints--

你应该有一个 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素应该有一个 `id` 值为 `track-artist`。

```js
assert.equal(document.querySelector('p')?.getAttribute('id'), 'track-artist');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>

        </div>
      </div>
    </div>
  );
}
```

