---
id: 674724ff8af92f7733e6d980
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `header` 元素之后，添加一个 `div` 元素，其 `id` 为 `player-container`。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应该有一个 `id` 值为 `player-container`。

```js
assert.equal(document.querySelector('#player-container')?.getAttribute('id'), 'player-container');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">

      </div>

    </div>
  );
}
```

