---
id: 6747299cbed5107c6a00d64c
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在 `.track-info` 元素内部，添加一个 `h2` 元素，其 `id` 为 `track-title`。

# --hints--

你应该有一个 `h2` 元素。

```js
assert.exists(document.querySelector('h2'));
```

你的 `h2` 元素应该有一个 `id` 值为 `track-title`。

```js
assert.equal(document.querySelector('h2')?.getAttribute('id'), 'track-title');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>
        </div>
      </div>
    </div>
  );
}
```

