---
id: 67472eef7b4b7ba77b50929b
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在 `.track-info` 元素之后，添加一个 `div` 元素，其 `className` 为 `controls`。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应该有一个 `className` 值为 `controls`。

```js
assert.equal(document.querySelector('#player-container div.controls')?.className, 'controls');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>
          <p id="track-artist"></p>
        </div>

      </div>
    </div>
  );
}
```

