---
id: 6747274264e8a5799fb3e0b5
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `.album-art-container` 元素内部，添加一个 `img` 元素，其 `id` 为 `album-art`。

# --hints--

你应该有一个 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素应该有一个 `id` 值为 `album-art`。

```js
assert.equal(document.querySelector('img')?.getAttribute('id'), 'album-art');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">

        </div>
      </div>
    </div>
  );
}
```

