---
id: 6748bab9827a356c39559214
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

在 `.time-info` 元素内部，添加一个 `span` 元素，其 `id` 为 `current-time`，文本为 `0:00`。

# --hints--

你应该有一个 `span` 元素。

```js
assert.exists(document.querySelector('span'));
```

你的 `span` 元素应该有一个 `id` 值为 `current-time`。

```js
assert.equal(document.querySelector('span')?.getAttribute('id'), 'current-time');
```

你的 `#current-time` 元素应该包含文本 `0:00`。

```js
assert.equal(document.querySelector('#current-time')?.textContent, '0:00');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>
          <p id="track-artist"></p>
        </div>
        <div className="controls">
          <button id="play-pause-btn">播放</button>
        </div>
        <div className="progress-container">
          <div className="progress-bar">
            <div className="progress"></div>
          </div>
        </div>
        <div className="time-info">

        </div>
      </div>
    </div>
  );
}
```

