---
id: 674850b437668dd08edddbe2
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

在 `.progress-bar` 元素内部，添加一个 `div` 元素，其 `className` 为 `progress`。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应该有一个 `className` 值为 `progress`。

```js
assert.equal(document.querySelector('.progress-bar div')?.className, 'progress');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>
          <p id="track-artist"></p>
        </div>
        <div className="controls">
          <button id="play-pause-btn">播放</button>
        </div>
        <div className="progress-container">
          <div className="progress-bar">
            <div className="progress"></div>
          </div>
        </div>
      </div>
    </div>
  );
}
```

