---
id: 674de4ee2039496c30a31b71
title: 步骤 18
challengeType: 0
dashedName: step-18
---

# --description--

在 `#player-container` 元素之后，添加一个 `div` 元素，其 `id` 为 `playlist-container`。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应该有一个 `id` 值为 `playlist-container`。

```js
assert.equal(document.querySelector('#music-player div#playlist-container')?.getAttribute('id'), 'playlist-container');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>
          <p id="track-artist"></p>
        </div>
        <div className="controls">
          <button id="play-pause-btn">播放</button>
        </div>
        <div className="progress-container">
          <div className="progress-bar">
            <div className="progress"></div>
          </div>
        </div>
        <div className="time-info">
          <span id="current-time">0:00</span>
          <span id="duration">0:00</span>
        </div>
      </div>

    </div>
  );
}
```

