---
id: 674dda6c6ae8bb538678e125
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

在 `#current-time` 元素之后，添加一个 `span` 元素，其 `id` 为 `duration`，文本为 `0:00`。

# --hints--

你应该有一个 `span` 元素。

```js
assert.exists(document.querySelector('span'));
```

你的 `span` 元素应该有一个 `id` 值为 `duration`。

```js
assert.equal(document.querySelector('#time-info span')?.getAttribute('id'), 'duration');
```

你的 `#duration` 元素应该包含文本 `0:00`。

```js
assert.equal(document.querySelector('#duration')?.textContent, '0:00');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>
          <p id="track-artist"></p>
        </div>
        <div className="controls">
          <button id="play-pause-btn">播放</button>
        </div>
        <div className="progress-container">
          <div className="progress-bar">
            <div className="progress"></div>
          </div>
        </div>
        <div className="time-info">
          <span id="current-time">0:00</span>

        </div>
      </div>
    </div>
  );
}
```

