---
id: 674e305b7bda7c1f4f37a930
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

在 `#playlist-container` 元素内部，在 `h3` 元素之后，添加一个 `ul` 元素，其 `id` 为 `playlist`。

# --hints--

你应该有一个 `ul` 元素。

```js
assert.exists(document.querySelector('ul'));
```

你的 `ul` 元素应该有一个 `id` 值为 `playlist`。

```js
assert.equal(document.querySelector('#playlist-container ul')?.getAttribute('id'), 'playlist');
```

# --seed--

## --seed-contents--

```jsx
const { useState, useRef, useEffect } = React;

export function MusicPlayer() {

  return (
    <div id="music-player">
      <header>
        <h1>音乐播放器</h1>
      </header>
      <div id="player-container">
        <div className="album-art-container">
          <img id="album-art" />
        </div>
        <div className="track-info">
          <h2 id="track-title"></h2>
          <p id="track-artist"></p>
        </div>
        <div className="controls">
          <button id="play-pause-btn">播放</button>
        </div>
        <div className="progress-container">
          <div className="progress-bar">
            <div className="progress"></div>
          </div>
        </div>
        <div className="time-info">
          <span id="current-time">0:00</span>
          <span id="duration">0:00</span>
        </div>
      </div>
      <div id="playlist-container">
        <h3>播放列表</h3>

      </div>
    </div>
  );
}