---
id: 671fa82dc39e4c320352f371
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--
为了装下不同类型的故事，创建一个对象 `storyObj`。它需要有三个属性：`scary`、`funny`、`adventure`。每个属性的值本身又是一个对象，里面包含：`story`（故事文本）以及 `borderColor`（边框颜色）。

请按照下面的内容进行赋值：

- `scary`:
    - `story`: `In the dark woods, a group of friends stumbled upon an old, abandoned cabin. They enter the cabin and awaken something malevolent that had been dormant for centuries.`
    - `borderColor`: `#ee4b2b`
- `funny`:
    - `story`: `During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire. However, his attempt caused him to burn the dinner as well as his eyebrows off.`
    - `borderColor`: `#f1be32`
- `adventure`:
    - `story`: `Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple. They braved deadly traps and encountered strange wildlife, all while deciphering cryptic clues left behind by a mysterious civilization.`
    - `borderColor`: `#acd157`

这个对象就像“故事资料库”，等按钮按下，就能快速拿出对应内容展示。再坚持一下，你的应用已经越来越完整啦～

# --hints--
需要一个对象 `storyObj`。（用来集中管理故事数据～）

```js
assert.isObject(storyObj);
```

`storyObj` 应包含三个属性：`scary`、`funny`、`adventure`。

```js
assert.containsAllKeys(storyObj, ['scary', 'funny', 'adventure']);
```

每个属性的对象里都要有 `story` 与 `borderColor` 两个键。

```js
assert.containsAllKeys(storyObj.scary, ['story', 'borderColor']);
assert.containsAllKeys(storyObj.funny, ['story', 'borderColor']);
assert.containsAllKeys(storyObj.adventure, ['story', 'borderColor']);
```

`scary.story` 要等于：`In the dark woods, a group of friends stumbled upon an old, abandoned cabin. They enter the cabin and awaken something malevolent that had been dormant for centuries.`

```js
assert.equal(
  storyObj.scary.story.replace(/\s+/g, ' '),
  "In the dark woods, a group of friends stumbled upon an old, abandoned cabin. They enter the cabin and awaken something malevolent that had been dormant for centuries."
);
```

`scary.borderColor` 要是：`#ee4b2b`

```js
assert.equal(storyObj.scary.borderColor, "#ee4b2b");
```

`funny.story` 要等于：`During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire. However, his attempt caused him to burn the dinner as well as his eyebrows off.`

```js
assert.equal(storyObj.funny.story, "During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire. However, his attempt caused him to burn the dinner as well as his eyebrows off.");
```

`funny.borderColor` 要是：`#f1be32`

```js
assert.equal(storyObj.funny.borderColor, "#f1be32");
```

`adventure.story` 要等于：`Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple. They braved deadly traps and encountered strange wildlife, all while deciphering cryptic clues left behind by a mysterious civilization.`

```js
assert.equal(storyObj.adventure.story, "Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple. They braved deadly traps and encountered strange wildlife, all while deciphering cryptic clues left behind by a mysterious civilization.");
```

`adventure.borderColor` 要是：`#acd157`

```js
assert.equal(storyObj.adventure.borderColor, "#acd157");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Storytelling App</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>

    <h1>Want to hear a short story?</h1>

    <main class="story-container">
        <div class="btn-container">
            <button class="btn" id="scary-btn">Scary Story</button>
            <button class="btn" id="funny-btn">Funny Story</button>
            <button class="btn" id="adventure-btn">Adventure Story</button>
        </div>
        <p id="result"></p>
    </main>
    <script src="./script.js"></script>
</body>
</html>
```

```js
const storyContainer = document.querySelector(".story-container");

const scaryStoryBtn = document.getElementById("scary-btn");
const funnyStoryBtn = document.getElementById("funny-btn");
const adventureStoryBtn = document.getElementById("adventure-btn");

const resultParagraph = document.getElementById("result");

--fcc-editable-region--

--fcc-editable-region--
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --dark-grey: #1b1b32;
    --black: #000;
    --white: #fff;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
}

body {
    background-color: var(--dark-grey);
    color: var(--white);
}

h1,
#result {
    text-align: center;
}

h1 {
    margin: 10px 0 15px;
}

.story-container {
    margin: auto;
    padding: 10px;
    width: 80%;
    border-style: double;
    border-width: 14px;
    border-color: var(--white);
}

.btn-container {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
}

@media (min-width: 760px) {
    .btn-container {
        flex-direction: row;
    }
}

#result {
    margin-top: 15px;
    font-size: 1.2rem;
    line-height: 30px;
}

.btn {
    cursor: pointer;
    width: 200px;
    margin: 10px 0 10px 0.5rem;
    color: var(--black);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}
```
