---
id: 66e028680eca7d21db7e1aee
title: 构建正则表达式沙盒
challengeType: 25
dashedName: lab-regex-sandbox
demoType: onClick
---

# --description--

对于这个实验，你从已经为你写好的 CSS 和 HTML 开始。你将使用 JavaScript 来启用正则表达式沙盒，以测试正则表达式与字符串的匹配并突出显示结果。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该访问 [#pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L278-L278)、[#test-string](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L293-L293)、[#test-btn](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L295-L295) 和 [#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 元素，并分别将它们保存在变量 [regexPattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675)、[stringToTest](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)、[testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 和 [testResult](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L678-L678) 中。
1. 你应该访问 [#i](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L284-L284) 和 [#g](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L287-L287) 标志复选框，并分别将它们保存在变量 [caseInsensitiveFlag](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L676-L676) 和 [globalFlag](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L677-L677) 中。
1. 你应该有一个名为 [getFlags](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L679-L686) 的函数，该函数返回一个包含选中标志复选框标志的字符串，如果没有选中任何复选框则返回空字符串。
1. 当你点击 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 按钮时，如果正则表达式模式与测试字符串匹配，则应突出显示匹配的文本。要突出显示匹配项，请将每个匹配项替换为被具有 [highlight](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L460-L462) 类的 [span](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L107-L107) 元素包围的相同文本。请注意，CSS 已经为你添加好了。
1. 当有匹配项时，匹配的文本应显示在 [#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 中。在多个匹配项的情况下，每个匹配的文本应通过逗号和空格与下一个匹配项分隔。
1. 当没有匹配项时，文本 [no match](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L702-L702) 应显示在 [#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 中。

# --hints--

你应该访问 [#pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L278-L278) 元素并将其保存在 [regexPattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675) 变量中。

```js
assert.equal(regexPattern, document.querySelector('#pattern'))
```

你应该访问 [#test-string](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L293-L293) 元素并将其保存在 [stringToTest](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 变量中。

```js
assert.equal(stringToTest, document.querySelector('#test-string'))
```

你应该访问 [#test-btn](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L295-L295) 元素并将其保存在 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 变量中。

```js
assert.equal(testButton, document.querySelector('#test-btn'))
```

你应该访问 [#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 元素并将其保存在 [testResult](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L678-L678) 变量中。

```js
assert.equal(testResult, document.querySelector('#result'))
```

你应该访问 [#i](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L284-L284) 标志复选框并将其保存在 [caseInsensitiveFlag](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L676-L676) 中。

```js
assert.equal(caseInsensitiveFlag, document.querySelector('#i'))
```

你应该访问 [#g](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L287-L287) 标志复选框并将其保存在 [globalFlag](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L677-L677) 变量中。

```js
assert.equal(globalFlag, document.querySelector('#g'))
```

你应该有一个 [getFlags](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L679-L686) 函数。

```js
assert.isFunction(getFlags)
```

如果选中 [#i](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L284-L284) 复选框，[getFlags](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L679-L686) 函数应返回 ["i"](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L676-L676)。

```js
caseInsensitiveFlag.checked = true;
globalFlag.checked = false;
assert.equal(getFlags(), "i")
```

如果选中 [#g](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L287-L287) 复选框，[getFlags](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L679-L686) 函数应返回 ["g"](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L677-L677)。

```js
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
assert.equal(getFlags(), "g")
```

如果两个复选框都被选中，[getFlags](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L679-L686) 函数应返回 ["ig"](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L676-L676) 或 ["gi"](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L677-L677)。

```js
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const flags = getFlags();
assert.lengthOf(flags, 2);
assert.include(flags, "i");
assert.include(flags, "g");
```

如果两个复选框都未被选中，[getFlags](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L679-L686) 函数应返回空字符串。

```js
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
assert.equal(getFlags(), "")
```

当 [stringToTest](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 的内部 HTML 为 [Gu1n34 P1g5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)，[regexPattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675) 的值为 [\d+](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675)，且未选中任何标志时，通过点击 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 按钮，[stringToTest.innerHTML](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 应变为 [Gu<span class="highlight">1</span>n34 P1g5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)。

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'Gu<span class="highlight">1</span>n34 P1g5';
assert.equal(stringToTest.innerHTML, expected);
```

当 [stringToTest](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 的内部 HTML 为 [Gu1n34 P1g5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)，[regexPattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675) 的值为 [\d+](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675)，且选中全局标志时，通过点击 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 按钮，[stringToTest.innerHTML](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 应变为 [Gu<span class="highlight">1</span>n<span class="highlight">34</span> P<span class="highlight">1</span>g<span class="highlight">5</span>](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)。

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'Gu<span class="highlight">1</span>n<span class="highlight">34</span> P<span class="highlight">1</span>g<span class="highlight">5</span>';
assert.equal(stringToTest.innerHTML, expected);
```

当 [stringToTest](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 的内部 HTML 为 [Gu1n34 P1g5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)，[regexPattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675) 的值为 [G](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675)，且同时选中全局和大小写不敏感标志时，通过点击 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 按钮，[stringToTest.innerHTML](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 应变为 [<span class="highlight">G</span>u1n34 P1<span class="highlight">g</span>5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)。

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "G";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = '<span class="highlight">G</span>u1n34 P1<span class="highlight">g</span>5';
assert.equal(stringToTest.innerHTML, expected);
```

当你点击 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 按钮时，如果正则表达式模式与测试字符串匹配，则匹配的文本应被具有 [highlight](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L460-L462) 类的 [span](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L107-L107) 元素包围。

```js
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "s[he]a?";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = '<span class="highlight">Sh</span>e <span class="highlight">se</span>lls <span class="highlight">sea</span><span class="highlight">sh</span>ells by the <span class="highlight">sea</span><span class="highlight">sh</span>ore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = 'She <span class="highlight">se</span>lls <span class="highlight">sea</span><span class="highlight">sh</span>ells by the <span class="highlight">sea</span><span class="highlight">sh</span>ore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
globalFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = 'She <span class="highlight">se</span>lls seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = '<span class="highlight">Sh</span>e sells seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);
```

当没有匹配项时，测试字符串不应被修改。

```js
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "sea\\s";
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'She sells seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);
```

当 [stringToTest](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 的内部 HTML 为 [Gu1n34 P1g5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)，[regexPattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675) 的值为 [\d+](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675)，且未选中任何标志时，通过点击 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 按钮，[#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 的内部文本应变为 [1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/moment@2.30.1/node_modules/moment/src/lib/duration/constructor.js#L6-L6)。

```js
const out = document.getElementById("result");
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "1"
assert.strictEqual(out.innerText, expected)
```

当 [stringToTest](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687) 的内部 HTML 为 [Gu1n34 P1g5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)，[regexPattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675) 的值为 [\d+](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L675-L675)，且选中全局标志时，通过点击 [testButton](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L688-L688) 按钮，[#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 的内部文本应变为 [1, 34, 1, 5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L687-L687)。

```js
const out = document.getElementById("result");
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "1, 34, 1, 5"
assert.strictEqual(out.innerText, expected)
```

当有匹配项时，匹配的文本应显示在 [#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 中。

```js
const out = document.getElementById("result");
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "s[he]a?";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "Sh, se, sea, sh, sea, sh"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = "se, sea, sh, sea, sh"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
globalFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = "se"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = "Sh"
assert.strictEqual(out.innerText, expected)
```

当没有匹配项时，文本 [no match](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L702-L702) 应显示在 [#result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-regex-sandbox/66e028680eca7d21db7e1aee.md#L298-L298) 中。

```js
const out = document.getElementById("result");
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "sea\\s";
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "no match";
assert.strictEqual(out.innerText, expected);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <link rel="stylesheet" href="styles.css">
    <title>Regex Sandbox</title>

</head>

<body>
    <h1>Regex Sandbox</h1>
    <main>
        <div id="regex-container">
            <label for="pattern">Regex Pattern:
                <div id="pattern-container">/<input type="text" id="pattern" name="pattern"
                        placeholder="Enter your regex pattern">/</div>
            </label>
            <div id="flags-container">
                <p>Flags: </p>
                <label for="i">
                    <input type="checkbox" name="flags" id="i"> i
                </label>
                <label for="g">
                    <input type="checkbox" name="flags" id="g"> g
                </label>
            </div>
        </div>
        <div id="test-container">
            <p>Test String:</p>
            <div id="test-string" placeholder="Enter your test string" contenteditable="true"></div>
        </div>
        <button class="btn" id="test-btn" type="button">Test Regex</button>
        <div id="result-container">
            <h2>Result:</h2>
            <p id="result">
            </p>
        </div>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

:root {
    --dark-grey: #1b1b32;
    --light-grey: #f5f6f7;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
    --border: 0.2rem solid darkgrey;
    --padding: 0.3rem;
}

body {
    background-color: var(--dark-grey);
    color: var(--light-grey);
    font-size: 20px;
    font-family: "Lato", Helvetica, Arial, sans-serif;
    padding: 5px;
}

h1 {
    margin: 5rem auto 2rem;
    text-align: center;
}

p {
    padding: var(--padding);
}

#regex-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
    border: var(--border);
}

#regex-container>label {
    padding: var(--padding);
    flex: 1 1 auto;
}

#pattern-container {
    display: inline-block;
    color: var(--dark-grey);
    background-color: var(--light-grey);
    margin: 5px;
    border: var(--border);
}

#pattern {
    margin: 0.2rem;
    border: 0;
    font-size: 1rem;
    width: calc(100% - 1.2rem);
}

#pattern:focus {
    outline: none;
}

#flags-container {
    display: flex;
    align-items: center;
    flex: 1 1 auto;
}

#flags-container>label {
    padding: var(--padding);
    margin-right: 0.3rem;
}

#test-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    flex-direction: column;
    flex: 0 0 auto;
    border: var(--border);
}

#test-string {
    background-color: var(--light-grey);
    min-height: 5rem;
    color: var(--dark-grey);
    border-top: var(--border);
    font-size: 1.2rem;
}

[contenteditable=true]:empty:before {
    content: attr(placeholder);
    pointer-events: none;
    color: var(--dark-grey);
}

::placeholder {
    color: var(--dark-grey);
}


button {
    display: block;
    cursor: pointer;
    width: 8rem;
    margin: 0.2rem auto;
    color: var(--dark-grey);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 0.2rem;
    font-size: 1.1rem;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
}

h2 {
    align-self: flex-start;
    margin: 0.4rem 0.2rem 0.2rem;
    flex: 0 1 auto;
}

#result {
    color: var(--dark-grey);
    background-color: var(--light-grey);
    font-size: 1.5rem;
    flex: 1 1 auto;
    margin: 0.2rem;
    border: var(--border);
    min-height: 3rem;
}

.highlight {
    background-color: lightgreen;
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <link rel="stylesheet" href="styles.css">
    <title>Regex Sandbox</title>

</head>

<body>
    <h1>Regex Sandbox</h1>
    <main>
        <div id="regex-container">
            <label for="pattern">Regex Pattern:
                <div id="pattern-container">/<input type="text" id="pattern" name="pattern"
                        placeholder="Enter your regex pattern">/</div>
            </label>
            <div id="flags-container">
                <p>Flags: </p>
                <label for="i">
                    <input type="checkbox" name="flags" id="i"> i
                </label>
                <label for="g">
                    <input type="checkbox" name="flags" id="g"> g
                </label>
            </div>
        </div>
        <div id="test-container">
            <p>Test String:</p>
            <div id="test-string" placeholder="Enter your test string" contenteditable="true"></div>
        </div>
        <button class="btn" id="test-btn" type="button">Test Regex</button>
        <div id="result-container">
            <h2>Result:</h2>
            <p id="result">
            </p>
        </div>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

:root {
    --dark-grey: #1b1b32;
    --light-grey: #f5f6f7;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
    --border: 0.2rem solid darkgrey;
    --padding: 0.3rem;
}

body {
    background-color: var(--dark-grey);
    color: var(--light-grey);
    font-size: 20px;
    font-family: "Lato", Helvetica, Arial, sans-serif;
}

h1 {
    margin: 5rem auto 2rem;
    text-align: center;
}

p {
    padding: var(--padding);
}

#regex-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
    border: var(--border);
}

#regex-container>label {
    padding: var(--padding);
    flex: 1 1 auto;
}

#pattern-container {
    display: inline-block;
    color: var(--dark-grey);
    background-color: var(--light-grey);
    margin: 5px;
    border: var(--border);
}

#pattern {
    margin: 0.2rem;
    border: 0;
    font-size: 1rem;
    width: calc(100% - 1.2rem);
}

#pattern:focus {
    outline: none;
}

#flags-container {
    display: flex;
    align-items: center;
    flex: 1 1 auto;
}

#flags-container>label {
    padding: var(--padding);
    margin-right: 0.3rem;
}

#test-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    flex-direction: column;
    flex: 0 0 auto;
    border: var(--border);
}

#test-string {
    background-color: var(--light-grey);
    min-height: 5rem;
    color: var(--dark-grey);
    border-top: var(--border);
    font-size: 1.2rem;
}

[contenteditable=true]:empty:before {
    content: attr(placeholder);
    pointer-events: none;
    color: var(--dark-grey);
}

::placeholder {
    color: var(--dark-grey);
}


button {
    display: block;
    cursor: pointer;
    width: 8rem;
    margin: 0.2rem auto;
    color: var(--dark-grey);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 0.2rem;
    font-size: 1.1rem;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
}

h2 {
    align-self: flex-start;
    margin: 0.4rem 0.2rem 0.2rem;
    flex: 0 1 auto;
}

#result {
    color: var(--dark-grey);
    background-color: var(--light-grey);
    font-size: 1.5rem;
    flex: 1 1 auto;
    margin: 0.2rem;
    border: var(--border);
    min-height: 3rem;
}

.highlight {
    background-color: lightgreen;
}
```

```js
const regexPattern = document.getElementById("pattern");
const caseInsensitiveFlag = document.getElementById("i");
const globalFlag = document.getElementById("g");
const testResult = document.getElementById("result");
const getFlags = () => {
    let flags = '';
    const flagArray = [caseInsensitiveFlag, globalFlag]
    flagArray.forEach((flag) => {
        if (flag.checked) flags += flag.id;
    })
    return flags
}
const stringToTest = document.getElementById("test-string");
const testButton = document.getElementById("test-btn")
testButton.addEventListener(
    "click",
    () => {
        let matched;
        let flags = getFlags();
        const regex = new RegExp(regexPattern.value, flags)
        if (flags.includes(globalFlag)) {
            stringToTest.innerHTML = stringToTest.innerText.replaceAll(regex, '<span class="highlight">$&</span>')
            matched = stringToTest.innerText.matchAll(regex)
        } else {
            stringToTest.innerHTML = stringToTest.innerText.replace(regex, '<span class="highlight">$&</span>')
            matched = stringToTest.innerText.match(regex);             
        }
        testResult.innerText = matched ? matched.join(", ") : 'no match';
    }
)
```
