---
id: 66bb3e20d3dc5b6d0a21f5dd
title: 构建实时计数器
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

在这个实验中，你将使用 JavaScript 构建一个实时字符计数器。字符计数器将显示在 [textarea](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/readline.d.ts#L345-L345) 元素中输入的字符数量。当用户在 [textarea](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/readline.d.ts#L345-L345) 中输入时，计数器将实时更新。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.typedarray.d.ts#L46-L46) 为 [text-input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 的 [textarea](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/readline.d.ts#L345-L345) 元素。
2. 应该有一个 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.typedarray.d.ts#L46-L46) 为 [char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L107-L107) 元素。它最初应包含文本 [Character Count: 0/50](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137)。这个占位符文本应该被 [textarea](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/readline.d.ts#L345-L345) 元素中当前字符计数替换。
3. 当 [#text-input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 元素包含文本 [hello coder](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 时，[#char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 元素应包含文本 "Character Count: 11/50"
4. 当 [#text-input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 元素包含文本 [i am a programmer](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 时，[#char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 元素应包含文本 "Character Count: 17/50"
5. 当 [#text-input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 元素包含文本 [hello world](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 时，[#char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 元素应包含文本 "Character Count: 11/50"
6. 当 [#text-input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 元素包含文本 [I am learning a new language and it is called c++.](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 时，[#char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 元素应包含文本 "Character Count: 50/50"。**注意**：虽然 [maxlength](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/readline.d.ts#L347-L347) 属性可以实现正确的功能，但此实验要求仅使用 JavaScript 解决方案。
7. 用户不应能够在 [textarea](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/readline.d.ts#L345-L345) 元素中输入超过 [50](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/ts4.8/test/node_tests/cpu_prof_interval.js#L4-L4) 个字符。当字符计数达到 [50](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/ts4.8/test/node_tests/cpu_prof_interval.js#L4-L4) 时，任何额外输入都应自动修剪，并且文本 [Character Count: 50/50](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 应以红色显示。

**注意：**确保在 HTML 中链接你的样式表和 JavaScript 文件。

# --hints--

你应该有一个 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.typedarray.d.ts#L46-L46) 为 [text-input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 的 [textarea](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/readline.d.ts#L345-L345) 元素。

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

你应该有一个 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.typedarray.d.ts#L46-L46) 为 [char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L107-L107) 元素。

```js
const charCount = document.querySelector('p');
assert.equal(charCount.id, 'char-count');
```

[#char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 元素最初应包含文本 [Character Count: 0/50](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137)。

```js
const charCount = document.querySelector('p')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

当 [#text-input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 元素包含文本 [hello coder](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L136-L136) 时，[#char-count](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-real-time-counter/66bb3e20d3dc5b6d0a21f5dd.md#L137-L137) 元素应包含文本 "Character Count: 11/50"

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

当字符计数为 [50](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/ts4.8/test/node_tests/cpu_prof_interval.js#L4-L4) 时，文本应以红色显示。

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

如果字符计数大于或等于 [50](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/ts4.8/test/node_tests/cpu_prof_interval.js#L4-L4)，则应修剪额外输入。

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
