---
id: 685dbd19200ad656d48b77ff
title: Build a Music Shopping Cart Page
challengeType: 25
dashedName: lab-music-shopping-cart-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该有一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)为[shopping-cart-container](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L484-L486)的元素。
2. 您的#shopping-cart-container元素应该使用正确的工具类将元素设置为Flexbox布局。
3. 您的#shopping-cart-container元素应该为较小的设备设置flex项目的方向为列，为较大的设备设置为行。请记住，Tailwind CSS采用移动优先的方法，您将使用[lg:](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L61-L62)断点前缀来针对较大的设备。
4. 在您的#shopping-cart-container元素内，您应该有一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)为[products-container](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L487-L487)的元素。
5. 您的#products-container元素应该至少有两个子元素，每个子元素都有一个名为[card](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L489-L490)的类。
6. 每个.card元素内部应该嵌套以下元素：
   - 一个h2元素，包含代表产品名称的文本，以及一个使用预定义大小类（如text-sm、text-md、text-lg、text-xl、text-2xl等）设置字体大小的工具类。
   - 一个带有[quantity](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L500-L500)类的元素和表示该产品购物车项目数量的文本。
   - 一个带有[price](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L509-L509)类的元素和表示价格的文本。
   - 一个带有[remove-button](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L511-L513)类和"Remove"文本的button。您的button应该有用于设置预定义红色背景颜色的工具类，以及用于悬停状态的不同预定义红色背景颜色的工具类。预定义红色背景颜色的示例包括bg-red-500、bg-red-600等。预定义悬停红色背景颜色的示例包括hover:bg-red-500、hover:bg-red-600等。
7. 在您的#shopping-cart-container元素内，您应该有一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)为[order-summary-container](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L610-L612)的元素。
8. 您的#order-summary-container元素应该具有以下样式：
   - 您选择的用于预定义圆角的工具类。示例包括rounded、rounded-lg、rounded-full等。
   - 您选择的用于在所有边上设置边框宽度的工具类。
9. 您的#order-summary-container元素内部应该嵌套以下元素：
   - 一个带有"Order Summary"文本的h2元素，以及您选择的用于设置字体大小的工具类。
   - 一个带有"Total:"文本且[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)设置为[total](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L624-L624)的元素。该元素还应使用工具类将字体粗细设置为预定义的Tailwind CSS字体粗细，以及您为元素选择的字体大小。
   - 一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)设置为[total-amount](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L625-L625)的元素，用于显示购物车中所有项目的总计。
   - 一个带有"Checkout"文本和用于居中文本的正确工具类的链接。href值应设置为#。您的链接还应有一个工具类，用于将背景颜色设置为您选择的预定义Tailwind CSS蓝色，并为悬停状态设置不同的蓝色背景颜色。

# --hints--

您应该有一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)设置为[shopping-cart-container](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L484-L486)的元素。

```js
assert.exists(document.getElementById("shopping-cart-container"));
```

您的#shopping-cart-container元素应该使用[flex](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L47-L48)工具类将元素设置为Flexbox布局。

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("flex"));
```

您的#shopping-cart-container元素应该使用正确的工具类将flex项目的方向设置为列。

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("flex-col"));
```

您的#shopping-cart-container元素应该使用[lg:](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L61-L62)断点前缀，后跟正确的工具类，为较大的设备将flex项目的方向设置为行。

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.isTrue(shoppingCartContainer.classList.contains("lg:flex-row"));
```

您应该在#shopping-cart-container元素内有一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)设置为[products-container](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L487-L487)的元素。

```js
const shoppingCartContainer = document.getElementById("shopping-cart-container");
assert.exists(shoppingCartContainer);
const productsContainer = document.getElementById("products-container");
assert.exists(productsContainer);
assert.isTrue(shoppingCartContainer.contains(productsContainer));
```

您的#products-container元素应该至少有两个子元素，每个子元素都有一个名为[card](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L489-L490)的类。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
assert.isAtLeast(cards.length, 2);
```

每个.card元素应该有一个包含代表产品名称文本的h2元素。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const productName = card.querySelector("h2");
  assert.exists(productName);
  assert.isNotEmpty(productName.textContent);
});
```

每个.card元素内的h2元素应该使用工具类设置字体大小，仅使用预定义的大小类，如text-sm、text-base、text-lg、text-xl、text-2xl等。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

cards?.forEach((card) => {
  const productName = card.querySelector("h2");
  assert.exists(productName);
  const hasPredefinedSize = predefinedSizes.some(size => productName.classList.contains(size));
  assert.isTrue(hasPredefinedSize);
});
```

每个.card元素应该有一个带有[quantity](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L500-L500)类的元素。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const quantity = card.querySelector(".quantity");
  assert.exists(quantity);
});
```

每个.quantity元素应该有代表该产品购物车项目数量的文本。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const quantity = card.querySelector(".quantity");
  assert.exists(quantity);
  assert.isNotEmpty(quantity.textContent.trim());
});
```

每个.card元素应该有一个带有[price](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L509-L509)类的元素。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");  

cards?.forEach((card) => {
  const price = card.querySelector(".price");
  assert.exists(price);
});
```

每个.price元素应该有代表价格的文本。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const price = card.querySelector(".price");
  assert.exists(price);
  assert.isNotEmpty(price.textContent.trim());
});
```

每个.card元素应该有一个带有[remove-button](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L511-L513)类和"Remove"文本的button元素。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);
  assert.equal(removeButton.textContent.trim(), "Remove");
});
```

每个button元素应该有您选择的红色背景颜色的工具类。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");
const predefinedRedColors = [
  "bg-red-50",
  "bg-red-100",
  "bg-red-200",
  "bg-red-300",
  "bg-red-400",
  "bg-red-500",
  "bg-red-600",
  "bg-red-700",
  "bg-red-800",
  "bg-red-900",
  "bg-red-950"
];

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);
  const hasRedBackground = predefinedRedColors.some(color => removeButton.classList.contains(color));
  assert.isTrue(hasRedBackground);
});
```

每个button元素应该有不同的悬停状态红色背景颜色。

```js
const productsContainer = document.getElementById("products-container");
const cards = productsContainer.querySelectorAll("#products-container .card");

const predefinedRedHoverColors = [
  "hover:bg-red-50",
  "hover:bg-red-100",
  "hover:bg-red-200",
  "hover:bg-red-300",
  "hover:bg-red-400",
  "hover:bg-red-500",
  "hover:bg-red-600",
  "hover:bg-red-700",
  "hover:bg-red-800",
  "hover:bg-red-900",
  "hover:bg-red-950"
];

cards?.forEach((card) => {
  const removeButton = card.querySelector("button.remove-button");
  assert.exists(removeButton);

  const classList = removeButton.classList;

  const hoverBgClass = [...classList].find(cls => cls.startsWith("hover:bg-red-"));
  const bgClass = [...classList].find(cls => cls.startsWith("bg-red-") && !cls.startsWith("hover:"));

  assert.notEqual(bgClass, hoverBgClass.replace("hover:", ""));
});
```

您应该在#shopping-cart-container元素内有一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)为[order-summary-container](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L610-L612)的元素。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
assert.exists(orderSummaryContainer);
```

您的#order-summary-container元素应该有您选择的用于设置圆角的工具类。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

const hasRoundedCorners = predefinedRoundedCorners.some(corner => orderSummaryContainer.classList.contains(corner));
assert.isTrue(hasRoundedCorners);
```

您的#order-summary-container元素应该有您选择的用于在所有边上设置边框宽度的工具类。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");

// border classes can vary in width, so this check is for the presence of any border class
const hasBorder = [...orderSummaryContainer.classList].some(className => className.startsWith("border-") || className === "border");
assert.isTrue(hasBorder);
```

您的#order-summary-container元素应该有一个带有"Order Summary"文本的h2元素。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const orderSummaryTitle = orderSummaryContainer.querySelector("h2");
assert.exists(orderSummaryTitle);
assert.equal(orderSummaryTitle.textContent.trim(), "Order Summary");
```

您的#order-summary-container元素内的h2应该使用您选择的用于设置字体大小的工具类。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

const orderSummaryTitle = orderSummaryContainer.querySelector("h2");
assert.exists(orderSummaryTitle);
const hasPredefinedSize = predefinedSizes.some(size => orderSummaryTitle.classList.contains(size));
assert.isTrue(hasPredefinedSize);
```

您的#order-summary-container元素应该有一个带有"Total:"文本且[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)设置为[total](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L624-L624)的元素。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
assert.exists(totalElement);
assert.equal(totalElement.textContent.trim(), "Total:");
```

您的#total元素应该使用工具类设置字体粗细。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

const hasFontWeight = predefinedFontWeights.some(weight => totalElement.classList.contains(weight));
assert.isTrue(hasFontWeight);
```

您的#total元素应该使用工具类设置字体大小。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total");
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

const hasPredefinedSize = predefinedSizes.some(size => totalElement.classList.contains(size));
assert.isTrue(hasPredefinedSize);
```

您的#order-summary-container元素应该有一个[id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L40-L41)设置为[total-amount](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L625-L625)的元素，用于显示购物车中所有项目的总计。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const totalElement = orderSummaryContainer.querySelector("#total-amount");
assert.exists(totalElement);
assert.isNotEmpty(totalElement.textContent.trim());
```

您的#order-summary-container元素应该有一个带有"Checkout"文本和[href](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-music-shopping-cart-page/685dbd19200ad656d48b77ff.md#L390-L390)值为#的链接。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
assert.exists(checkoutLink);
assert.equal(checkoutLink.textContent.trim(), "Checkout");
assert.equal(checkoutLink.getAttribute("href"), "#");
```

您的Checkout链接应该有用于居中文本的工具类。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const hasTextCentering = checkoutLink.classList.contains("text-center");
assert.isTrue(hasTextCentering);
```

您的Checkout链接应该有用于将背景颜色设置为您选择的预定义Tailwind CSS蓝色的工具类。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const predefinedBlueColors = [
  "bg-blue-50",
  "bg-blue-100",
  "bg-blue-200",
  "bg-blue-300",
  "bg-blue-400",
  "bg-blue-500",
  "bg-blue-600",
  "bg-blue-700",
  "bg-blue-800",
  "bg-blue-900",
  "bg-blue-950"
];

const hasBlueBackground = predefinedBlueColors.some(color => checkoutLink.classList.contains(color));
assert.isTrue(hasBlueBackground);
```

您的Checkout链接应该有不同的悬停状态蓝色背景颜色。

```js
const orderSummaryContainer = document.querySelector("#shopping-cart-container #order-summary-container");
const checkoutLink = orderSummaryContainer.querySelector("a");
const predefinedBlueHoverColors = [
  "hover:bg-blue-50",
  "hover:bg-blue-100",
  "hover:bg-blue-200",
  "hover:bg-blue-300",
  "hover:bg-blue-400",
  "hover:bg-blue-500",
  "hover:bg-blue-600",
  "hover:bg-blue-700",
  "hover:bg-blue-800",
  "hover:bg-blue-900",
  "hover:bg-blue-950"
];

const classList = checkoutLink.classList;
const hoverBgClass = [...classList].find(cls => cls.startsWith("hover:bg-blue-"));
const bgClass = [...classList].find(cls => cls.startsWith("bg-blue-") && !cls.startsWith("hover:"));

assert.notEqual(bgClass, hoverBgClass.replace("hover:", ""));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music Shopping Cart Page</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music Shopping Cart Page</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-900 text-white">
    <h1 class="text-center text-3xl my-8">Shopping Cart Page</h1>
    <main
      id="shopping-cart-container"
      class="flex flex-col lg:flex-row items-center gap-8"
    >
      <div id="products-container" class="w-full md:w-2/3 px-4">
        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">Ludwig Supraphonic Snare Drum Chrome</h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">2</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1398.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">Shure SM58S Mic With Switch Standard</h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$119.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">
            Fender American Professional II Stratocaster
          </h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1700.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>

        <div
          class="card flex flex-col lg:flex-row items-center gap-6 rounded-lg border border-gray-200 p-4 my-8 bg-gray-800"
        >
          <h2 class="text-lg lg:w-1/2">
            Moog Subsequent 37 Analog Synthesizer
          </h2>
          <div class="btn-container flex items-center lg:w-1/4">
            <button
              class="decrement-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Decrease quantity"
            >
              <span aria-hidden="true">-</span>
            </button>
            <span class="quantity mx-2">1</span>
            <button
              class="increment-btn bg-gray-800 hover:bg-gray-700 border border-gray-200 py-1 px-2 rounded-md"
              type="button"
              aria-label="Increase quantity"
            >
              <span aria-hidden="true">+</span>
            </button>
          </div>
          <p class="price lg:w-1/6 text-right">$1799.00</p>
          <button
            class="remove-button bg-red-700 hover:bg-red-800 py-1 px-2 rounded-md"
          >
            Remove
          </button>
        </div>
      </div>

      <div
        id="order-summary-container"
        class="rounded-lg border border-gray-200 p-8 bg-gray-800 w-full md:w-1/3 mt-8 lg:mt-0"
      >
        <h2 class="text-2xl mb-4">Order Summary</h2>
        <ul>
          <li>
            <span id="subtotal" class="font-semibold text-lg">Subtotal:</span>
            <span>$5,016.00</span>
          </li>
          <li class="my-6">
            <span id="shipping" class="font-semibold text-lg">Shipping:</span>
            <span>Calculated in checkout</span>
          </li>
          <li>
            <span id="total" class="font-semibold text-lg">Total:</span>
            <span id="total-amount">$5,016.00</span>
          </li>
        </ul>
        <a
          href="#"
          class="text-center rounded-lg mx-auto mt-4 block px-5 py-2.5 font-medium bg-blue-600 hover:bg-blue-700"
          >Checkout</a
        >
      </div>
    </main>
  </body>
</html>
```
