---
id: 66ea8239b9eeb76f12299030
title: 步骤 4
challengeType: 1
dashedName: step-4
---

# --description--

工作坊的下一部分内容是学习如何在两个特定数值之间生成随机数。

首先添加另一个 `console.log()`，将信息 "现在，让我们在两个数值之间生成一个随机数。" 输出到控制台。

然后，创建一个名为 `min` 的变量，为其赋值 `1`，再创建一个名为 `max` 的变量，为其赋值 `100`。

下一步，我们将在这些值之间生成一个随机数！

# --hints--

你应该有一个 `console.log()`，输出信息 "现在，让我们在两个数值之间生成一个随机数。"

```js
assert.match(code, /console\.log\((`|'|")\s*现在，\s*让\s*我们\s*在\s*两\s*个\s*数\s*值\s*之\s*间\s*生\s*成\s*一\s*个\s*随\s*机\s*数\.\s*\1\);?/);
```

你应该有一个名为 `min` 的变量。

```js
assert.isNotNull(min);
```

你的 `min` 变量应该是一个数字。

```js
assert.isNumber(min);
```

你应将数字 `1` 赋值给 `min` 变量。

```js
assert.equal(min, 1);
```

你应该有一个名为 `max` 的变量。

```js
assert.isNotNull(max);
```

你的 `max` 变量应该是一个数字。

```js
assert.isNumber(max);
```

你应将数字 `100` 赋值给 `max` 变量。

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```