---
id: 66ea7e7399244c6d4b6052a1
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

在工作坊的这一部分，我们将一起探索 `Math.random()` 方法的奇妙之处。

让我们先添加一行代码，用 `console.log()` 将这句话输出到控制台："Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。"

# --hints--

你应该使用 `console.log()` 输出字符串 "Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。"

```js
assert.match(code, /console\.log\((`|'|")\s*Math\.random\(\)\s+方法\s+会\s+返回\s+一个\s+在\s+0\s+到\s+小于\s+1\s+范围\s+内\s+的\s+伪\s+随机数\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```