---
id: 66eb3ec3b07dfe8f898646c0
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

正如你在之前的课程中学到的，`Math.floor()` 方法会将数值向下舍入到最接近的整数。

```js
const price = 10.99;
Math.floor(price); // 10
```

创建一个名为 `numRoundedDown` 的变量，将浮点数 `6.7` 向下舍入到最接近的整数后的结果赋值给它。

然后，将 `numRoundedDown` 变量输出到控制台以查看结果。

# --hints--

你应该有一个 `numRoundedDown` 变量。

```js
assert.isNotNull(numRoundedDown);
```

`numRoundedDown` 变量应为数字。

```js
assert.isNumber(numRoundedDown);
```

你的 `numRoundedDown` 变量应等于将浮点数 `6.7` 向下舍入到最接近的整数。

```js
assert.equal(numRoundedDown, 6);
```

你不应将 `numRoundedDown` 变量的值硬编码为 `6`。

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

你应将 `numRoundedDown` 变量输出到控制台。

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("Math.floor() 方法会将数值向下舍入到最接近的整数。");
--fcc-editable-region--

--fcc-editable-region--
```