---
id: 66eb488c0129e192542f0865
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

在之前的课程中，你已经学会了如何使用 `Math.ceil()` 方法，就像这样：

```js
const price = 10.01;
Math.ceil(price); // 11
```

`Math.ceil()` 方法会将数值向上舍入到最接近的整数。

创建一个名为 `numRoundedUp` 的变量，将浮点数 `3.2` 向上舍入到最接近的整数后的结果赋值给它。

然后，将 `numRoundedUp` 变量输出到控制台以查看结果。

# --hints--

你应该有一个 `numRoundedUp` 变量。

```js
assert.isNotNull(numRoundedUp);
```

`numRoundedUp` 变量应为数字。

```js
assert.isNumber(numRoundedUp);
```

你的 `numRoundedUp` 变量应等于将浮点数 `3.2` 向上舍入到最接近的整数。

```js
assert.equal(numRoundedUp, 4);
```

你不应将 `numRoundedUp` 变量的值硬编码为 `4`。

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

你应将 `numRoundedUp` 变量输出到控制台。

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("Math.floor() 方法会将数值向下舍入到最接近的整数。");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Math.ceil() 方法会将数值向上舍入到最接近的整数。");

--fcc-editable-region--

--fcc-editable-region--
```