---
id: 66eb499511810e932ca7eecc
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

工作坊的下一部分将学习如何将数字四舍五入到最接近的整数。

让我们先添加一个 `console.log()`，输出这句话到控制台："Math.round() 方法会将数值四舍五入到最接近的整数。"

# --hints--

你应该将这句话输出到控制台："Math.round() 方法会将数值四舍五入到最接近的整数。"

```js
assert.match(code, /console\.log\((['"])Math\.round\(\)\s+方法\s+会\s+将\s+数\s+值\s+四\s+舍\s+五\s+入\s+到\s+最\s+接\s+近\s+的\s+整\s+数\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("Math.floor() 方法会将数值向下舍入到最接近的整数。");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Math.ceil() 方法会将数值向上舍入到最接近的整数。");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

--fcc-editable-region--

--fcc-editable-region--
```