---
id: 66ece147323c91f654bf8dd9
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

在之前的课程中，你已经学会了如何从一组数字中获取最大值和最小值，就像这样：

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

创建一个名为 `maxNum` 的变量，将 `3, 125, 55, 24` 中的最大数字赋值给它。然后，将 `maxNum` 的值输出到控制台。

在下方创建一个名为 `minNum` 的变量，将 `6, 90, 14, 90, 2` 中的最小数字赋值给它。然后，将 `minNum` 的值输出到控制台。

# --hints--

你应该有一个名为 `maxNum` 的变量。

```js
assert.isNotNull(maxNum);
```

你的 `maxNum` 变量应该等于 `3, 125, 55, 24` 中的最大数字。

```js
assert.equal(maxNum, 125);
```

你不应该将 `125` 这个值直接写入 `maxNum` 变量中。请确保使用 `Math.max()` 方法。

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

你应该将 `maxNum` 的值输出到控制台。

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

你应该有一个名为 `minNum` 的变量。

```js
assert.isNotNull(minNum);
```

你的 `minNum` 变量应该等于 `6, 90, 14, 90, 2` 中的最小数字。

```js
assert.equal(minNum, 2);
```

你不应该将 `2` 这个值直接写入 `minNum` 变量中。请确保使用 `Math.min()` 方法。

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

你应该将 `minNum` 的值输出到控制台。

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("Math.floor() 方法会将数值向下舍入到最接近的整数。");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Math.ceil() 方法会将数值向上舍入到最接近的整数。");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "Math.round() 方法会将数值四舍五入到最接近的整数。"
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("Math.max() 和 Math.min() 方法用于从一组数字中获取最大值和最小值。");

--fcc-editable-region--

--fcc-editable-region--
```