---
id: 66ece0b518f5dcf5930612ab
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

在工作坊的最后一部分，我们将学习如何获取一组数字中的最小值和最大值。

让我们先添加一个 `console.log()`，输出这句话到控制台："Math.max() 和 Math.min() 方法用于从一组数字中获取最大值和最小值。"

# --hints--

你应该将这句话输出到控制台："Math.max() 和 Math.min() 方法用于从一组数字中获取最大值和最小值。"

```js
assert.match(code, /console\.log\((['"`])Math\.max\(\)\s+和\s+Math\.min\(\)\s+方法\s+用\s+于\s+从\s+一\s+组\s+数\s+字\s+中\s+获\s+取\s+最\s+大\s+值\s+和\s+最\s+小\s+值\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("Math.floor() 方法会将数值向下舍入到最接近的整数。");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Math.ceil() 方法会将数值向上舍入到最接近的整数。");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "Math.round() 方法会将数值四舍五入到最接近的整数。"
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);
--fcc-editable-region--

--fcc-editable-region--
```