---
id: 66ece75c0b840df77c7e20f0
title: 步骤 14
challengeType: 1
dashedName: step-14
---

# --description--

在工作坊的最后一步，你需要将这句话输出到控制台："和你一起学习各种 Math 方法真是太有趣了！"

完成了这个最后的修改，你就成功完成了 MathBot 工作坊！

# --hints--

你应该将这句话输出到控制台："和你一起学习各种 Math 方法真是太有趣了！"

```js
assert.match(code, /console\.log\((['"`])和\s*你\s*一\s*起\s*学\s*习\s*各\s*种\s*Math\s*方\s*法\s*真\s*是\s*太\s*有\s*趣\s*了\!\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("Math.floor() 方法会将数值向下舍入到最接近的整数。");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Math.ceil() 方法会将数值向上舍入到最接近的整数。");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "Math.round() 方法会将数值四舍五入到最接近的整数。"
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("Math.max() 和 Math.min() 方法用于从一组数字中获取最大值和最小值。");

const maxNum = Math.max(3, 125, 55, 24);
console.log(maxNum);
const minNum = Math.min(6, 90, 14, 90, 2);
console.log(minNum);

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("Math.floor() 方法会将数值向下舍入到最接近的整数。");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Math.ceil() 方法会将数值向上舍入到最接近的整数。");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log("Math.round() 方法会将数值四舍五入到最接近的整数。");

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("Math.max() 和 Math.min() 方法用于从一组数字中获取最大值和最小值。");

const maxNum = Math.max(3, 125, 55, 24);
console.log(maxNum);
const minNum = Math.min(6, 90, 14, 90, 2);
console.log(minNum);

console.log("和你一起学习各种 Math 方法真是太有趣了！");
```