---
id: a3bfc1673c0526e06d3ac698
title: 构建一个质数和计算器
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

在这个实验中，您将构建一个计算器，它接受一个数字并返回小于或等于该数字的所有质数的和。

**目标**：实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该有一个 `sumPrimes` 函数，该函数接受一个数字作为参数。
2. `sumPrimes` 函数应该返回小于或等于所提供数字的所有质数的和。
3. 如果输入数字小于 `2`，函数应该返回 `0`。

# --hints--

您应该有一个 `sumPrimes` 函数。

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` 应该返回 `17`。

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` 应该返回 `10`。

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` 应该返回 `2`。

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` 应该返回 `0`。

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` 应该返回 `73156`。

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```