---
id: 68ce56d54d3cad1149a1c4cf
title: Design a Feature Selection Page
challengeType: 25
dashedName: design-a-feature-selection-page
demoType: onClick
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你的页面应该有一个包含文本`Feature Selection`的`h1`元素。
2. 你的页面应该有一个带有`feature-card-container`类的`div`元素。
3. 你的页面应该在`div`内至少有两个带有`feature-card`类的`label`元素。
4. 你的每个`label`元素都应该包含标签文本和一个类型为`checkbox`的`input`元素。
5. 创建一个选择器来定位复选框，并应用以下样式：
    - 你的所有复选框元素在CSS中应该设置为`appearance: none;`。
    - 你的所有复选框元素都应该有你选择的边框宽度、颜色和样式。
6. 当复选框被选中时，应在复选框中心显示一个勾号`✓`。
7. 当复选框被选中时，复选框的背景颜色应更改为你选择的颜色。
8. 当复选框被选中时，复选框的边框颜色应更改为你选择的颜色。

**注意：** 确保在HTML中链接你的CSS文件。

# --hints--

你应该有一个`h1`元素。

```js
assert.exists(document.querySelector('h1'));
```

你的`h1`应该有`Feature Selection`文本。

```js
assert.equal(document.querySelector('h1').innerText.trim(), `Feature Selection`);
```

你应该有一个带有`feature-card-container`类的`div`元素。

```js
assert.exists(document.querySelector('div.feature-card-container'));
```

你应该在`.feature-card-container`内至少有两个带有`feature-card`类的`label`元素。

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isAtLeast(labels.length, 2);
```

在每个`label`元素内你应该有标签文本。

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isNotEmpty(label.innerText.trim()));
```

在每个label元素内应该有一个类型为`checkbox`的`input`元素。

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.exists(label.querySelector('input[type="checkbox"]')));
```

你应该有一个复选框的选择器。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你的复选框应该设置为`appearance: none`。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.equal(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('appearance'), 'none');
```

你的复选框应该有你选择的边框宽度。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-width'));
```

你的复选框应该有你选择的边框颜色。

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-color'));
```

你的复选框应该有你选择的边框样式。

```js
const selectors = [  
  '.feature-card [type="checkbox"]',
  '.feature-card input[type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input', // 没有其他input元素，所以这是一个有效的方法
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-style'));
```

当选中复选框时，复选框中心应该有一个勾号`✓`。

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;

    checkbox.checked = true;
    const after = window.getComputedStyle(checkbox, ':after');
    assert.equal(after.content, '"✓"');
});
```

当复选框被选中时，复选框的背景颜色应更改为你选择的颜色。

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBG = window.getComputedStyle(checkbox).backgroundColor;
    checkbox.checked = true;
    const checkedBG = window.getComputedStyle(checkbox).backgroundColor;

    assert.notEqual(checkedBG, uncheckedBG);
});
```

当复选框被选中时，复选框的边框颜色应更改为你选择的颜色。

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBC = window.getComputedStyle(checkbox).borderColor;
    checkbox.checked = true;
    const checkedBC = window.getComputedStyle(checkbox).borderColor;

    assert.notEqual(checkedBC, uncheckedBC);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <title>功能选择页面</title>
    <link rel="stylesheet" href="styles.css" />
</head>
<body>
    <h1>Feature Selection</h1>
    <div class="feature-card-container">
        <label class="feature-card">
            <input type="checkbox" />
            <h4>云存储</h4>
            <p>100吉字节安全存储</p>
        </label>
        
        <label class="feature-card">
            <input type="checkbox" />
            <h4>专属支持</h4>
            <p>24/7客户帮助</p>
        </label>
        
        <label class="feature-card">
            <input type="checkbox" />
            <h4>高级分析</h4>
            <p>洞察与报告</p>
        </label>
    </div>
</body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  flex-direction: column;
  height: 100vh;
  font-family: Arial, sans-serif;
  margin: 0;
  background-color: #f5f5f5;
}

.feature-card-container {
  display: flex;
  flex-wrap: wrap;
  gap: 20px;
  max-width: 800px;
  justify-content: center;
  padding: 20px;
}

.feature-card-container > * {
  flex: 1 1 200px;
  max-width: 250px;
}

.feature-card {
  border: 2px solid #ddd;
  border-radius: 12px;
  padding: 20px;
  text-align: center;
  position: relative;
  cursor: pointer;
  background-color: white;
  box-shadow: 0 2px 5px rgba(0,0,0,0.1);
  transition: all 0.3s ease;
}

.feature-card:hover {
  border-color: #4CAF50;
  box-shadow: 0 4px 12px rgba(76, 175, 80, 0.3);
  transform: translateY(-3px);
}

.feature-card input[type='checkbox'] {
  position: absolute;
  top: 10px;
  right: 10px;
  width: 24px;
  height: 24px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #ccc;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s ease;
}

.feature-card input[type='checkbox']:checked {
  background-color: #4CAF50;
  border-color: #4CAF50;
}

.feature-card input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  font-size: 16px;
  line-height: 20px;
}

.feature-card h4 {
  margin: 15px 0 10px;
  font-size: 1.3rem;
  color: #333;
}

.feature-card p {
  margin: 0;
  color: #666;
  font-size: 0.9rem;
}

@media (max-width: 600px) {
  body {
    height: auto;
  }
  
  .feature-card-container {
    flex-direction: column;
    align-items: center;
  }
  
  .feature-card-container > * {
    flex: none;
    width: 90%;
  }
}
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>功能选择</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>功能选择</h1>
    <div class="feature-card-container">
      <label class="feature-card">
        <input type="checkbox" />
        <h4>云存储</h4>
        <p>100吉字节安全存储</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>专属支持</h4>
        <p>24/7客户帮助</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>高级分析</h4>
        <p>洞察与报告</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>自定义用户主题</h4>
        <p>个性化仪表板设计</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>多用户协作</h4>
        <p>团队访问和共享</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>API访问</h4>
        <p>与您自定义构建的工具集成</p>
      </label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  flex-direction: column;
  height: 100vh;
  font-family: Arial, sans-serif;
  margin: 0;
  background-color: #f5f5f5;
}

.feature-card-container {
  display: flex;
  flex-wrap: wrap;
  gap: 20px;
  max-width: 800px;
  justify-content: center;
  padding: 20px;
}

.feature-card-container > * {
  flex: 1 1 200px;
  max-width: 250px;
}

.feature-card {
  border: 2px solid #ddd;
  border-radius: 12px;
  padding: 20px;
  text-align: center;
  position: relative;
  cursor: pointer;
  background-color: white;
  box-shadow: 0 2px 5px rgba(0,0,0,0.1);
  transition: all 0.3s ease;
}

.feature-card:hover {
  border-color: #4CAF50;
  box-shadow: 0 4px 12px rgba(76, 175, 80, 0.3);
  transform: translateY(-3px);
}

.feature-card input[type='checkbox'] {
  position: absolute;
  top: 10px;
  right: 10px;
  width: 24px;
  height: 24px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #ccc;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s ease;
}

.feature-card input[type='checkbox']:checked {
  background-color: #4CAF50;
  border-color: #4CAF50;
}

.feature-card input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  font-size: 16px;
  line-height: 20px;
}

.feature-card h4 {
  margin: 15px 0 10px;
  font-size: 1.3rem;
  color: #333;
}

.feature-card p {
  margin: 0;
  color: #666;
  font-size: 0.9rem;
}

@media (max-width: 600px) {
  body {
    height: auto;
  }
  
  .feature-card-container {
    flex-direction: column;
    align-items: center;
  }
  
  .feature-card-container > * {
    flex: none;
    width: 90%;
  }
}
```