---
id: 68b704862dd6559e890c2a86
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

在 `if` 语句之后，创建一个名为 `result` 的变量。这将是从此方法返回的字符串。

所以给它分配一个格式为 `Media Catalogue (N items):` 的 f-string。将 `N` 替换为 `items` 列表中的项目数量，并使字符串以两个换行符结尾。

# --hints--

你应该有一个变量名为 `result` 在你的 `__str__` 方法中。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_variable("result")`)
})
```

你应该将 `f'Media Catalogue ({len(self.items)} items):\n\n'` 分配给 `result` 变量。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_stmt("result = f'Media Catalogue ({len(self.items)} items):\\\\n\\\\n'")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)

    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'

        result = f'Media Catalogue ({len(self.items)} items):\n\n'

        for i, movie in enumerate(self.items, 1):
            result += f'{i}. {movie}\n'
        return result

try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')