---
id: 68b7e1818bd294974b9099bf
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

在声明 `result` 变量之后，创建一个 `for` 循环，使用枚举来遍历目录中的项目，并将格式为 `N. movie` 的字符串追加到 `result` 中，其中 `N` 是从 `1` 开始的数字索引，`movie` 是项目的字符串表示。

使字符串以换行符结尾，并在循环后返回 `result`。

# --hints--

你的 `__str__` 方法中应该有一个 `for` 循环。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").find_for_loops()[0]`)
})
```

你的 `for` 循环应该使用 `enumerate` 函数。

```js
({
  test: () => runPython(`assert "enumerate(" in str(_Node(_code).find_class("MediaCatalogue").find_function("__str__").find_for_loops()[0].find_for_iter())`)
})
```

你的 `__str__` 方法应该返回 `result`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_return("result")`)
})
```

当 `self.items` 包含 `Movie('Dances with Wolves', 1990, 'Kevin Costner', 224)` 和 `Movie('Annie Hall', 1977, 'Woody Allen', 93)` 时，你的 `__str__` 方法应该返回 `Media Catalogue (2 items):

1. Dances with Wolves (1990) - 224 min, Kevin Costner
2. Annie Hall (1977) - 93 min, Woody Allen
`。

```js
({
  test: () => runPython(`
    import io
    import sys

    captured_output = io.StringIO()
    sys.stdout = captured_output

    _catalogue = MediaCatalogue()
    _catalogue.add(Movie('Dances with Wolves', 1990, 'Kevin Costner', 224))
    _catalogue.add(Movie('Annie Hall', 1977, 'Woody Allen', 93))
    print(_catalogue)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    _expected = """Media Catalogue (2 items):

    1. Dances with Wolves (1990) - 224 min, Kevin Costner
    2. Annie Hall (1977) - 93 min, Woody Allen
    """
    assert _expected in output
  `)
})
```

当 `self.items` 包含 `Movie('Barry Lyndon', 1975, 'Stanley Kubrick', 184)` 和 `Movie('Fahrenheit 451', 1966, 'Francois Truffaut', 112)` 时，你的 `__str__` 方法应该返回 `Media Catalogue (2 items):

1. Barry Lyndon (1975) - 184 min, Stanley Kubrick
2. Fahrenheit 451 (1966) - 112 min, Francois Truffaut
`。

```js
({
  test: () => runPython(`
    import io
    import sys

    captured_output = io.StringIO()
    sys.stdout = captured_output

    _catalogue = MediaCatalogue()
    _catalogue.add(Movie('Barry Lyndon', 1975, 'Stanley Kubrick', 184))
    _catalogue.add(Movie('Fahrenheit 451', 1966, 'Francois Truffaut', 112))
    print(_catalogue)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    _expected = """Media Catalogue (2 items):

    1. Barry Lyndon (1975) - 184 min, Stanley Kubrick
    2. Fahrenheit 451 (1966) - 112 min, Francois Truffaut
    """
    assert _expected in output
  `)
})
```

当 `self.items` 为空时，你的 `__str__` 方法应该返回 `Media Catalogue (empty)`。

```js
({
  test: () => runPython(`
    import io
    import sys

    captured_output = io.StringIO()
    sys.stdout = captured_output

    print(MediaCatalogue())

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Media Catalogue (empty)" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)
--fcc-editable-region--
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        result = f'Media Catalogue ({len(self.items)} items):\n\n'

--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```