---
id: 66e2e99caa0cb3570f4822b4
title: 步骤 4
challengeType: 1
dashedName: step-4
---

# --description--
来测试一下你的辅音统计功能吧！

创建变量 `consonantCount`，赋值为调用 `getConsonantCount("Coding is fun")` 的结果。

然后在控制台打印：`Consonant Count: <变量>`，可以选择模板字符串或 `+` 拼接。看见正确输出，就说明阶段性成功啦～

# --hints--
需要一个变量 `consonantCount`。

```js
assert.isNotNull(consonantCount)
```

它的值要等于 `getConsonantCount("Coding is fun")`。

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

控制台输出格式需形如：`Consonant Count: ${consonantCount}`（或使用拼接）。

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
