---
id: 66e2edb23013ca596f8d9568
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--
测试一下你的标点统计功能：

创建变量 `punctuationCount`，赋值为调用 `getPunctuationCount("WHAT?!?!?!?!?")` 的结果。

然后在控制台打印：`Punctuation Count: <变量>`。模板字符串或拼接方式任选。看到输出后再开心一点点～

# --hints--
需要一个变量 `punctuationCount`。

```js
assert.isNotNull(punctuationCount)
```

它的值要等于 `getPunctuationCount("WHAT?!?!?!?!?")`。

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

控制台输出需包含：`Punctuation Count: ${punctuationCount}`（或使用 `+` 拼接）。

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

