---
id: 66e2d680e129e1423116a541
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--
在这个小可爱工作坊里，你要做一个“句子分析器”——输入一段句子，就能统计它有多少单词、元音、辅音以及标点符号。像给句子做一次体检，超有成就感！

第一步：请创建一个 `getVowelCount` 函数，带一个参数 `sentence`。它要返回该句子里所有元音字母的总数量。放轻松，这只是暖身～

# --hints--
你需要创建一个 `getVowelCount` 函数。（分析元音的核心助手！）

```js
assert.isFunction(getVowelCount);
```

`getVowelCount` 函数要带上一个参数 `sentence`。（用来接收用户输入的句子）

```js
assert.match(getVowelCount.toString(), /sentence/);
```

`getVowelCount` 必须返回一个数字（表示元音个数）。

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

当句子是 `"Apples are tasty fruits"` 时，结果应该是 `7`。

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

当句子是 `"Hello, World!"` 时，应返回 `3`。

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

统计应大小写不敏感。（A 和 a 都要乖乖算进去）

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("MasterPuti is awesome"), 10);
```

你的函数要适用于任意句子，包括空字符串。继续加油，你已经踏出第一步！

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```

