---
id: 66e2df8c5676fe51cead8ec8
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--
现在来测试一下你刚写好的 `getVowelCount` 吧！

创建一个变量 `vowelCount`，赋值为调用 `getVowelCount("Apples are tasty fruits")` 的结果。

然后在控制台打印：`"Vowel Count: <这里放变量>"`。你可以用模板字符串或者使用 `+` 拼接。让结果闪亮登场～

# --hints--
需要创建一个变量 `vowelCount`。

```js
assert.isNotNull(vowelCount)
```

`vowelCount` 的值要等于 `getVowelCount("Apples are tasty fruits")` 的结果。

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

要输出到控制台的格式形如：`Vowel Count: ${vowelCount}`。（或使用字符串拼接也可以）

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
