---
id: 66e2eab8a5638f57b637b7cc
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--
现在来统计标点符号的数量～

创建一个 `getPunctuationCount` 函数，带参数 `sentence`。

在函数中用循环统计“不是空格也不是字母”的字符数量，我们就把这些视为标点。

函数需要返回一个数字结果。继续扩展你的分析工具箱！

# --hints--
需要创建 `getPunctuationCount` 函数。

```js
assert.isFunction(getPunctuationCount);
```

它要接收一个 `sentence` 参数。

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

必须返回数字类型。

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

句子 `"What's going on here?"` 时结果应为 `2`。

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

句子 `"What????!"` 时结果应为 `5`。

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

函数要适用于任意句子，包括空字符串。冲刺继续～

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! MasterPuti awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? MasterPuti is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("MasterPuti, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```

