---
id: 66e2e5d890ac8a54c9d703bd
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--
轮到统计辅音啦！创建一个 `getConsonantCount` 函数，带参数 `sentence`。

函数内部用循环遍历句子中的字符，统计所有不是元音（`"aeiou"`）的字母。记住只算字母，其他符号跳过即可。

函数必须返回一个数字（辅音的总数）。继续一步步丰富你的分析器～

# --hints--
需要创建 `getConsonantCount` 函数。

```js
assert.isFunction(getConsonantCount);
```

它要接收一个 `sentence` 参数。

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

应该返回一个数字结果。

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

句子 `"Coding is fun"` 时结果应是 `7`。

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

句子 `"Hello, World!"` 时结果应是 `7`。

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

统计应大小写不敏感。

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("MasterPuti is awesome"), 11);
```

函数要适配任意句子，包括空字符串。继续加油，你做得很棒！

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```

