﻿---
id: 6823d6244511f252c8300eed
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

如果需要在网页上添加段落，你可以使用 `p` 元素，就像这样：

```html
<p>欢迎来到全栈课程学习</p>
```

请将 `学习成为全栈开发人员的技能` 转换为 `p` 元素。
# --hints--

你的 `p` 元素应该有一个开始标签。开始标签的语法是： `<p>`。
```js
assert.exists(document.querySelector("p"));
```

你的 `p` 元素应该有一个结束标签。结束标签的语法是： `</p>`。
```js
assert.match(code, /<\/p\s*\>/);
```

你的 `p` 元素的文本内容应该为 `学习成为全栈开发人员的技能`。
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer\s*\<\/p\s*\>/i);
```

你的 `p` 元素应该在 `h2` 元素后面。
```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈课程学习</h2>
--fcc-editable-region--
学习成为全栈开发人员的技能
--fcc-editable-region--
```