﻿---
id: 6823e637c1c0ed56f781b4fc
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

然后是学习JavaScript，它能够使网页具有交互性，例如当用户点击按钮、提交表单或执行其他操作时，网页能够响应。接下来需要你在页面中添加一个`h3`和`p`元素，用`h3`元素描述JavaScript。
请在编辑器的底部添加一个 `h3` 元素，文本内容为：
```md
JavaScript介绍
```

然后，在`h3`元素下方添加一个`p`元素，文本内容为：
```md
JavaScript 为网页提供交互性
```

# --hints--

你需要在页面中添加一个`h3`元素?
```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

你的`h3`元素的文本内容应该为`JavaScript介绍`?
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*JavaScript\s*介绍\s*\<\/h3\s*\>/i);
```

你需要添加第四个`p`元素?
```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

你的`p`元素的文本内容应该为`JavaScript 为网页提供交互性`?
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*为网页提供交互性\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈课程学习</h2>
<p>学习成为全栈开发人员的技?/p>

<h3>HTML介绍</h3>
<p>HTML表示网页的内容和结构</p>

<h3>CSS介绍</h3>
<p>CSS用于为网页添加样?/p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈课程学习</h2>
<p>学习成为全栈开发人员的技?/p>

<h3>HTML介绍</h3>
<p>HTML表示网页的内容和结构</p>

<h3>CSS介绍</h3>
<p>CSS用于为网页添加样?/p>

<h3>JavaScript介绍</h3>
<p>JavaScript 为网页提供交互?/p>
```