﻿---
id: 6823e169fda14755fbf00445
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

恭喜你，你已经学会了HTML。现在，在编辑器的底部添加一个 `h3` 元素，文本内容为：
```md
CSS介绍
```

# --hints--

你应该添加第二个 `h3` 元素到页面。
```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

你第二个 `h3` 元素的文本内容应该为 `CSS介绍`。你可以从指令中复制文本。
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈课程学习</h2>
<p>学习成为全栈开发人员的技能</p>

<h3>HTML介绍</h3>
<p>HTML表示网页的内容和结构</p>

--fcc-editable-region--

--fcc-editable-region--
```