---
id: 6823e036ea4b71553558c01b
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

你有看到吗？每个标题在预览页面的显示都不同，段落也是如此。请在 `h3` 元素后面添加一个 `p` 元素。
```md
HTML表示网页的内容和结构
```

# --hints--

然后你需要添加第二个 `p` 元素。
```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

你的第二个 `p` 元素的文本内容应该为 `HTML表示网页的内容和结构`。你可以从指令中复制文本。
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈课程学习</h2>
<p>学习成为全栈开发人员的技能</p>

<h3>HTML介绍</h3>
--fcc-editable-region--

--fcc-editable-region--
```