---
id: 682ba2318000b62f179bdf04
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML 由元素组成。你将学习第一个元素是 h1 元素。
```html
<h1>Welcome to MasterPuti</h1>
```

它由一个开始标签 (`<h1>`), 结束标签 (`</h1>`)组成，而在标签之间应该显示你需要的文本内容。
请在你的 `Welcome to MasterPuti` 文本前添加一个 `h1` 开始标签，并在后面添加一个结束标签，组成一个 `h1` 元素。
# --hints--

你的 `h1`元素的开始标签应该是 `<h1>`。
```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 应该含有一个结束标签 `</h1>`。
```js
assert.match(code, /<\/h1\s*\>/);
```

你的 `h1` 元素应该是这样的：`<h1>Welcome to MasterPuti</h1>`。
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*MasterPuti\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to MasterPuti
--fcc-editable-region--
```