---
id: 682cd20b883fc7b25eb539c6
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

请在这两行文本下面，添加?
```md
我们一起学习全栈课程，成为一个全栈开发人?```

# --hints--

在你的编辑器中应该包含文?`我们一起学习全栈课程，成为一个全栈开发人员` ，并检查文本内?
```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /我们一起学习全栈课程，成为一个全栈开发人?i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈开发学?/h2>
--fcc-editable-region--

--fcc-editable-region--
```
