---
id: 682cd206883fc7b25eb539c5
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

看到了吗，你在编辑器中输入的Html内容会在预览窗口中显示出来。在这个课程中，你将学习一部分网页的HTML代码?
请在你的 `h1` 元素下方,输入一下内?

```md
全栈课程学习
```

# --hints--

你应该在编辑器中包含文本 `全栈课程学习` ，并检查文本的内容?
```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /全栈课程学习/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
--fcc-editable-region--

--fcc-editable-region--
```
