﻿---
id: 683921c4769fd23dbadec2fe
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

HTML 定义了网页的结构和内容，而CSS 用于为网页添加样式，能够改变网页的外观和布局。
请在 `h3` 元素下面添加一个段落元素，文本为：

```md
CSS用于为网页添加样式
```

# --hints--

你需要添加第二个 `h3` 元素。
```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

你的 `h3` 元素的文本内容应该为 `CSS介绍`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*CSS介绍\s*\<\/h3\s*\>/i);
```

你应该在页面中添加第三个 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

你的 `p` 元素的文本内容应该为 `CSS用于为网页添加样式`。你可以从指令中复制文本。
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈课程学习</h2>
<p>学习成为全栈开发人员的技能</p>

<h3>HTML介绍</h3>
<p>HTML表示网页的内容和结构</p>

--fcc-editable-region--
<h3>CSS介绍</h3>

--fcc-editable-region--
```