---
id: 672b4434dd6d6a1edc405419
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 [.card-links](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素内添加两个 [a](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素。

第一个应该有文本 `Send Card`，[class](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 为 [send-link](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 和 [href](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 为 [#send](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

第二个应该有文本 `Share on Social Media`，[class](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 为 [share-link](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 和 [href](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 为 [#share](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

# --hints--

[.card-links](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 内的第一个 [a](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该有文本 `Send Card`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

第一个 [a](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该有一个 [class](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 值为 [send-link](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

第一个 [a](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该有一个 [href](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 值为 [#send](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

[.card-links](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 内的第二个 [a](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该有文本 `Share on Social Media`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

第二个 [a](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该有一个 [class](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 值为 [share-link](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

第二个 [a](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该有一个 [href](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 值为 [#share](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

--fcc-editable-region--

      <div class="card-links"></div>

--fcc-editable-region--

    </div>
  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
