---
id: 66b102ee0bdbad34a9f85ba0
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

欢迎来到小小教师机器人工作坊！我们将通过制作一个可爱的聊天机器人来练习字符串。先在控制台打印一句友好的问候 —— `Hi there!`，你做得到，加油！

# --hints--

小提示：在文件里添加一个 `console.log`，输出一句简单的问候吧。

```js
assert.match(code, /console\.log(.*)/);
```

你的 `console.log` 应该输出：`"Hi there!"`。

```js
assert.match(code, /console\.log\((['"])Hi\s+there!\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
