---
id: 66b59be6ab830800c4df9146
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

现在把刚才创建的变量组合成一句完整的话吧！使用模板字面量创建 `sentence`，例如：`Today, you will learn about ${topic} in ${subject}.`，然后把它打印到控制台。做得好，继续前进～

# --hints--

提示：

- 创建 `sentence` 并确保它是字符串：

```js
assert.isNotNull(sentence);
assert.isString(sentence);
```

- 使用模板字面量包含 `${topic}` 和 `${subject}`：

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

- 最后把 `sentence` 打印到控制台：

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--
const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--
