---
id: 66b5928b86e907fdfbf34e56
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

太好了！现在我们要为小机器人起个名字。创建一个变量 `botName`，赋值为字符串 `"teacherBot"`。这样你的机器人就有名字啦～

# --hints--

提示：创建 `botName` 变量并赋值为字符串 `"teacherBot"`。

```js
assert.isNotNull(botName);
```

确保它是字符串类型：

```js
assert.isString(botName);
```

并且它的值应为 `"teacherBot"`。

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
