---
id: 66b59b12a745e10011158f7b
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

太棒了！现在我们来创建两个变量，它们会出现在机器人之后的句子中：

- `subject`，值为 `"JavaScript"`
- `topic`，值为 `"strings"`

你已经在路上了，继续保持好奇心～

# --hints--

提示：

- `subject` 应存在且为字符串：

```js
assert.isNotNull(subject);
assert.isString(subject);
assert.equal(subject, "JavaScript");
```

- `topic` 也应存在且为字符串：

```js
assert.isNotNull(topic);
assert.isString(topic);
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--
const subject = "JavaScript";

const topic = "strings";

--fcc-editable-region--
