---
id: 66b6efddeca35833cd6f0b03
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

复习一下如何通过索引访问字符串字符：索引从 0 开始，所以第一个字符在索引 0。请添加一条 `console` 语句，用括号表示法输出 `subject` 的第一个字母。继续，你做得很棒！

# --hints--

提示：使用 `subject[0]` 并通过 `console.log` 输出。

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);
console.log(subject[0]);
