---
id: a0b5010f579e69b815e7c5d6
title: Build a Smart Word Replacement Function
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

在这个实验中，你将创建一个函数，使用提供的参数在句子中执行搜索和替换，然后返回新句子。

**注意：** 替换时要保持原单词首字母的大小写。例如，如果你想用单词 `dog` 替换单词 `Book`，它应该被替换为 `Dog`

**目标**：完成下面的用户故事并让所有测试通过以完成实验。

**用户故事：**

1. 你应该创建一个名为 `myReplace` 的函数。
1. `myReplace` 函数应接受三个参数：一个字符串、要被替换的单词和用来替换的单词。
1. `myReplace` 函数应返回替换后的新字符串。
1. 替换时你应该保持原单词首字母的大小写。

# --hints--

你应该创建一个名为 `myReplace` 的函数。

```js
assert.isFunction(myReplace);
```

`myReplace` 应该接受三个参数。

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` 应该返回字符串 `Let us go to the mall`。

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` 应该返回字符串 `He is Sitting on the couch`。

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` 应该返回字符串 `I think we should look down there`。

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` 应该返回字符串 `This has a spelling error`。

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` 应该返回字符串 `His name is John`。

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` 应该返回字符串 `Let us get back to more Algorithms`。

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
