---
id: 683ec39565dbc7e3375263b1
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在之前的学习中，你已经了解了二分查找以及它与线性查找的不同之处。

在这个练习中，你将通过构建一个二分查找算法来巩固这些知识。

首先，创建一个带有 [search_list] 和 [value] 参数的 [binary_search] 函数。

[search_list] 是函数将要在其中进行搜索的元素列表，[value] 是函数将尝试在 [search_list] 中查找的目标值。

现在在函数内部添加 [pass] 关键字，这样测试就可以通过啦。

# --hints--

你应该创建一个 [binary_search] 函数。

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

你的 [binary_search] 函数应该有 [search_list] 和 [value] 参数。

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```