---
id: 683ec95f2f8781355556ff7b
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在二分查找中，有最小可能索引和最大可能索引。

最小可能索引表示在当前搜索值范围内的最左侧边界，在这种情况下是 [search_list]。而最大可能索引则表示正在搜索的值范围内的最右侧边界。

为了考虑到这两点，你必须考虑整个被搜索的列表。所以在 [binary_search] 函数内部，定义一个值为 [0] 的 [low] 变量，以及一个值为被搜索列表最后一个索引的 [high] 变量。

# --hints--

你的 [binary_search] 函数中应该有一个名为 [low] 的变量。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

你的 [low] 变量应该被设置为 [0]。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

你的 [binary_search] 函数中应该有一个名为 [high] 的变量。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

你的 [high] 变量应该被设置为 [len(search_list) - 1]。

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []

--fcc-editable-region--
```