---
id: 683ec95f2f8781355556ff81
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

你需要测试一下，这样你就能理解算法在这个初始阶段的流程。

要做到这一点，你需要首先跳出循环。这是因为当前的实现只允许一次迭代，所以如果条件不满足，就会出现无限循环。

就在 [if] 语句之后，使用 [break] 关键字跳出 [while] 循环。然后，在 [while] 循环之后，返回一个空列表以表示未找到该值。

# --hints--

你应该使用 [break] 关键字在 [if] 语句之后跳出 [while] 循环。

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

你应该在 [while] 循环之后返回一个空列表。

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```