---
id: 683ec95f2f8781355556ff7a
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你正在构建的二分查找算法不仅仅会返回 [True] 或 [False] 来表示是否找到了目标值。它还会返回查找目标值所经过的路径。为此，你需要一个 [path_to_target] 变量来跟踪查找路径。

在你的 [binary_search] 函数内部，用 [path_to_target] 变量替换 [pass] 关键字，并将其初始化为空列表，这是一个很好的开始。

# --hints--

你应该创建一个 [path_to_target] 变量并将其设置为一个空列表在你的 [binary_search] 函数中。不要忘记删除 [pass] 关键字哦。

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```