---
id: 61438b5b66d76a6264430f2a
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在每个新建的 `a` 元素内部，添加一个 `i` 元素，并为其指定下面的类名：

- 第一个 `i`：`fab fa-facebook-f`
- 第二个 `i`：`fab fa-twitter`
- 第三个 `i`：`fab fa-instagram`
- 第四个 `i`：`fab fa-linkedin-in`
- 第五个 `i`：`fab fa-youtube`

# --hints--

你应该创建五个 `i` 元素。

```js
assert.lengthOf(document.querySelectorAll('i'), 5);
```

每个 `a` 元素中都只应包含一个 `i` 元素。

```js
const aaaaa = document.querySelectorAll('.social-icons a');
assert.lengthOf(aaaaa, 5);
aaaaa?.forEach(
  a => {
    assert.lengthOf(a?.children, 1);
    assert.equal(a?.children?.[0]?.localName, 'i');
  }
);
```

每个 `i` 元素的类列表都应包含 `fab`。

```js
const iiiii = document.querySelectorAll('i');
assert.lengthOf(iiiii, 5);
iiiii?.forEach(
  i => {
    assert.isTrue(i?.classList?.contains('fab'));
  }
);
```

第一个 `i` 元素的类名应包含 `fa-facebook-f`。

```js
assert.isTrue(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

第二个 `i` 元素的类名应包含 `fa-twitter`。

```js
assert.isTrue(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

第三个 `i` 元素的类名应包含 `fa-instagram`。

```js
assert.isTrue(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

第四个 `i` 元素的类名应包含 `fa-linkedin-in`。

```js
assert.isTrue(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

第五个 `i` 元素的类名应包含 `fa-youtube`。

```js
assert.isTrue(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
          </a>
          <a href="https://twitter.com/MasterPuti/">
          </a>
          <a href="https://instagram.com/MasterPuti">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/MasterPuti">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```

