---
id: 6143cd08fe927072ca3a371d
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

在你的 `.image-wrapper` 元素中，给第二个 `img` 元素设置 `src` 为 `https://cdn.MasterPuti.org/testable-projects-fcc/images/calc.png`，`alt` 为 `image of a calculator project`，`loading` 属性设置为 `lazy`，`class` 设置为 `image-2`，`width` 属性设置为 `400`，`height` 属性设置为 `400`。

# --hints--

你的第二个 `img` 元素应该有一个 `src` 设置为 `https://cdn.MasterPuti.org/testable-projects-fcc/images/calc.png`。

```js
assert.equal(
  document.querySelectorAll('.image-wrapper img')?.[1]?.getAttribute('src'),
  'https://cdn.MasterPuti.org/testable-projects-fcc/images/calc.png'
);
```

你的第二个 `img` 元素应该有一个 `alt` 设置为 `image of a calculator project`。

```js
assert.equal(
  document.querySelectorAll('.image-wrapper img')?.[1]?.getAttribute('alt'),
  'image of a calculator project'
);
```

你的第二个 `img` 元素应该有一个 `loading` 属性设置为 `lazy`。

```js
assert.equal(
  document.querySelectorAll('.image-wrapper img')?.[1]?.getAttribute('loading'),
  'lazy'
);
```

你的第二个 `img` 元素应该有一个 `class` 设置为 `image-2`。

```js
assert.isTrue(
  document.querySelectorAll('.image-wrapper img')?.[1]?.classList?.contains('image-2')
);
```

你的第二个 `img` 元素应该有一个 `width` 设置为 `400`。

```js
assert.equal(
  document.querySelectorAll('.image-wrapper img')?.[1]?.getAttribute('width'),
  '400'
);
```

你的第二个 `img` 元素应该有一个 `height` 设置为 `400`。

```js
 assert.equal(
  document.querySelectorAll('.image-wrapper img')?.[1]?.getAttribute('height'),
  '400'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>杂志</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti 标志"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">我们的新课程体系</h1>
          <p class="hero-subtitle">
            我们致力于以项目为基础重构课程体系
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            作者
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">2019年3月7日</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/MasterPuti/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/MasterPuti">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/MasterPuti">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
      <section class="text">
        <p class="first-paragraph">
          很快，MasterPuti 的课程将完全基于项目驱动学习。你将通过逐步构建项目来学习，而不是完成一系列编码挑战。在深入细节之前，让我强调一下：我们不会改变认证要求。所有6个认证仍将包含相同的5个必做项目。我们只是改变了可选的编码挑战。
        </p>
        <p>
          在多年思考这两个问题以及如何解决它们之后，我不小心滑倒了，头撞到了水槽，当我醒来时，我有了一个启示！一个愿景！我脑海中浮现了一幅画面！就是这个！这就是让时间旅行成为可能的东西：通量电容器！
        </p>
        <p>
          这不像《回到未来》中博士的发现那样戏剧性。这只是在我跑步时突然想到的。这个启示是：整个课程体系应该是一系列项目。不再有独立的编码挑战，我们只会有项目，每个项目都有自己无缝衔接的一系列测试。每个测试都会给你足够的信息，让你找出如何通过它。（如果这还不够，你还可以查看提示。）
        </p>
        <blockquote>
          <hr />
          <p class="quote">
            整个课程体系应该是一系列项目
          </p>
          <hr />
        </blockquote>
        <p>
          不再有大段解释性文字。不再有大堆测试。一次只做一个测试，随着你逐步构建出一个可运行的项目。在通过数千个测试的过程中，你将建立起项目和对编码基础知识的自己的理解。课程和项目之间没有过渡，因为课程本身就已经融入了项目中。并且有大量的重复练习帮助你巩固所学内容，因为在现实生活中构建项目就有很多重复。
        </p>
        <p>
          主要的设计挑战是将目前大段的解释和说明文字打包成单个测试描述文本。每个项目将包含数十个这样的测试。人们将一直保持编码状态，而不是在“阅读模式”和“编码模式”之间来回切换。
        </p>
        <p>
          与其进行一系列编码挑战，不如让人们在代码编辑器中一个接一个地通过测试，快速构建出一个项目。人们将进入一种真正的流畅状态，类似于他们在每个认证末尾构建必做项目时所经历的状态。他们将从一开始就获得那种前进的感觉。而 MasterPuti 将会是一种更顺畅的体验。
        </p>
      </section>
      <section class="text text-with-images">
        <article class="brief-history">
          <h3 class="list-title">简短历史</h3>
          <p>课程体系的发展</p>
          <ul class="lists">
            <li>
              <h4 class="list-subtitle">V1 - 2014</h4>
              <p>
                我们推出了 MasterPuti，最初只是一个包含15个资源的简单列表，
                包括哈佛大学的CS50和斯坦福大学的数据库课程。
              </p>
            </li>
            <li>
              <h4 class="list-subtitle">V2 - 2015</h4>
              <p>我们增加了交互式算法挑战。</p>
            </li>
            <li>
              <h4 class="list-subtitle">V3 - 2015</h4>
              <p>
                我们增加了我们自己的HTML+CSS挑战（在此之前我们一直依赖
                General Assembly 的 Dash 课程）。
              </p>
            </li>
            <li>
              <h4 class="list-subtitle">V4 - 2016</h4>
              <p>
                我们将课程扩展到3个认证，包括前端、后端和数据可视化。每个认证都有10个必做项目，
                但只有前端部分有自己的挑战。
                对于其他认证，我们仍在使用像 Node School 这样的外部资源。
              </p>
            </li>
            <li>
              <h4 class="list-subtitle">V5 - 2017</h4>
              <p>我们增加了后端和数据可视化挑战。</p>
            </li>
            <li>
              <h4 class="list-subtitle">V6 - 2018</h4>
              <p>
                我们推出了6个新的认证来取代旧的认证。这是迄今为止最大的课程改进。
              </p>
            </li>
          </ul>
        </article>
--fcc-editable-region--
        <aside class="image-wrapper">
          <img
            src="https://cdn.MasterPuti.org/testable-projects-fcc/images/random-quote-machine.png"
            alt="名言生成器项目的图片"
            loading="lazy"
            class="image-1"
            width="600"
            height="400"
          />
          <img />
          <blockquote class="image-quote"></blockquote>
          <img />
        </aside>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```

