---
id: 61439dfc811e12666b04be6f
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

在 `.text` 元素末尾添加第三个 `p` 元素，并给予以下文本：

```
这并不像《回到未来》中博士的启示那样戏剧性。我只是在跑步时突然想到的。启示是：整个课程体系应该是一系列项目。我们将不再有单独的编程挑战，而是直接做项目，每个项目都有一系列无缝衔接的测试。每个测试都会给你足够的信息来弄清楚如何通过测试。（如果这些信息还不够，你可以查看提示。）
```

# --hints--

你应该在 `.text` 元素中创建第三个 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('.text p'), 3);
```

你的第三个 `p` 元素应包含所提供的文本。

```js
assert.equal(document.querySelectorAll('.text p')?.[2]?.innerText, "这并不像《回到未来》中博士的启示那样戏剧性。我只是在跑步时突然想到的。启示是：整个课程体系应该是一系列项目。我们将不再有单独的编程挑战，而是直接做项目，每个项目都有一系列无缝衔接的测试。每个测试都会给你足够的信息来弄清楚如何通过测试。（如果这些信息还不够，你可以查看提示。）");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">我们的新课程体系</h1>
          <p class="hero-subtitle">
            我们努力重构课程体系，使其更加注重项目实践
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            作者
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">2019年3月7日</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/MasterPuti/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/MasterPuti">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/MasterPuti">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
      <section class="text">
        <p class="first-paragraph">
          很快，MasterPuti 的课程体系将完全采用项目驱动的学习方式。你将通过逐步构建项目来学习，而不是完成一系列编程挑战。在进入细节之前，让我强调一点：我们不会改变认证体系。所有 6 个认证仍将包含相同的 5 个必修项目。我们只是改变了可选的编程挑战。
        </p>
        <p>
          经过多年——多年的思考这两个问题以及如何解决它们，我滑倒了，头撞到水槽上，当我醒来时，我有了一个启示！一个愿景！脑海中浮现出一幅画面！就是这个！这就是时间旅行成为可能的原因：通量电容器！
        </p>
        <p>
          这并不像《回到未来》中博士的启示那样戏剧性。我只是在跑步时突然想到的。启示是：整个课程体系应该是一系列项目。我们将不再有单独的编程挑战，而是直接做项目，每个项目都有一系列无缝衔接的测试。每个测试都会给你足够的信息来弄清楚如何通过测试。（如果这些信息还不够，你可以查看提示。）
        </p>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css
