---
id: 6148e2dcdd60306dd77d41cc
title: 步骤 71
challengeType: 0
dashedName: step-71
---

# --description--

是时候为杂志的最后一部分——图片添加样式了。

图片被包裹在使用 `image-wrapper` 类的 `aside` 元素中，所以创建一个 `.image-wrapper` 选择器。将 `display` 属性设置为 `grid`。

# --hints--

你应该有一个 `.image-wrapper` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.image-wrapper'));
```

你的 `.image-wrapper` 选择器应该有一个 `display` 属性设置为 `grid`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.image-wrapper')?.display, 'grid');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>杂志</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti 标志"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">我们的新课程</h1>
          <p class="hero-subtitle">
            我们致力于以项目为中心重构课程体系
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            作者：
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">2019年3月7日</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/MasterPuti/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/MasterPuti">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/MasterPuti">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
      <section class="text">
        <p class="first-paragraph">
          很快，MasterPuti 的课程将完全基于项目驱动学习。你将通过逐步构建项目来学习，而不是完成一系列编程挑战。在深入细节之前，让我强调一下：我们不会改变认证要求。所有6个认证仍将保持相同的5个必做项目。我们只改变可选的编程挑战。
        </p>
        <p>
          在多年思考这两个问题以及如何解决它们之后，我不小心滑倒，头撞到了水槽，当我醒来时，我有了一个启示！一个愿景！我脑海中浮现了一幅画面！就是这个！这就是让时间旅行成为可能的东西：通量电容器！
        </p>
        <p>
          这不像《回到未来》中博士的启示那样戏剧性。这只是在我跑步时突然想到的。启示是：整个课程应该是一系列项目。不再是单独的编程挑战，我们将只有项目，每个项目都有其无缝衔接的一系列测试。每个测试都会给你足够的信息，让你找出如何通过它。（如果这还不够，你可以查看提示。）
        </p>
        <blockquote>
          <hr />
          <p class="quote">
            整个课程应该是一系列项目
          </p>
          <hr />
        </blockquote>
        <p>
          不再有大段的解释性文字，不再有成堆的测试。一次只做一个测试，在构建工作项目的过程中逐步推进。在通过数千个测试的过程中，你将建立起项目并加深对编码基础的理解。课程和项目之间没有过渡，因为课程本身已经融入了项目中。而且有大量的重复练习帮助你巩固所学，毕竟——现实生活中构建项目本身就充满了重复。
        </p>
        <p>
          主要的设计挑战是将目前大段的解释和说明整合到单个测试描述文本中。每个项目都将包含数十个这样的测试。人们将一直专注于编码，而不是在“阅读模式”和“编码模式”之间来回切换。
        </p>
        <p>
          与其进行一系列编程挑战，不如让人们在代码编辑器中一个接一个地通过测试，快速构建项目。人们将进入真正的专注状态，就像他们在每个认证末尾构建必做项目时所体验到的那样。他们从一开始就能感受到前进的动力。而 MasterPuti 将提供更流畅的学习体验。
        </p>
      </section>
      <section class="text text-with-images">
        <article class="brief-history">
          <h3 class="list-title">课程简史</h3>
          <p>关于课程的发展</p>
          <ul class="lists">
            <li>
              <h4 class="list-subtitle">V1 - 2014</h4>
              <p>
                我们推出了 MasterPuti，最初只是一个包含15个资源的简单列表，
                包括哈佛大学的CS50和斯坦福大学的数据库课程。
              </p>
            </li>
            <li>
              <h4 class="list-subtitle">V2 - 2015</h4>
              <p>我们增加了交互式算法挑战。</p>
            </li>
            <li>
              <h4 class="list-subtitle">V3 - 2015</h4>
              <p>
                我们增加了自己的HTML+CSS挑战（在此之前我们一直依赖于
                General Assembly的Dash课程）。
              </p>
            </li>
            <li>
              <h4 class="list-subtitle">V4 - 2016</h4>
              <p>
                我们将课程扩展到3个认证，包括前端、后端和数据可视化。每个认证都有10个必做项目，
                但只有前端部分有自己的挑战。
                对于其他认证，我们仍在使用Node School等外部资源。
              </p>
            </li>
            <li>
              <h4 class="list-subtitle">V5 - 2017</h4>
              <p>我们增加了后端和数据可视化挑战。</p>
            </li>
            <li>
              <h4 class="list-subtitle">V6 - 2018</h4>
              <p>
                我们推出了6个新的认证来取代旧的认证。这是迄今为止最大的课程改进。
              </p>
            </li>
          </ul>
        </article>
        <aside class="image-wrapper">
          <img
            src="https://cdn.MasterPuti.org/testable-projects-fcc/images/random-quote-machine.png"
            alt="引语生成器项目图片"
            loading="lazy"
            class="image-1"
            width="600"
            height="400"
          />
          <img
            src="https://cdn.MasterPuti.org/testable-projects-fcc/images/calc.png"
            alt="计算器项目图片"
            loading="lazy"
            class="image-2"
            width="400"
            height="400"
          />
          <blockquote class="image-quote">
            <hr />
            <p class="quote">
              数百万通过 MasterPuti 学习编程的人将拥有更好的资源来帮助他们学习这些基础知识。
            </p>
            <hr />
          </blockquote>
          <img
            src="https://cdn.MasterPuti.org/testable-projects-fcc/images/survey-form-background.jpeg"
            alt="四个人正在编写代码"
            loading="lazy"
            class="image-3"
            width="600"
            height="400"
          />
        </aside>
      </section>
    </main>
  </body>
</html>
```

```css
*,
::before,
::after {
  padding: 0;
  margin: 0;
  box-sizing: border-box;
}

html {
  font-size: 62.5%;
}

body {
  font-family: 'Baskervville', serif;
  color: linen;
  background-color: rgb(20, 30, 40);
}

h1 {
  font-family: 'Anton', sans-serif;
}

h2, h3, h4, h5, h6 {
  font-family: 'Raleway', sans-serif;
}

a {
  text-decoration: none;
  color: linen;
}

main {
  display: grid;
  grid-template-columns: minmax(2rem, 1fr) minmax(min-content, 94rem) minmax(2rem, 1fr);
  row-gap: 3rem;
}

img {
  width: 100%;
  object-fit: cover;
}

hr {
  margin: 1.5rem 0;
  border: 1px solid rgba(120, 120, 120, 0.6);
}

.heading {
  grid-column: 2 / 3;
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  row-gap: 1.5rem;
}

.text {
  grid-column: 2 / 3;
  font-size: 1.8rem;
  letter-spacing: 0.6px;
  column-width: 25rem;
  text-align: justify;
}

.hero {
  grid-column: 1 / -1;
  position: relative;
}

.hero-title {
  text-align: center;
  color: orangered;
  font-size: 8rem;
}

.hero-subtitle {
  font-size: 2.4rem;
  color: orangered;
  text-align: center;
}

.author {
  font-size: 2rem;
  font-family: "Raleway", sans-serif;
}

.author-name a:hover {
  background-color: #306203;
}

.publish-date {
  color: rgba(255, 255, 255, 0.5);
}

.social-icons {
  display: grid;
  font-size: 3rem;
  grid-template-columns: repeat(5, 1fr);
  grid-auto-flow: column;
  grid-auto-columns: 1fr;
  align-items: center;
}

.first-paragraph::first-letter {
  font-size: 6rem;
  color: orangered;
  float: left;
  margin-right: 1rem;
}

.quote {
  color: #00beef;
  font-size: 2.4rem;
  text-align: center;
  font-family: "Raleway", sans-serif;
}

.quote::before {
  content: '" ';
}

.quote::after {
  content: ' "';
}

.text-with-images {
  display: grid;
  grid-template-columns: 1fr 2fr;
  column-gap: 3rem;
  margin-bottom: 3rem;
}

.lists {
  list-style-type: none;
  margin-top: 2rem;
}

.lists li {
  margin-bottom: 1.5rem;
}

.list-title, .list-subtitle {
  color: #00beef;
}

--fcc-editable-region--

--fcc-editable-region--
```

