---
id: 6148d1f9eb63c252e1f8acc4
title: 步骤 60
challengeType: 0
dashedName: step-60
---

# --description--

要给 `hr` 添加颜色，你需要调整 `border` 属性。给 `hr` 选择器添加一个 `border` 属性，设置为 `1px solid rgba(120, 120, 120, 0.6)`。

# --hints--

你的 `hr` 应该有一个 `border` 属性设置为 `1px solid rgba(120, 120, 120, 0.6)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('hr')?.borderWidth, '1px');
assert.equal(new __helpers.CSSHelp(document).getStyle('hr')?.borderStyle, 'solid');
assert.equal(new __helpers.CSSHelp(document).getStyle('hr')?.borderColor, 'rgba(120, 120, 120, 0.6)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>杂志</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.MasterPuti.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="MasterPuti 标志"
            loading="lazy"
            class="hero-img"
          />
          <h1 class="hero-title">我们的新课程体系</h1>
          <p class="hero-subtitle">
            我们致力于以项目为中心重构课程体系
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            作者：
            <a href="https://MasterPuti.org" target="_blank" rel="noreferrer"
              >MasterPuti</a
            >
          </p>
          <p class="publish-date">2019年3月7日</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/MasterPuti/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/MasterPuti/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/MasterPuti">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/MasterPuti">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
      <section class="text">
        <p class="first-paragraph">
          很快，MasterPuti 的课程体系将完全转变为以项目驱动的学习方式。你将通过逐步构建项目来学习，而不是完成一系列编码挑战。在深入了解细节之前，让我强调一下：我们不会改变认证要求。所有6个认证仍将包含相同的5个必做项目。我们只改变可选的编码挑战。
        </p>
        <p>
          经过多年——多年的思考这两个问题以及如何解决它们，我滑倒了，头撞到了水槽上，当我醒来时，我有了一个启示！一个愿景！我脑海中浮现了一幅画面！就是这个！这就是让时间旅行成为可能的东西：通量电容器！
        </p>
        <p>
          这不像《回到未来》中博士的启示那么戏剧性。这只是在我跑步时突然想到的。这个启示是：整个课程体系应该是一系列项目。不再是个别的编码挑战，我们将只有项目，每个项目都有其无缝衔接的一系列测试。每个测试都会给你足够的信息来弄清楚如何通过它。（如果这些信息不够，你也可以查看提示。）
        </p>
        <blockquote>
          <hr />
          <p class="quote">
            整个课程体系应该是一系列项目
          </p>