---
id: 669fde2081f65141ad703fe4
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

页面的第一个部分将是关于部分。该部分将介绍 Whiskers 先生，让用户体验了解博客的内涵。

在你的 `main` 元素中，添加一个 `section` 元素，其 `id` 属性设置为 `"about"`。

在 `section` 元素中，添加一个 `h2`，其文本为 `About`。

# --hints--

你的 `section` 元素应该有一个开始标签。

```js
assert.match(code, /<section.*>/);
```

你的 `section` 元素应该有一个结束标签。

```js
assert.match(code, /<\/section\>/);
```

你的 `section` 元素应该有一个 `id`，设置为 `"about"`。

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

你的 `section` 元素应该在 `main` 元素中。

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

你的 `h2` 元素应该有一个开始标签。

```js
assert.match(code, /<h2\>/);
```

你的 `h2` 元素应该有一个结束标签。

```js
assert.match(code, /<\/h2\>/);
```

你的 `h2` 元素应该在 `section` 元素中。

```js
const h2Element = document.querySelector('h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

你的 `h2` 元素应该有文本 `About`。

```js
const h2Element = document.querySelector('h2');
assert.strictEqual(h2Element?.textContent, "关于");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Mr. Whiskers' Blog</title>
  </head>
  <body>
    <header>
      <h1>Mr. Whiskers' Blog</h1>
      <figure>
        <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg" alt="a cat in the garden" />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      
    </main>
  </body>
</html>
```
