---
id: 669fcb5106eeb13ab7afcd09
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

你要创建的第一个部分是<dfn>页头</dfn>。

`header` 元素用于表示介绍性内容，如页面导航和其他介绍性信息。

下面是一个使用 `header` 元素的示例：

```html
<header>
  <h1>主页面标题放在这里</h1>
  <img src="example-logo.png" alt="示例标志" />
</header>
```

在 `body` 元素内，添加一个 `header` 元素。

# --hints--

你应该有一个 `<header>` 开始标签。

```js
assert.match(code, /<header>/i);
```

你应该有一个 `</header>` 结束标签。

```js
assert.match(code, /<\/header>/i);
```

你的 `<header>` 开始标签应该在 `<body>` 开始标签之后。

```js
assert.match(code, /<body>\s*<header>/i);
```

你的 `</header>` 结束标签应该在 `</body>` 结束标签之前。

```js
assert.match(code, /<\/header>\s*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Mr. Whiskers' Blog</title>
  </head>
  <body>
    
  </body>
</html>
```
