---
id: 669fdd3965f36f40af9615e5
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

现在你已经完成了页头的制作，需要开始添加主体内容。

在你的 `header` 下方，添加一个 `main` 元素。

# --hints--

你的 `main` 元素应该有一个开始标签。

```js
assert.match(code, /<main\>/);
```

你的 `main` 元素应该有一个结束标签。

```js
assert.match(code, /<\/main\>/);
```

你的 `main` 元素应在 `header` 元素下方。

```js
const headerElement = document.querySelector('header');
const mainElement = document.querySelector('main');
assert.equal(headerElement?.nextElementSibling, mainElement);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Mr. Whiskers' Blog</title>
  </head>
  <body>
    <header>
      <h1>Mr. Whiskers' Blog</h1>
      <figure>
        <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg" alt="a cat in the garden" />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    
  </body>
</html>
```
