---
id: 669fdfc9a5e3da42d2376609
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

现在你已添加了关于部分，请尝试点击 `About` 链接，查看页面是否能跳转到该部分。

博客页面的下一部分将是关于 Whiskers 先生的博文列表。

添加另一个 `section` 元素，其 `id` 设置为 `"posts"`。

在 `section` 元素中，添加一个 `h2` 元素，其文本为 `Posts`。

# --hints--

你的代码应有一个 `section` 开始标签。

```js
assert.match(code, /<section.*>/i);
```

你的代码应有一个 `section` 结束标签。

```js
assert.match(code, /<\/section>/i);
```

你的 `section` 元素应有一个 `id`，设置为 `"posts"`。

```js
assert.match(code, /<section.*id="posts".*>/i);
```

你的 `section` 元素应在 `main` 元素中。

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

你的 `h2` 元素应在 `section` 元素中。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

你的 `h2` 元素应有文本 `Posts`。

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.textContent, "文章");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Mr. Whiskers' Blog</title>
  </head>
  <body>
    <header>
      <h1>Mr. Whiskers' Blog</h1>
      <figure>
        <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg" alt="a cat in the garden" />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
      </section>
      
    </main>
  </body>
</html>
```
