---
id: 669fc938d38e6e38ace9251e
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `head` 元素内，嵌套一个 `meta` 元素，其 `charset` 属性的值设置为 `"UTF-8"`。

在该 `meta` 元素下方，添加一个 `title` 元素。

`title` 元素的文本应为 `Mr. Whiskers' Blog`。

# --hints--

你应该有一个 `meta` 元素。

```js
assert.isNotNull(document.querySelector("meta"));
```

`meta` 元素是一个空元素，它不应该有一个结束标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 标签应该有一个 `charset` 属性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 属性的值应为 `"UTF-8"`。

```js
assert.match(code, /<meta\s+charset\s*=\s*["']UTF-8["']/i);
```

你应该有一个 `title` 元素。

```js
assert.isNotNull(document.querySelector("title"));
```

你的 `title` 元素应该有文本 `Mr. Whiskers' Blog`。

```js
assert.equal(document.querySelector("title").textContent, "Mr. Whiskers' Blog");
```

你的 `meta` 元素应该在你的 `head` 元素内。

```js
assert.exists(document.querySelector("head meta"));
```

你的 `title` 元素应该在你的 `head` 元素内。

```js
assert.exists(document.querySelector("head title"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
</html>
```
