---
id: 66a335e5888ffb367633200a
title: 步骤 15
challengeType: 0
dashedName: step-15
---

# --description--

博文将包含几段 <dfn>lorem ipsum</dfn> 文本。

Lorem ipsum 通常在网络开发中用作占位文本。当你想集中精力构建网页的基本结构，而不必担心实际内容时，它就非常有用。

下面是一个使用 lorem ipsum 的例子：

```html
<p>
  Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
  quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
  quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
  quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
</p>
```

在 `h3` 元素下方，添加两段 lorem ipsum 文本。

# --hints--

你应该在 `article` 元素中添加两个 `p` 元素。

```js
const paragraphElements = document.querySelectorAll('#posts article p');
assert.strictEqual(paragraphElements.length, 2);
```

第一个 `p` 元素应使用 lorem ipsum 文字。请仔细检查拼写和标点符号是否有误。

```js
const paragraphElement = document.querySelector('article p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

第二个 `p` 元素应使用 lorem ipsum 文字。请仔细检查拼写和标点符号是否有误。

```js
const paragraphElement = document.querySelector('article p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

你的段落应位于 `h3` 元素之后。

```js
let h3 = document.querySelectorAll('article > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('article > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>欢迎来到 Whiskers 先生的博客页面！</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="花园里的一只猫"
        />
        <figcaption>花园里的 Whiskers 先生</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">关于</a></li>
          <li><a href="#posts">文章</a></li>
          <li><a href="#contact">联系</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>关于</h2>
      </section>
      <section id="posts">
        <h2>文章</h2>
        <article>
          <h3>Whiskers 先生回家的第一天</h3>
          
        </article>
      </section>
    </main>
  </body>
</html>
```
