---
id: 6894c1b715b8d18a18ef3203
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次练习中，你将通过构建一个乐器库存来练习 Python 的类和对象概念，轻轻松松上手面向对象的基础。

在 Python 中定义类时，请使用 `class` 关键字，后面跟上类名和冒号：

```python
class ClassName:
    # class contents go here
```

提示：Python 的类名通常使用 PascalCase（每个单词首字母大写且不加空格）。

请先创建一个名为 `MusicalInstrument` 的类，暂时在类体内使用 `pass` 作为占位。

# --hints--

你应该创建一个名为 `MusicalInstrument` 的类。

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
