---
id: 6895e0bdb0348bf46952ab67
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

现在是创建 `MusicalInstrument` 类实例的时候了。

要创建类的实例，可以像调用函数一样使用类名，并传入 `__init__` 方法所需的参数（`self` 会由 Python 自动处理）：

```python
my_instance = ClassName(arg1, arg2)
```

当你创建实例时，Python 会调用 `__init__` 方法并把传入的值存为该实例独有的属性。

在类定义之外，创建一个名为 `instrument_1` 的实例，表示一种名为 `Oboe` 的 `woodwind`（木管）乐器。

# --hints--

你应该创建一个名为 `instrument_1` 的 `MusicalInstrument` 实例。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

创建 `instrument_1` 时，请传入 `Oboe` 和 `woodwind` 作为参数。

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
