---
id: 6895e0452b3370f26f365a46
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

在 `__init__` 方法中，你需要把传入的参数赋值给实例属性（instance attributes）。实例属性是属于某个特定类实例的变量。

在你的 `__init__` 方法里，把 `name` 参数赋值给 `self.name`，并把 `instrument_type` 参数赋值给 `self.instrument_type`。

# --hints--

你应该在 `__init__` 方法中把 `name` 参数赋给 `self.name`。

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

你还应该把 `instrument_type` 参数赋给 `self.instrument_type`。

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
