---
id: 6895e356fd877efd6e3657ee
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--


在 `MusicalInstrument` 类中，定义另一个名为 `play` 的方法，该方法只接收 `self` 参数。`self` 使方法可以访问该实例的属性。

# --hints--

你应该定义一个名为 `play`、接收 `self` 参数的方法。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
