---
id: 6895e4940e39950288e406c2
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

接下来，在类中添加一个名为 `get_fact` 的方法，用来返回关于该乐器的一条信息。

此方法仅接收 `self` 参数，并应返回一个 f-string，格式为：`The [instrument name] is part of the [instrument type] family of instruments.`。

与直接打印信息的 `play` 方法不同，`get_fact` 应返回字符串以便在其他地方使用。

# --hints--

你应该定义一个名为 `get_fact`、接收 `self` 参数的方法。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

你的 `get_fact` 方法应返回包含乐器名称与类型的 f-string。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
