---
id: 6895e235935106f88281b648
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

你可以使用点号（dot notation）来访问实例的属性：

```py
instance_name.attribute_name
```

这样可以检索在初始化时存储的属性值。

试着打印出两个实例的 `name` 和 `instrument_type` 属性，你就能看到当初赋予它们的值啦。

# --hints--

你应该打印 `instrument_1` 的 `name` 属性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

你应该打印 `instrument_1` 的 `instrument_type` 属性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

你应该打印 `instrument_2` 的 `name` 属性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

你应该打印 `instrument_2` 的 `instrument_type` 属性。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
