---
id: 6895dffc64e97df0e62e9ef9
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

方法（methods）是定义在类内部的函数，可以使用类的数据来执行操作。下面是方法的定义方式：

```python
def method_name(parameter1, parameter2,...):
  # Method contents will go here
```

在 Python 中，`__init__` 是一个特殊方法，称为初始化方法（initializer），当创建类的新实例时会被调用。

在你的 `MusicalInstrument` 类中，添加一个带有三个参数的 `__init__` 方法：`self`、`name` 和 `instrument_type`。

`self` 参数是引用类实例的惯用方式，它必须作为所有实例方法的第一个参数。

# --hints--

你应该在 `MusicalInstrument` 类中定义 `__init__` 方法。

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

你的 `__init__` 方法应包含三个参数：`self`、`name` 和 `instrument_type`。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
