---
id: 6895e197180c9af6988e1e19
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

创建第二个实例 `instrument_2`，表示一种名为 `Trumpet` 的 `brass`（铜管）乐器。

你会看到，同一个类模板可以创建多个实例，每个实例都有自己独立的属性值，互不干扰。

# --hints--

你应该创建一个名为 `instrument_2` 的 `MusicalInstrument` 实例。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

创建 `instrument_2` 时，请传入 `Trumpet` 和 `brass` 作为参数。

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
