---
id: 6895e3f91c069e0068a42582
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

在 `play` 方法中，使用 `print()` 输出一条关于演奏该乐器的提示信息。

请使用 f-string 在信息中包含乐器名称，格式为：`The [instrument name] is fun to play!`，其中 `[instrument name]` 请替换为该乐器的真实名字。

# --hints--

你的 `play` 方法应使用 f-string 输出 `The [instrument name] is fun to play!`，并通过 `self.name` 访问乐器名称。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
