---
id: 66be24cb4144f955b6bcc550
title: Build a Page of Playing Cards
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该构建一个显示至少三张扑克牌的网页。
2. 您应该有一个ID为[playing-cards](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L307-L313)的[main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L197-L287)元素。
3. 在您的#playing-cards元素内，您应该有至少三个[div](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L198-L210)元素，每个元素都有[card](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L315-L330)类。
4. 在每个.card元素内，您应该有三个[div](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L198-L210)元素，第一个带有[left](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L332-L334)类，第二个带有[middle](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L336-L341)类，第三个带有[right](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L343-L345)类。
5. 您的#playing-cards元素应该使用flexbox水平居中其子元素，允许它们换行，并在它们之间放置[20px](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L311-L311)的空间。
6. 每个.card元素应该使用flexbox通过[space-between](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L328-L328)来调整其子元素。
7. 每个.left元素应该使用flexbox项目属性将其自身对齐到父容器的起始位置。
8. 每个.middle元素应该使用flexbox项目属性将其自身对齐到父容器的中心。
9. 每个.right元素应该使用flexbox项目属性将其自身对齐到父容器的末尾。
10. 每个.middle元素应该使用flexbox以列的形式显示其子元素。

这里有一些字符您可以复制粘贴到卡片中使用（如果需要）：♠、♣、♥和♦。

# --hints--

您应该有一个ID为[playing-cards](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L307-L313)的[main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L197-L287)元素。

```js
assert.exists(document.querySelector('main#playing-cards'));
```

您应该在#playing-cards元素内有至少三个带有[card](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L315-L330)类的[div](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L198-L210)元素。

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

每个.card元素都应该有宽度和高度。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

每个.card元素应该恰好有三个[div](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L198-L210)元素作为子元素。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

您应该在每个.card元素内有一个带有[left](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L332-L334)类的[div](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L198-L210)元素。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

您应该在每个.card元素内有一个带有[middle](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L336-L341)类的[div](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L198-L210)元素。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

您应该在每个.card元素内有一个带有[right](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L343-L345)类的[div](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L198-L210)元素。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

您应该有一个#playing-cards选择器，将其元素的[display](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L309-L309)设置为[flex](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L309-L309)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

您应该有一个#playing-cards选择器，将其元素的[justify-content](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L310-L310)设置为[center](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L310-L310)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

您应该有一个#playing-cards选择器，将其元素的[flex-wrap](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L312-L312)设置为[wrap](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L312-L312)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

您应该有一个#playing-cards选择器，将其元素的[gap](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L311-L311)设置为[20px](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L311-L311)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

您应该有一个.card选择器，将其元素的[display](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L327-L327)设置为[flex](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L327-L327)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

您应该有一个.card选择器，将其元素的[justify-content](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L328-L328)设置为[space-between](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L328-L328)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

您应该有一个.left选择器，将其元素的[align-self](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L333-L333)设置为[flex-start](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L333-L333)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

您应该有一个.middle选择器，将其元素的[align-self](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L337-L337)设置为[center](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L310-L310)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

您应该有一个.right选择器，将其元素的[align-self](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L344-L344)设置为[flex-end](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L344-L344)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

您应该有一个.middle选择器，将其元素的[display](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L338-L338)设置为[flex](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L338-L338)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

您应该在.middle类中将[flex-direction](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L340-L340)设置为[column](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-page-of-playing-cards/66be24cb4144f955b6bcc550.md#L340-L340)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}