---
id: 67051431a73f1ca25d9a6f25
title: 构建一个项目想法板
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该定义一个名为 `projectStatus` 的对象常量，包含三个键：`PENDING`、`SUCCESS` 和 `FAILURE`。每个状态都应该分配一个带有 `description` 键的对象，值分别为 `Pending Execution`、`Executed Successfully` 和 `Execution Failed`。
2. 您应该定义一个名为 `ProjectIdea` 的类，其 `constructor` 接受 `title` 和 `description` 作为参数。使用提供的参数初始化 `title` 和 `description` 属性。该类还应该有一个名为 `status` 的属性，默认设置为 `projectStatus.PENDING` 的值。
3. 您应该在 `ProjectIdea` 类中定义一个名为 `updateProjectStatus` 的方法。此方法应该接受一个 `newStatus` 参数，并将 `status` 属性更新为给定值。
4. 您应该定义一个 `ProjectIdeaBoard` 类，其 `constructor` 接受一个 `title` 并初始化一个名为 `ideas` 的空数组，用于保存 `ProjectIdea` 类的实例。
5. 您应该在 `ProjectIdeaBoard` 类中定义一个名为 `pin` 的方法，该方法接受 `ProjectIdea` 类的一个实例，并将给定实例推送到 `ideas` 数组中。
6. 您应该在 `ProjectIdeaBoard` 类中定义一个名为 `unpin` 的方法。此方法应该接受 `ProjectIdea` 类的一个实例，并将其从 `ideas` 数组中移除。
7. 您应该定义一个名为 `count` 的方法，返回给定 `ProjectIdeaBoard` 数组中的项目想法数量。
8. 您应该定义一个名为 `formatToString` 的方法，以以下格式返回项目想法的名称、描述和状态：

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

您应该定义一个名为 `projectStatus` 的对象常量，包含正确的项目状态和描述。

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

您应该有一个 `ProjectIdea` 类。

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

您的 `ProjectIdea` 类应该根据传递的参数初始化 `title` 属性和 `description` 属性。它还应该将状态设置为 `projectStatus.PENDING`。

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

您的 `ProjectIdea` 类应该有一个 `updateProjectStatus` 方法。

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

您应该有一个 `ProjectIdeaBoard` 类。

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

您的 `ProjectIdeaBoard` 应该根据传递的参数初始化 `title` 属性，并初始化一个名为 `ideas` 的空数组，用于保存 `ProjectIdea` 类的实例。

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

您的 `ProjectIdeaBoard` 类应该有一个 `pin` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

您的 `ProjectIdeaBoard` 类应该有一个 `unpin` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

您的 `ProjectIdeaBoard` 类应该有一个 `count` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

您的 `ProjectIdeaBoard` 类应该有一个 `formatToString` 方法。

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` 应该返回 `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`。

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

在 `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` 上调用 `updateProjectStatus(projectStatus.SUCCESS)` 应该将状态更新为 `{ description: 'Executed Successfully' }`。

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

在 `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` 上调用 `updateProjectStatus(projectStatus.FAILURE)` 应该将状态更新为 `{ description: 'Execution Failed' }`。

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

在 `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` 上调用 `updateProjectStatus(projectStatus.SUCCESS)` 应该将状态更新为 `{ description: 'Executed Successfully' }`。

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

您应该能够使用 `pin` 方法将 `ProjectIdea` 对象固定到您的 `ProjectIdeaBoard` 上。

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

您应该能够使用 `unpin` 方法从您的 `ProjectIdeaBoard` 上取消固定 `ProjectIdea` 对象。

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

当 `new ProjectIdeaBoard("Empty Board")` 为空时，`emptyBoard.formatToString()` 应该返回 `Empty Board has 0 idea(s)\n`。

```js

let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

// Create an empty board
let emptyBoard = new ProjectIdeaBoard("Empty Board");

// Attempt to unpin a project from an empty board
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

当您将 `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` 固定到 `new ProjectIdeaBoard("Tech Projects Board")` 上时，`techProjects.formatToString()` 应该返回 `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`。

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```