---
id: 68da6138d70960c197ec5eee
title: 步骤 4
challengeType: 20
dashedName: step-4
---

# --description--

现在你将设置数据结构来存储结果。创建一个名为 [result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eed.md#L6-L6) 的变量并将其初始化为空列表。这个列表将存储你生成的所有有效括号组合。

更新你的 return 语句以返回 [result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eed.md#L6-L6) 而不是空列表。

# --hints--

你应该在你的 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 函数中声明一个名为 [result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eed.md#L6-L6) 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

你应该将 [result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eed.md#L6-L6) 初始化为空列表。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("result").is_equivalent("result = []")
`) })
```

你应该在函数末尾返回 [result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eed.md#L6-L6)。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'

    return []
--fcc-editable-region--
```
