---
id: 68e03c0b7327c7012f12e333
title: 步骤 14
challengeType: 20
dashedName: step-14
---

# --description--

最后，使用 `3` 作为参数调用 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 生成三对括号的所有五个有效组合，并将结果打印到控制台。这样工作坊就完成了。

# --hints--

你应该打印 `gen_parentheses(3)`。

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(gen_parentheses(3))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'

    queue = [('', 0, 0)]
    result = []
    while queue:
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            if closes_used < opens_used:
                queue.append((current + ')', opens_used, closes_used + 1))

    return result
--fcc-editable-region--
print(gen_parentheses(2))

--fcc-editable-region--
```

# --solutions--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'

    queue = [('', 0, 0)]
    result = []
    while queue:
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            if closes_used < opens_used:
                queue.append((current + ')', opens_used, closes_used + 1))

    return result

print(gen_parentheses(2))
print(gen_parentheses(3))
```
