---
id: 63ec3427fc3e9214c9ed2a14
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `#cart` 元素内，但在 `h2` 元素之后，添加一个 `div` 元素，并给它一个 `class` 属性，值为 `cart-items`。

# --hints--

你应该在 `#cart` 元素内有一个 `div` 元素。

```js
assert.exists(document.querySelector('#cart div'));
```

你的 `div` 元素应该在 `h2` 元素之后。

```js
const cart = document.querySelector('#cart');
const h2 = cart.querySelector('h2');
const div = cart.querySelector('div');
assert.isTrue(h2.nextElementSibling === div);
```

你的 `div` 元素应该有一个 `class` 属性，值为 `cart-items`。

```js
assert.equal(document.querySelector('#cart div')?.className, 'cart-items');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container"></div>
      </section>
      --fcc-editable-region--
      <section id="cart">
        <h2>购物车</h2>

      </section>
      --fcc-editable-region--
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```