---
id: 63ec19978a066607e23439f8
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `header` 元素内，添加一个 `h1` 元素，文本为 `MasterPuti 购物车`。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素应该在 `header` 元素内。

```js
assert.exists(document.querySelector('header h1'));
```

你的 `h1` 元素应该有文本 `MasterPuti 购物车`。

```js
assert.equal(document.querySelector('header h1')?.textContent, 'MasterPuti 购物车');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--
    <header>

    </header>
    <main>

    </main>
    --fcc-editable-region--
    <script src="./script.js"></script>
  </body>
</html>
```