---
id: 63eea817673c8e1c22927fa6
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

在第二个 `.product` 元素之后，添加另一个 `div` 元素，并给它一个 `class` 属性，值为 `product`。

# --hints--

你应该在 `.products-container` 元素内有三个 `div` 元素。

```js
assert.equal(document.querySelectorAll('.products-container div')?.length, 3);
```

新的 `div` 元素应该有一个 `class` 属性，值为 `product`。

```js
assert.equal(document.querySelectorAll('.products-container div')[2]?.className, 'product');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container">
          <div class="product">
            <h3>西瓜</h3>
            <p>$12.99</p>
            <button>添加到购物车</button>
          </div>
          <div class="product">
            <h3>香蕉</h3>
            <p>$0.99</p>
            <button>添加到购物车</button>
          </div>
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        <div class="cart-total">
          <h3>总计: $0</h3>
        </div>
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```