---
id: 63ee5e0f08e82208364c4128
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

在 `.product` 元素内，但在 `h3` 元素之后，添加一个 `p` 元素，文本为 `$12.99`。

# --hints--

你应该在 `.product` 元素内有一个 `p` 元素。

```js
assert.exists(document.querySelector('.product p'));
```

你的 `p` 元素应该在 `h3` 元素之后。

```js
const product = document.querySelector('.product');
const h3 = product.querySelector('h3');
const p = product.querySelector('p');
assert.isTrue(h3.nextElementSibling === p);
```

你的 `p` 元素应该有文本 `$12.99`。

```js
assert.equal(document.querySelector('.product p')?.textContent, '$12.99');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container">
          <div class="product">
            --fcc-editable-region--
            <h3>西瓜</h3>

            --fcc-editable-region--
          </div>
        </div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        <div class="cart-total">
          <h3>总计: $0</h3>
        </div>
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```