---
id: 63ee611d478dca0b77f6a393
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

在新的 `.product` 元素内，添加一个 `h3` 元素，文本为 `香蕉`。

# --hints--

你应该在新的 `.product` 元素内有一个 `h3` 元素。

```js
assert.exists(document.querySelectorAll('.product')[1]?.querySelector('h3'));
```

你的 `h3` 元素应该有文本 `香蕉`。

```js
assert.equal(document.querySelectorAll('.product')[1]?.querySelector('h3')?.textContent, '香蕉');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container">
          <div class="product">
            <h3>西瓜</h3>
            <p>$12.99</p>
            <button>添加到购物车</button>
          </div>
          --fcc-editable-region--
          <div class="product">

          </div>
          --fcc-editable-region--
        </div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        <div class="cart-total">
          <h3>总计: $0</h3>
        </div>
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```