---
id: 63f017b4ad028a148eb713c0
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

现在你需要添加一些 JavaScript 来使购物车功能正常工作。在 `script.js` 文件中，声明一个 `cart` 变量，并将其初始化为空数组。

# --hints--

你应该声明一个 `cart` 变量。

```js
assert.match(code, /var\s+cart|let\s+cart|const\s+cart/);
```

你的 `cart` 变量应该初始化为空数组。

```js
assert.match(code, /var\s+cart\s*=\s*\[\s*\]|let\s+cart\s*=\s*\[\s*\]|const\s+cart\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js

```