---
id: 63f0348a54a177272071a595
title: 步骤 43
challengeType: 0
dashedName: step-43
---

# --description--

在 `li` 变量下方，将 `li` 元素的 `textContent` 设置为 `${item.name} - $${item.price}`。

# --hints--

你应该将 `li` 元素的 `textContent` 设置为 `${item.name} - $${item.price}`。

```js
assert.match(code, /li\.textContent\s*=\s*`.*\$\{item\.name\}.*-.*\$\$.*\$\{item\.price\}.*`/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

  for (let i = 0; i < cart.length; i++) {
    const item = cart[i];
    const li = document.createElement("li");

  }

}

```