---
id: 63f0311f5ea9382388d6124f
title: 步骤 40
challengeType: 0
dashedName: step-40
---

# --description--

在 `displayCart` 函数内，使用 `document.getElementById()` 获取 `cart-items` 元素，并将其 `innerHTML` 设置为空字符串。这将清除购物车中任何现有的内容。

# --hints--

你应该使用 `document.getElementById()` 获取 `cart-items` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)cart-items\1\s*\)/);
```

你应该将元素的 `innerHTML` 设置为空字符串。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)cart-items\1\s*\)\.innerHTML\s*=\s*('|"|`)\2/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {

}

```