---
id: 63f0224ceb16dc196d2c860a
title: 步骤 30
challengeType: 0
dashedName: step-30
---

# --description--

在 `product` 变量下方，将该商品添加到 `cart` 数组中。

# --hints--

你应该使用 `cart.push()` 方法将商品添加到购物车中。

```js
assert.match(code, /cart\.push\(/);
```

你应该将 `product` 变量作为参数传递给 `cart.push()`。

```js
assert.match(code, /cart\.push\(\s*product\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);

}
```