---
id: 63f0378e173e3c3b7638b528
title: 步骤 49
challengeType: 0
dashedName: step-49
---

# --description--

在将 `button` 元素添加到 `li` 元素的代码下方，为 `button` 元素添加一个点击事件监听器。当按钮被点击时，它应该从 `cart` 数组中移除相应的商品，并调用 `displayCart` 和 `calculateTotal` 函数来更新页面。

# --hints--

你应该为 `button` 元素添加一个点击事件监听器。

```js
assert.match(code, /button\.addEventListener\(\s*('|"|`)click\1/);
```

当按钮被点击时，你应该从 `cart` 数组中移除相应的商品。

```js
assert.match(code, /button\.addEventListener\(\s*('|"|`)click\1\s*,\s*\(\s*\)\s*=>\s*{\s*cart\.splice\(\s*i\s*,\s*1\s*\)/);
```

当按钮被点击时，你应该调用 `displayCart` 函数。

```js
assert.match(code, /button\.addEventListener\(\s*('|"|`)click\1\s*,\s*\(\s*\)\s*=>\s*{\s*cart\.splice\(\s*i\s*,\s*1\s*\)\s*;\s*displayCart\(\s*\)/);
```

当按钮被点击时，你应该调用 `calculateTotal` 函数。

```js
assert.match(code, /button\.addEventListener\(\s*('|"|`)click\1\s*,\s*\(\s*\)\s*=>\s*{\s*cart\.splice\(\s*i\s*,\s*1\s*\)\s*;\s*displayCart\(\s*\)\s*;\s*calculateTotal\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
  displayCart();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

  for (let i = 0; i < cart.length; i++) {
    const item = cart[i];
    const li = document.createElement("li");
    li.textContent = `${item.name} - $${item.price}`;
    const button = document.createElement("button");
    button.textContent = "删除";
    li.appendChild(button);

  }

}

```