---
id: 63f0295e673b661ccb299e8a
title: 步骤 34
challengeType: 0
dashedName: step-34
---

# --description--

现在你需要添加一个函数来计算购物车中商品的总价。在 `addItem` 函数下方，声明一个 `calculateTotal` 函数。

# --hints--

你应该声明一个 `calculateTotal` 函数。

```js
assert.match(code, /function\s+calculateTotal|const\s+calculateTotal\s*=\s*\(.*\)\s*=>|let\s+calculateTotal\s*=\s*\(.*\)\s*=>|var\s+calculateTotal\s*=\s*\(.*\)\s*=>/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

document.getElementById("add-1").addEventListener("click", () => {
  addItem(1);
});

document.getElementById("add-2").addEventListener("click", () => {
  addItem(2);
});

document.getElementById("add-3").addEventListener("click", () => {
  addItem(3);
});

```