---
id: 63f03446c2ed3e264be6c7fc
title: 步骤 42
challengeType: 0
dashedName: step-42
---

# --description--

在 `item` 变量下方，声明一个 `li` 变量，并将其设置为一个新的 `li` 元素。

# --hints--

你应该声明一个 `li` 变量。

```js
assert.match(code, /var\s+li|let\s+li|const\s+li/);
```

你应该将 `li` 变量设置为一个新的 `li` 元素。

```js
assert.match(code, /var\s+li\s*=\s*document\.createElement\(\s*('|"|`)li\1\s*\)|let\s+li\s*=\s*document\.createElement\(\s*('|"|`)li\1\s*\)|const\s+li\s*=\s*document\.createElement\(\s*('|"|`)li\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

  for (let i = 0; i < cart.length; i++) {
    const item = cart[i];

  }

}

```