---
id: 63f0374d5351223a747c301d
title: 步骤 48
challengeType: 0
dashedName: step-48
---

# --description--

在设置 `button` 元素的 `textContent` 的代码下方，将 `button` 元素添加到 `li` 元素中。

# --hints--

你应该将 `button` 元素添加到 `li` 元素中。

```js
assert.match(code, /li\.appendChild\(\s*button\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
  displayCart();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

  for (let i = 0; i < cart.length; i++) {
    const item = cart[i];
    const li = document.createElement("li");
    li.textContent = `${item.name} - $${item.price}`;
    const button = document.createElement("button");
    button.textContent = "删除";

  }

}

```