---
id: 63f02bdeb9b428208b97eb6b
title: 步骤 38
challengeType: 0
dashedName: step-38
---

# --description--

现在你需要在每次添加商品到购物车时调用 `calculateTotal` 函数。在 `addItem` 函数的末尾，调用 `calculateTotal()`。

# --hints--

你应该在 `addItem` 函数的末尾调用 `calculateTotal()`。

```js
assert.match(code, /function\s+addItem\s*\(\s*id\s*\)\s*{\s*const\s+product\s*=\s*products\.find\(\s*$$\s*item\s*$$\s*=>\s*item\.id\s*===\s*id\s*$$\s*;\s*cart\.push\(\s*product\s*$$\s*;\s*calculateTotal\(\s*$$\s*}/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

document.getElementById("add-1").addEventListener("click", () => {
  addItem(1);
});

document.getElementById("add-2").addEventListener("click", () => {
  addItem(2);
});

document.getElementById("add-3").addEventListener("click", () => {
  addItem(3);
});

```