---
id: 63ec20a06fff670d37befbd9
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `#products` 元素之后，添加另一个 `section` 元素，并给它一个 `id` 属性，值为 `cart`。

# --hints--

你应该有另一个 `section` 元素。

```js
assert.equal(document.querySelectorAll('section')?.length, 2);
```

新的 `section` 元素应该在 `#products` 元素之后。

```js
const products = document.querySelector('#products');
const cart = document.querySelector('#cart');
assert.isTrue(products.nextElementSibling === cart);
```

新的 `section` 元素应该有一个 `id` 属性，值为 `cart`。

```js
assert.equal(document.querySelector('#cart')?.id, 'cart');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      --fcc-editable-region--
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container"></div>
      </section>

      --fcc-editable-region--
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```