---
id: 63ec47b454495519739486a7
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在 `.cart-total` 元素内，添加一个 `h3` 元素，文本为 `总计: $0`。

# --hints--

你应该在 `.cart-total` 元素内有一个 `h3` 元素。

```js
assert.exists(document.querySelector('.cart-total h3'));
```

你的 `h3` 元素应该有文本 `总计: $0`。

```js
assert.equal(document.querySelector('.cart-total h3')?.textContent, '总计: $0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container"></div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        --fcc-editable-region--
        <div class="cart-total">

        </div>
        --fcc-editable-region--
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```