---
id: 63ec36f6133df7160be3ec66
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在 `#cart` 元素内，但在 `.cart-items` 元素之后，添加一个 `div` 元素，并给它一个 `class` 属性，值为 `cart-total`。

# --hints--

你应该在 `#cart` 元素内有另一个 `div` 元素。

```js
assert.equal(document.querySelectorAll('#cart div')?.length, 2);
```

新的 `div` 元素应该在 `.cart-items` 元素之后。

```js
const cartItems = document.querySelector('.cart-items');
const cartTotal = document.querySelector('.cart-total');
assert.isTrue(cartItems.nextElementSibling === cartTotal);
```

新的 `div` 元素应该有一个 `class` 属性，值为 `cart-total`。

```js
assert.equal(document.querySelector('.cart-total')?.className, 'cart-total');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container"></div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        --fcc-editable-region--
        <div class="cart-items"></div>

        --fcc-editable-region--
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```