---
id: 63f038a0ae041d3c5b0cdf23
title: 步骤 50
challengeType: 0
dashedName: step-50
---

# --description--

现在你需要为总价元素添加一个 `id` 属性，这样你就可以在 JavaScript 中引用它。在 HTML 文件中，给包含总价的 `h3` 元素添加一个 `id` 属性，值为 `total`。

# --hints--

你的 `h3` 元素应该有一个 `id` 属性，值为 `total`。

```js
assert.equal(document.getElementById('total')?.tagName, 'H3');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container">
          <div class="product" data-id="1" data-price="12.99">
            <h3>西瓜</h3>
            <p>$12.99</p>
            <button id="add-1">添加到购物车</button>
          </div>
          <div class="product" data-id="2" data-price="0.99">
            <h3>香蕉</h3>
            <p>$0.99</p>
            <button id="add-2">添加到购物车</button>
          </div>
          <div class="product" data-id="3" data-price="0.89">
            <h3>橙子</h3>
            <p>$0.89</p>
            <button id="add-3">添加到购物车</button>
          </div>
        </div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        <div class="cart-total">
          --fcc-editable-region--
          <h3>总计: $0</h3>
          --fcc-editable-region--
        </div>
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```