---
id: 63f02a4ef92d711ec1ff618c
title: 步骤 36
challengeType: 0
dashedName: step-36
---

# --description--

在 `total` 变量下方，使用 `for` 循环遍历 `cart` 数组。在循环内部，将每个商品的价格加到 `total` 变量上。

# --hints--

你应该使用 `for` 循环遍历 `cart` 数组。

```js
assert.match(code, /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*cart\.length\s*;\s*i\s*\+\+\s*\)/);
```

你应该将每个商品的价格加到 `total` 变量上。

```js
assert.match(code, /total\s*\+=\s*cart\s*\[\s*i\s*\]\.price/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

function calculateTotal() {
  let total = 0;

}

document.getElementById("add-1").addEventListener("click", () => {
  addItem(1);
});

document.getElementById("add-2").addEventListener("click", () => {
  addItem(2);
});

document.getElementById("add-3").addEventListener("click", () => {
  addItem(3);
});

```