---
id: 63f018f04e487e164dc27bd9
title: 步骤 28
challengeType: 0
dashedName: step-28
---

# --description--

现在你需要添加一个函数来处理将商品添加到购物车的操作。在 `products` 变量下方，声明一个 `addItem` 函数，它接受一个 `id` 参数。

# --hints--

你应该声明一个 `addItem` 函数。

```js
assert.match(code, /function\s+addItem|const\s+addItem\s*=\s*\(.*\)\s*=>|let\s+addItem\s*=\s*\(.*\)\s*=>|var\s+addItem\s*=\s*\(.*\)\s*=>/);
```

你的 `addItem` 函数应该接受一个 `id` 参数。

```js
assert.match(code, /function\s+addItem\s*\(\s*id\s*\)|const\s+addItem\s*=\s*\(\s*id\s*\)\s*=>|let\s+addItem\s*=\s*\(\s*id\s*\)\s*=>|var\s+addItem\s*=\s*\(\s*id\s*\)\s*=>/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

```