---
id: 63f029b96b9e9e1df93be951
title: 步骤 35
challengeType: 0
dashedName: step-35
---

# --description--

在 `calculateTotal` 函数内，声明一个 `total` 变量，并将其初始化为 `0`。

# --hints--

你应该声明一个 `total` 变量。

```js
assert.match(code, /var\s+total|let\s+total|const\s+total/);
```

你的 `total` 变量应该初始化为 `0`。

```js
assert.match(code, /var\s+total\s*=\s*0|let\s+total\s*=\s*0|const\s+total\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

function calculateTotal() {

}

document.getElementById("add-1").addEventListener("click", () => {
  addItem(1);
});

document.getElementById("add-2").addEventListener("click", () => {
  addItem(2);
});

document.getElementById("add-3").addEventListener("click", () => {
  addItem(3);
});

```