---
id: 63f02b22cce1c11fe9604381
title: 步骤 37
challengeType: 0
dashedName: step-37
---

# --description--

在 `for` 循环下方，使用 `document.getElementById()` 获取 `total` 元素，并将其 `textContent` 设置为 `total` 变量的值。确保使用 `toFixed(2)` 方法将总价格式化为保留两位小数。

# --hints--

你应该使用 `document.getElementById()` 获取 `total` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)total\1\s*\)/);
```

你应该将元素的 `textContent` 设置为 `total` 变量的值。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)total\1\s*\)\.textContent\s*=\s*total/);
```

你应该使用 `toFixed(2)` 方法将总价格式化为保留两位小数。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)total\1\s*\)\.textContent\s*=\s*total\.toFixed\(\s*2\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

}

document.getElementById("add-1").addEventListener("click", () => {
  addItem(1);
});

document.getElementById("add-2").addEventListener("click", () => {
  addItem(2);
});

document.getElementById("add-3").addEventListener("click", () => {
  addItem(3);
});

```