---
id: 6852dd8486c7ee0171322950
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你需要一个变量来保存解码过程中的秘密代码，所以移除 `pass`，在函数内创建一个名为 `secret_code` 的变量，并赋给它一个空字符串。

# --hints--

你的 `pin_extractor` 函数内应该有一个 `secret_code` 变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

变量 `secret_code` 应该有一个空字符串的值。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

你的代码中不应该再有 `pass`。

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    pass
--fcc-editable-region--
```
