---
id: 68531fbb934d735415053274
title: 步骤 4
challengeType: 20
dashedName: step-4
---

# --description--

现在你已经添加了函数调用，你可以在终端中看到函数内部产生的任何错误，并且你总是可以使用 `print` 来查看事物的值。

每行中都隐藏着PIN码的一个数字，所以在函数内使用 `split` 方法将字符串分割成行列表。在换行符（`\n`）处分割行，并将结果列表赋给一个名为 `lines` 的变量。

# --hints--

你的 `pin_extractor` 函数内应该有一个名为 `lines` 的变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

变量 `lines` 应该有 `poem.split('\n')` 的值。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--
poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
