---
id: 68530231c05b93324e071104
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在你处理函数的过程中，有一个函数调用用于调试会很有用，所以在函数外部创建一个包含这首诗的 `poem` 变量：

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

你可以使用多行字符串来实现。

然后，使用 `poem` 变量作为参数调用函数。

# --hints--

你应该有一个名为 `poem` 的全局变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

`poem` 变量的值应该等于 `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`。

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

You should call the function `pin_extractor` with `poem` as argument.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
```
