---
id: 6855918e87c33c6979d21d3f
title: 步骤 14
challengeType: 20
dashedName: step-14
---

# --description--

现在这个函数在从诗歌中提取代码方面做得很好，但有一个边界情况需要考虑：诗歌的行可能比预期的短，这会使函数出错（你可以通过从诗歌中删除单词 `bright` 来尝试这个情况）。

将 `secret_code += str(len(words[line_index]))` 这一行放在一个 `if` 语句中，检查 `words` 列表中是否有足够的单词。

# --hints--

在循环内你应该有一个 `if` 语句。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

if 语句的条件应该是 `len(words) > line_index`。


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Inside the `if` statement you should have `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        secret_code += str(len(words[line_index]))
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
