---
id: 685535bf527b2c77236b3a9a
title: 步骤 6
challengeType: 20
dashedName: step-6
---

# --description--

PIN码的第n位数字隐藏在第n行的第n个单词的长度中。

为了找到第n个单词的长度，下一步是将诗歌的行分割成单词列表：在循环内，创建一个变量 `words`，并将 `line` 使用 `split` 方法分割成单词的值赋给它。

然后，添加一个 `print` 调用，以 `words` 作为参数。

# --hints--

你的循环内应该有一个 `words` 变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

`words` 变量应该有 `line.split()` 的值。

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

You should have `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
