---
id: 68559f3561d7df287586142e
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

现在PIN码提取器可以工作了，但它可以更高效！它可以同时从许多诗歌中提取PIN码！

在全局作用域中添加两首更多的诗歌，`poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` 和 `poem3 = 'There\nonce\nwas\na\ndragon'`，现在，将调用 `pin_extractor` 的 `print` 语句注释掉。


# --hints--

变量 `poem2` 和 `poem3` 应该存在并且有正确的值。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\nhere and there\\nhoping for rain\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\nonce\\nwas\\na\\ndragon'")
    `)
})
```

你应该注释掉 `print(pin_extractor(poem))` 调用。

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
