---
id: 68557bd29f7f3c1880f98836
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

在循环的最后一行，更新 `print`，使其打印 `words` 列表中索引为 `line_index` 的单词。

# --hints--

你的循环末尾应该有 `print(words[line_index])`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words[line_index])')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
        print(words)


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
