---
id: 685578aae65e4106ad2b4300
title: 步骤 7
challengeType: 20
dashedName: step-7
---

# --description--

正如你在之前的课程中学到的，`enumerate` 函数允许在遍历可迭代对象时跟踪索引：

```py
fruits = ['apple', 'plum', 'bananas'] 

for index, item in enumerate(fruits):  
    print(index, item)  

# 0 apple  
# 1 plum
# 2 bananas
```

秘密代码的第n位数字来自第n行的第n个单词，所以你需要知道每一行的编号。

修改循环，使其在 `enumerate(lines)` 上迭代，并在 `line` 之前添加另一个名为 `line_index` 的循环变量。

同时，将 `print` 调用更新为 `print(line_index, line)`。

# --hints--

循环应该在 `enumerate(lines)` 上迭代。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

循环应该有 `line_index, line` 作为循环变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

You should have `print(line_index, line)` inside the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
