---
id: 68592b42b9c7f27cbe123203
title: 步骤 15
challengeType: 20
dashedName: step-15
---

# --description--

诗歌的第三行缺少第三个单词，所以PIN码比预期的短。

添加一个 `else` 子句，当单词不够时向 `secret_code` 追加 `'0'`，这样诗歌的所有行都会被用来查找数字。

# --hints--

你应该有一个 `else` 子句。


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
})
```

`else` 子句应该包含 `secret_code += '0'`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
})
```

`pin_extractor(poem)` should return `5202`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5202'
    `)
})
```

`secret_code('There\nonce\nwas\na\ndragon')` should return `50000`.

```js
({
    test: () => runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
