---
id: 6859a037b3d6721c7fbb8055
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

在循环之前，创建一个新变量 `secret_codes`，其初始值为空列表。将 `return secret_code` 这一行替换为将 `secret_code` 添加到 `secret_codes` 列表的一行。

# --hints--

你的 `pin_extractor` 函数内应该有一个 `secret_codes` 变量。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

变量 `secret_codes` 应该有一个空列表的值。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

At the end of the loop you should have `secret_codes.append(secret_code)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

You should not return `secret_code`.


```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poems):

    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))
--fcc-editable-region--

```
