---
id: 685926db99041c5e619ead41
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

`len` 函数给出字符串的长度。更新第二个 `print` 调用，打印那个单词的长度。

# --hints--

你的循环末尾应该有 `print(len(words[line_index]))`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
        print(words[line_index])


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
