---
id: 68592730aaf34b61ae279168
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

`len` 返回的长度是一个整数，你需要使用 `str` 函数将其转换为字符串。

更新 `print` 调用，将长度作为字符串打印。

# --hints--

你的循环末尾应该有 `print(str(len(words[line_index])))`。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
        print(len(words[line_index]))


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
