---
id: 68e448ec4ac4205968c0a805
title: ?9 ?
challengeType: 1
dashedName: step-9
---

# --description--

如果句子?`"love"` 这个词能重复三次，是不是更充满热情呢？让我们来实现它?

还记得之前学过的 `repeat()` 方法吗？它可以让一个字符串重复指定的次数。示例如下：

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

因为字符串是不可变的，这个方法不会直接修改原始内容，而是返回一个带有重复结果的新字符串?

现在创建一个名?`repeatedLove` 的变量，并赋值为 `"love ".repeat(3)`。再输出它看看成果，多出来的热情正在闪光?

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该创建一个名?`repeatedLove` 的变量，承载满满的“love”?

```js
assert.exists(repeatedLove);
```

你的 `repeatedLove` 变量应该是一个字符串?

```js
assert.isString(repeatedLove);
```

你需要将 `"love ".repeat(3)` 的结果赋值给 `repeatedLove`?

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

你还需要输?`repeatedLove`，欣赏一下它的样子。太可爱啦！

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
