---
id: 68e44519c769bf545bb3f22a
title: ?7 ?
challengeType: 1
dashedName: step-7
---

# --description--

要得到理想的结果 `"I love dogs and dogs are so much fun!"`，我们需要使用更强力、可爱又贴心?`replaceAll()` 方法来替换所有出现的 `cats`。它会为每一次匹配生成一个新的字符串?

请将 `exampleSentence.replace("cats", "dogs")` 更新为使?`replaceAll()`。完成后看看控制台，完美的句子已经呈现，太棒了！

# --hints--

你需要把那一?`exampleSentence.replace("cats", "dogs")` 改为使用 `replaceAll()`，全力替换！

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
