---
id: 68e43f94365a964f26c88f78
title: ?5 ?
challengeType: 1
dashedName: step-5
---

# --description--

现在我们来认识一个新的场景，继续进行愉快的字符串旅程?

创建一个名?`exampleSentence` 的新变量，并赋值为 `"I love cats and cats are so much fun!"`?

然后输出 `"Original sentence:"`，再输出 `exampleSentence` 的内容。慢慢来，每一步都让你更熟练！

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该创建一?`exampleSentence` 变量，让它闪亮登场?

```js
assert.exists(exampleSentence);
```

你的 `exampleSentence` 变量应该是一个字符串，保持内容纯净?

```js
assert.isString(exampleSentence);
```

你需要将字符?`"I love cats and cats are so much fun!"` 赋值给 `exampleSentence`?

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

你需要在控制台输?`"Original sentence:"` 作为标签?

```js
assert.equal(getLogs()[4], "Original sentence:");
```

你还需要输?`exampleSentence` 的内容，看看它的初始状态。继续冲鸭！

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
