---
id: 68e41084bb275539f72de9fe
title: ?1 ?
challengeType: 1
dashedName: step-1
---

# --description--

欢迎来到这个小巧又有趣的字符串变形工作坊！你将通过构建一个“字符串变换”小应用，开心地练习 `replace()`、`replaceAll()` 以及 `repeat()` 这些方法。它们就像三位魔法助手，帮你温柔地替换、批量替换，和可爱地重复字符串?

首先，让我们踏出冒险的第一步：创建一个名?`originalString` 的变量，并把可爱内容 `"I love cats."` 赋值给它。加油，这一步超级简单，你一定可以！

# --hints--

你应该创建一个名?`originalString` 的变量——从这里开始，我们就逐步点亮这段练习旅程啦！

```js
assert.exists(originalString);
```

你的 `originalString` 变量应该是一个字符串，小心地保持它的文字形态哦?

```js
assert.isString(originalString);
```

你需要把字符?`"I love cats."` 赋值给变量 `originalString`。完成后请微笑一下，你又前进了一步！

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
