---
id: 68e4417b7317835208145323
title: ?6 ?
challengeType: 1
dashedName: step-6
---

# --description--

继续前进：创建一个名?`dogsOnlySentence` 的变量，并赋值为 `exampleSentence.replace("cats", "dogs")`?

然后输出 `"Replacing all occurrences of cats with dogs:"`，再输出 `dogsOnlySentence` 的内容?

现在控制台会显示 `"I love dogs and cats are so much fun!"`。但还差一点点理想效果——别担心，下一步我们会把所有的 `cats` 都换掉，让句子真正变得整齐可爱！

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该创建一?`dogsOnlySentence` 变量?

```js
assert.exists(dogsOnlySentence);
```

你的 `dogsOnlySentence` 变量应该是一个字符串?

```js
assert.isString(dogsOnlySentence);
```

你需要将 `exampleSentence.replace("cats", "dogs")` 的结果赋值给 `dogsOnlySentence`?

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

你需要在控制台输?`"Replacing all occurrences of cats with dogs:"`?

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

你还需要输?`dogsOnlySentence` 的内容，观察当前的替换效果。马上就要完美啦?

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
