---
id: 68e43e884ef09d4d2ab05631
title: ?4 ?
challengeType: 1
dashedName: step-4
---

# --description--

鏄椂鍊欐楠屽垰鎵嶇殑鏇挎崲鎴愭灉鍟︼?

鍏堟坊鍔犱竴?`console.log()`锛屽唴瀹逛负 `"After using the replace() method:"`锛屽儚鏄湪瀹ｅ竷榄旀硶瀹屾垚銆?

鍐嶆坊鍔犱竴鏉¤緭?`replacedString` ?`console.log()`?

鐪嬬湅鎺у埗鍙帮紝浣犲簲璇ョ湅鍒伴偅鍙ョ啛鎮夌殑璇濆凡缁忓彉鎴愪簡 `"I love dogs."`锛屽師鏉ョ殑灏忕尗鍜灏忕嫍鐙楁俯鏌旀浛鎹簡銆傚共寰楁紓浜?

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

浣犻渶瑕佸湪鎺у埗鍙拌緭鍑?`"After using the replace() method:"` 杩欎竴琛岋紝璁╃粨鏋滀华寮忔劅婊℃弧銆?

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

浣犺繕闇€瑕佽緭?`replacedString` 鍙橀噺锛屼互灞曠ず鎴愬姛鐨勫彉鍖栥€傜户缁繚鎸佽繖浠借妭濂忥?

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--
const replacedString = originalString.replace("cats", "dogs");

--fcc-editable-region--
```
