---
id: 68e44abea678a65b877864dc
title: ?10 ?
challengeType: 1
dashedName: step-10
---

# --description--

让我们把刚才重复的单词放进一句完整的句子里吧！创建一个名?`newSentence` 的变量：先拼?`"I "`，再拼接 `repeatedLove`，最后再拼接 `" learning."`。可以使?`+` 运算符，也可以使用模板字符串，随你选择?

完成后你会看?`"I love love love  learning."` 这句显得元气满满?

*注意*：最后一?`"love"` 后面多了一个空格，这是故意的，最后一步我们会把它优雅地去掉?

# --hints--

你应该创建一个名?`newSentence` 的变量?

```js
assert.exists(newSentence);
```

你的 `newSentence` 变量应该是一个字符串?

```js
assert.isString(newSentence);
```

你需要按顺序拼接：`"I "`、`repeatedLove`、`" learning."` 并把结果赋值给 `newSentence`。注意空格位置，让句子保持可爱整洁?

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

你不应该直接把整句字面量 `"I love love love  learning."` 赋值给 `newSentence`，要通过拼接来完成它哦。再读一遍说明，会更踏实?

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
