---
id: 68e44c305f1f935d949d76bb
title: ?11 ?
challengeType: 1
dashedName: step-11
---

# --description--

正如上一步所说，句子 `"I love love love  learning."` 里最后一?`"love"` 后有多余的空格。现在让我们?`trimEnd()` 方法小心地把它修剪掉，让句子精神抖擞?

?`"love ".repeat(3)` 之后链式调用 `trimEnd()`。方法链就像积木，可以这样堆叠：

```js
.firstMethod().secondMethod()
```

现在查看控制台，你会看到多余的空格已经消失啦?

伴随着这一小步的优雅收尾，你已经完成了整个工作坊！超级棒，继续保持热爱与好奇心?

# --hints--

你需要在 `"love ".repeat(3)` 后面链式调用 `trimEnd()`，让结尾干净利落。继续闪闪发光！

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
