---
id: 68e9a6bf1a74721a4e43e217
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

现在你将移除浏览器应用的默认复选框。在你的`input[type='checkbox']`选择器中，添加`appearance`属性，值为`none`。

将`appearance`属性设置为`none`将清除浏览器对复选框应用的外观，让你能够显示你想要的样式。

这样做之后，由于复选框将不再可见，设置一个`border`属性，值为`2px solid #f1be32`。

# --hints--

你应该为`input[type="checkbox]`选择器添加`appearance`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

你应该有一个值为`none`的`appearance`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

你应该为`input[type="checkbox]`选择器添加`border`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

你应该有一个值为`2px solid #f1be32`的`border`属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;

}
```
