---
id: 6853e58d52ec64349a5760ca
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个工作坊中，你将实现最短路径算法。你将编写一个Python函数来计算图中节点之间的最短路径，并返回所经过的路径。

例如，给定一个城市之间通过不同距离的道路连接的图，算法将找到从一个城市到另一个城市的最短路线。如果你想从城市A到城市D旅行，算法可能会发现经过A ⇨ B ⇨ C ⇨ D（总计：15公里）比直接经过A ⇨ D（20公里）更短。

要开始，请定义一个名为`INF`的变量，并将其赋值为`float('inf')`，这表示正无穷大。稍后，你将使用它来表示两个节点之间的无限距离。

# --hints--

你应该在文件顶部定义变量`INF`。

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

你应该将`float('inf')`赋值给变量`INF`。

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
