---
id: 61329d52e5010e08d9b9d66b
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

另一个对可访问性与 SEO 都很重要的 `meta` 是页面描述。搜索引擎会使用它的 `content` 值来展示你的页面简介。

请添加一个 `meta`，`name="description"`，并提供有意义且精炼的 `content` 内容。描述不宜过长，控制在合理范围，更利于搜索展示。继续，你的元信息越来越完整！

# --hints--

你应该在 `head` 中再添加一个新的 `meta`。

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

你应该给此 `meta` 设置 `name="description"`。

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

你应该给此 `meta` 设置一个非空的 `content`。

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

`content` 的长度不应超过 165 个字符。（这是常见的搜索展示上限）

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
