---
id: 61329b210dac0b08047fd6ab
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

继续了解 `meta`：`viewport` 可以告诉浏览器如何在不同设备上渲染页面，它提升移动端的视觉可访问性，也有利于 SEO。

请添加一个用于定义视口的 `meta`，通过 `content` 属性说明页面的宽度与初始缩放比例。继续加油，页面正在变得更友好！

# --hints--

你应该在 `head` 中再创建一个新的 `meta` 元素。

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

你应该给这个 `meta` 设置 `name="viewport"`。

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

你应该给这个 `meta` 设置 `content="width=device-width, initial-scale=1"`（或等价包含 `1.0`）。

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
