---
id: 613297a923965e0703b64796
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

你可能已经熟悉 `meta` 元素：它可以声明页面的标题、描述、关键词、作者等信息。

现在请添加一个 `meta` 元素，并给它的 `charset` 属性设置合理的值。

`charset` 用来指定页面的字符编码；现代浏览器基本都使用 `UTF-8` —— 选它就对啦！继续前进，稳稳提升～

# --hints--

你应该在 `head` 中创建一个 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你应该给该 `meta` 标签添加 `charset="UTF-8"`。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
