---
id: 6133acc353338c0bba9cb553
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `head` 中最后一个关键元素：`title`。它帮助屏幕阅读器快速了解页面主题，也是 SEO 的重要组成部分。

请为页面设定一个既清晰又简洁的标题，足够描述内容又不过分冗长。继续，你的页面结构正在完善！

# --hints--

你应该在 `head` 中添加一个 `title` 元素。

```js
assert.exists(document.querySelector('head > title'));
```

标题长度不应超过 60 个字符。

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

标题要足够描述性（提示：至少 15 个字符）。

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```

