---
id: 6133d11ef548f51f876149e3
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

良好的导航结构是可访问性的核心之一，屏幕阅读器需要你用语义化 HTML 来表达页面结构。

现在请添加一个 `header` 和一个 `main` 元素。

`header` 将放置页面引导与导航菜单；`main` 则承载页面的核心内容。一步步构筑结构，你做得超级棒！

# --hints--

你应该在 `body` 中添加一个 `header` 元素。

```js
assert.exists(document.querySelector('body > header'));
```

你应该在 `body` 中添加一个 `main` 元素。

```js
assert.exists(document.querySelector('body > main'));
```

`header` 应该位于 `main` 之前。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```

