---
id: 613e275749ebd008e74bb62e
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

当前这张 `img` 使用默认尺寸，看起来过大。CSS 的 `max()` 函数会返回一组逗号分隔值中最大的一个。例如：

```css
img {
  width: max(250px, 25vw);
}
```

如果视口宽度小于 1000px，宽度为 250px；否则为 25vw（视口宽度的 25%）。

现在请通过 `id` 选择器缩放图片，将它的 `width` 设置为 `max(10rem, 18vw)`。继续，响应式也在一步步完善！

# --hints--

你应该使用 `#logo` 选择器定位该图片。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

你应该给该图片设置 `width: max(10rem, 18vw)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```

