---
id: 614202874ca576084fca625f
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

每个拥有 `region` 角色的区域都应该有一个可识别的标签，帮助屏幕阅读器用户理解其用途。常见做法：区域内放置一个标题，再用 `aria-labelledby` 指向该标题的 `id`。

请分别为三个 `section` 添加如下 `aria-labelledby`：

- `student-info`
- `html-questions`
- `css-questions`

然后在对应的 `section` 中各嵌套一个 `h2`，其 `id` 要与上面设置的值匹配，并填入合适的文本内容（例如：`Student Info`、`HTML`、`CSS`）。继续，语义标签更加完善！

# --hints--

你需要给第一个 `section` 添加 `aria-labelledby="student-info"`。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

你需要给第二个 `section` 添加 `aria-labelledby="html-questions"`。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

你需要给第三个 `section` 添加 `aria-labelledby="css-questions"`。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

你需要在第一个 `section` 中嵌套一个 `h2`。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

你需要在第二个 `section` 中嵌套一个 `h2`。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

你需要在第三个 `section` 中嵌套一个 `h2`。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

你需要让第一个 `h2` 的 `id` 为 `student-info`。

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

你需要让第二个 `h2` 的 `id` 为 `html-questions`。

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

你需要让第三个 `h2` 的 `id` 为 `css-questions`。

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

你需要给第一个 `h2` 添加合适的文本（例如 `Student Info`）。

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

你需要给第二个 `h2` 添加合适的文本（例如 `HTML`）。

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

你需要给第三个 `h2` 添加合适的文本（例如 `CSS`）。

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```

