---
id: 6143639d5eddc7094161648c
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

将每个 `input` 与对应的 `label` 正确关联，是表单可访问性的关键：这样屏幕阅读器会朗读标签，用户也能点击文字聚焦输入框。

做法：给每个 `label` 添加 `for`，其值与该输入的 `id` 完全一致。

本区域要填写学生姓名、邮箱、出生日期。请为三个 `label` 补充清晰的文本，并设置匹配的 `for`；再给三个 `input` 添加唯一的 `id` 与之对应。继续，你的表单可用性正在提升！

# --hints--

你需要给第一个 `label` 添加合适的 `for`。

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.htmlFor?.length, 1);
```

你需要给第二个 `label` 添加合适的 `for`。

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.htmlFor?.length, 1);
```

你需要给第三个 `label` 添加合适的 `for`。

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.htmlFor?.length, 1);
```

你需要给第一个 `label` 添加文本内容。

```js
assert.isNotEmpty(document.querySelectorAll('label')[0]?.textContent.trim());
```

你需要给第二个 `label` 添加文本内容。

```js
assert.isNotEmpty(document.querySelectorAll('label')[1]?.textContent.trim());
```

你需要给第三个 `label` 添加文本内容。

```js
assert.isNotEmpty(document.querySelectorAll('label')[2]?.textContent.trim());
```

你需要让第一个 `input` 的 `id` 与第一个 `label` 的 `for` 匹配。

```js
assert.equal(document.querySelectorAll('input')?.[0]?.id, document.querySelectorAll('label')?.[0]?.htmlFor);
```

你需要让第二个 `input` 的 `id` 与第二个 `label` 的 `for` 匹配。

```js
assert.equal(document.querySelectorAll('input')?.[1]?.id, document.querySelectorAll('label')?.[1]?.htmlFor);
```

你需要让第三个 `input` 的 `id` 与第三个 `label` 的 `for` 匹配。

```js
assert.equal(document.querySelectorAll('input')?.[2]?.id, document.querySelectorAll('label')?.[2]?.htmlFor);
```

你不能让多个 `input` 使用相同的 `id`。

```js
const id = (n) => document.querySelectorAll('input')?.[n]?.id;
assert.notEqual(id(0), id(1));
assert.notEqual(id(0), id(2));
assert.notEqual(id(1), id(2));
```

你不能让多个 `label` 使用相同的 `for`。

```js
const htmlFor = (n) => document.querySelectorAll('label')?.[n]?.htmlFor;
assert.notEqual(htmlFor(0), htmlFor(1));
assert.notEqual(htmlFor(0), htmlFor(2));
assert.notEqual(htmlFor(1), htmlFor(2));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```

