---
id: 6143931a113bb80c45546287
title: 步骤-23
challengeType: 0
dashedName: step-23
---

# --description--

在第二个 `section` 中添加两个 `div`，类名为 `question-block`，用于承载每一道题。

接着在每个 `div.question-block` 中加入一个从 `1` 开始递增数字的 `h3` 作为题号，以及一个带 `question` 类的 `fieldset` 用于后续题目选项。继续，题目结构马上成型！

# --hints--

你需要在第二个 `section` 中放入两个 `div`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.length, 2);
```

你需要让第一个新 `div` 的类名为 `question-block`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[0]?.className, 'question-block');
```

你需要让第二个新 `div` 的类名为 `question-block`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[1]?.className, 'question-block');
```

你需要在每个 `div.question-block` 中放入一个 `h3`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')?.length, 2);
```

你需要让第一个 `h3` 的文本为 `1`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')[0]?.textContent.trim(), '1');
```

你需要让第二个 `h3` 的文本为 `2`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > h3')[1]?.textContent.trim(), '2');
```

你需要在每个 `div.question-block` 中放入一个 `fieldset`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.length, 2);
```

你需要让第一个题块中 `h3` 后紧跟 `fieldset`。

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block > h3 + fieldset'));
```

你需要让第二个题块中 `h3` 后紧跟 `fieldset`。

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block:nth-of-type(2) > h3 + fieldset'));
```

你需要让第一个 `fieldset` 的类名为 `question`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[0]?.className, 'question');
```

你需要让第二个 `fieldset` 的类名为 `question`。

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[1]?.className, 'question');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>

        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}


```

