---
id: 6144e818fd5ea704fe56081d
title: 步骤 28
challengeType: 0
dashedName: step-28
---

# --description--

为每个题目的 `fieldset` 添加语义性的 `name` 属性（例如第一题可用 `html-question-one`），再给两个 `ul` 都添加类名 `answers-list` 以便统一样式。

最后在每个 `fieldset` 的 `legend` 中加入题干文本，让屏幕阅读器和视觉用户都能理解题目内容。继续，题目结构更规范啦！

# --hints--

你需要给第一个 `fieldset` 一个合适的 `name`（例如 `html-question-one`）。

```js
assert.notEmpty(document.querySelectorAll('fieldset')?.[0]?.name);
```

你需要给第二个 `fieldset` 一个合适的 `name`（例如 `html-question-two`）。

```js
assert.notEmpty(document.querySelectorAll('fieldset')?.[1]?.name);
```

你需要让第一个 `ul` 的类名为 `answers-list`。

```js
assert.equal(document.querySelectorAll('fieldset > ul')?.[0]?.className, 'answers-list');
```

你需要让第二个 `ul` 的类名为 `answers-list`。

```js
assert.equal(document.querySelectorAll('fieldset > ul')?.[1]?.className, 'answers-list');
```

你需要给第一个 `legend` 添加题干文本。

```js
assert.isNotEmpty(document.querySelectorAll('legend')[0]?.textContent.trim());
```

你需要给第二个 `legend` 添加题干文本。

```js
assert.isNotEmpty(document.querySelectorAll('legend')[1]?.textContent.trim());
```

你不能让两个 `legend` 的文本完全相同。

```js
const legends = document.querySelectorAll('legend');
assert.lengthOf(legends, 2);
assert.notEqual(legends[0]?.textContent.trim(), legends[1]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="MasterPuti Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="MasterPuti" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://MasterPuti.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
          <div class="question-block">
            <h3><span class="sr-only">Question</span>1</h3>
            <fieldset class="question">
              <legend></legend>
              <ul>
                <li></li>
                <li></li>
              </ul>
            </fieldset>
          </div>
          <div class="question-block">
            <h3><span class="sr-only">Question</span>2</h3>
            <fieldset class="question">
              <legend></legend>
              <ul>
                <li></li>
                <li></li>
              </ul>
            </fieldset>
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

.sr-only {
  position: absolute;
  width: 1px;
  height: 1px;
  padding: 0;
  margin: -1px;
  overflow: hidden;
  clip: rect(0, 0, 0, 0);
  white-space: nowrap;
  border: 0;
}

```

