---
id: 614ccc21ea91ef1736b9b578
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

欢迎来到这个可访问性练习！你已经拥有一个基础模板。

第一步：给 `html` 元素添加一个 `lang` 属性，让屏幕阅读器能立刻知道页面语言。（英文可用 `en`）。

别忘了把样式表链接上，这样页面就能美美地展示啦～继续，加油，你已经开始闪光！

# --hints--

你需要给 `html` 元素添加一个 `lang` 属性。（提示：英文页面可用 `en`。）

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

你需要在 `head` 中拥有一个 `link` 元素。

```js
assert.exists(document.querySelector("head>link"));
```

你的 `link` 元素需要具备 `rel="stylesheet"` 和 `href="styles.css"`。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
