---
id: 68e51d8f163d3f8b70f9c4e5
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个工作坊中，你将练习使用 [includes()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 和 [slice()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 方法。

首先，创建一个名为 [fccSentence](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e51d8f163d3f8b70f9c4e5.md#L12-L12) 的变量，并赋值字符串 `"freeCodeCamp is a great place to learn web development."`。

# --hints--

你应该声明一个名为 [fccSentence](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e51d8f163d3f8b70f9c4e5.md#L12-L12) 的变量。

```js
assert.exists(fccSentence);
```

你的 [fccSentence](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e51d8f163d3f8b70f9c4e5.md#L12-L12) 变量应该是一个字符串。

```js
assert.isString(fccSentence);
```

你的 [fccSentence](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e51d8f163d3f8b70f9c4e5.md#L12-L12) 变量应该被赋值字符串 `"freeCodeCamp is a great place to learn web development."`。

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
