---
id: 68e6172875008fea742a63c9
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

正如你在之前的课程中回忆的那样，[includes()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 方法检查字符串是否包含特定的子字符串，并返回 [true](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/node/globals.d.ts#L107-L107) 或 [false](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/node/globals.d.ts#L108-L108)。

这里是一个使用 [includes()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 方法的例子：

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

创建一个名为 [hasFreeCodeCamp](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 的变量。然后将使用 [includes()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 方法检查 [fccSentence](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e51d8f163d3f8b70f9c4e5.md#L12-L12) 是否包含 `"freeCodeCamp"` 的结果赋值给它。

# --hints--

你应该声明一个名为 [hasFreeCodeCamp](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 的变量。

```js
assert.exists(hasFreeCodeCamp);
```

你应该使用 [fccSentence.includes("freeCodeCamp")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12)。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

你应该将 [fccSentence.includes("freeCodeCamp")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 的结果赋值给你的 [hasFreeCodeCamp](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 变量。

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

[fccSentence.includes("freeCodeCamp")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 的结果应该是 [true](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/node/globals.d.ts#L107-L107)。

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
