---
id: 68e6172975008fea742a63cd
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

现在你将探索 [slice()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 方法，它可以提取字符串的一部分。

创建一个名为 [message](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 的变量，并赋值字符串 `"Welcome to freeCodeCamp!"`。

然后使用 [console.log()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/node/console.d.ts#L82-L82) 向控制台输出 `"Here are some examples of the slice() method:"`。

# --hints--

你应该声明一个名为 [message](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 的变量。

```js
assert.exists(message);
```

你的 [message](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 变量应该是一个字符串。

```js
assert.isString(message);
```

你的 [message](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 变量应该是 `"Welcome to freeCodeCamp!"`。

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

你应该向控制台输出 `"Here are some examples of the slice() method:"`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
