---
id: 68e6172975008fea742a63cb
title: 步骤 5
challengeType: 1
dashedName: step-5
---

# --description--

现在是时候看看当子字符串未找到时会发生什么了。

创建一个名为 [hasJavaScript](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cb.md#L12-L12) 的变量。然后将使用 [includes()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 方法检查 [fccSentence](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e51d8f163d3f8b70f9c4e5.md#L12-L12) 是否包含 `"JavaScript"` 的结果赋值给它。

然后使用模板字面量或字符串连接向控制台输出消息 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`。将 `<hasJavaScript>` 替换为变量的实际值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该声明一个名为 [hasJavaScript](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cb.md#L12-L12) 的变量。

```js
assert.exists(hasJavaScript);
```

你应该使用 [fccSentence.includes("JavaScript")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cb.md#L12-L12)。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

你应该将 [fccSentence.includes("JavaScript")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cb.md#L12-L12) 的结果赋值给你的 [hasJavaScript](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cb.md#L12-L12) 变量。

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

[fccSentence.includes("JavaScript")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cb.md#L12-L12) 的结果应该是 [false](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/node/globals.d.ts#L108-L108)。

```js
assert.isFalse(hasJavaScript);
```

你应该输出消息 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`，其中 `<hasJavaScript>` 应该被替换为变量的实际值。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
