---
id: 68e6172975008fea742a63cc
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

[includes()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 方法是区分大小写的，这意味着它能区分大写和小写字母。

创建一个名为 [hasLowercaseFCC](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cc.md#L12-L12) 的变量。然后将使用 [includes()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172875008fea742a63c9.md#L12-L12) 方法检查 [fccSentence](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e51d8f163d3f8b70f9c4e5.md#L12-L12) 是否包含 `"freecodecamp"`（全小写）的结果赋值给它。

然后使用模板字面量或字符串连接向控制台输出消息 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`。将 `<hasLowercaseFCC>` 替换为变量的实际值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该声明一个名为 [hasLowercaseFCC](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cc.md#L12-L12) 的变量。

```js
assert.exists(hasLowercaseFCC);
```

你应该使用 [fccSentence.includes("freecodecamp")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cc.md#L12-L12)。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

你应该将 [fccSentence.includes("freecodecamp")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cc.md#L12-L12) 的结果赋值给你的 [hasLowercaseFCC](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cc.md#L12-L12) 变量。

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

[fccSentence.includes("freecodecamp")](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cc.md#L12-L12) 的结果应该是 [false](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/node/globals.d.ts#L108-L108)。

```js
assert.isFalse(hasLowercaseFCC);
```

你应该输出包含 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` 的消息，其中 `<hasLowercaseFCC>` 应该被替换为变量的实际值。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
