---
id: 68e6172975008fea742a63d0
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

你可以使用[slice()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12)从字符串的开头提取内容。

创建一个名为[greetingWord](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)的变量。然后将使用[message.slice(0, 7)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12)提取的第一个单词[Welcome](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)的结果赋值给它。

然后使用模板字面量或字符串连接将消息[The first word is "<greetingWord>".](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)记录到控制台，将[<greetingWord>](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)替换为变量的实际值。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该声明一个名为[greetingWord](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)的变量。

```js
assert.exists(greetingWord);
```

你的[greetingWord](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)变量应该是一个字符串。

```js
assert.isString(greetingWord);
```

你应该使用[message.slice(0, 7)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12)。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

你应该将[message.slice(0, 7)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12)的结果赋值给你的[greetingWord](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)变量。

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

你的[greetingWord](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)变量应该等于[Welcome](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)。

```js
assert.strictEqual(greetingWord, "Welcome");
```

你应该记录消息[The first word is "<greetingWord>".](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)，其中[<greetingWord>](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_chinese/workshop-string-inspector/68e6172975008fea742a63cf.md#L12-L12)应由变量的实际值替换。

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```