---
id: 68e6172975008fea742a63ce
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

[slice()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 方法提取字符串的一部分。它接受两个参数：起始索引（包含）和结束索引（不包含）。

这里有一个例子：

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

字符串索引从0开始。在 [message](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63cd.md#L12-L12) 字符串中，索引11处的字符是 `f`（"freeCodeCamp"的开始）。

创建一个名为 [platform](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12) 的变量。然后将使用 [message.slice(11, 23)](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12) 提取 `"freeCodeCamp"` 的结果赋值给它。

# --hints--

你应该声明一个名为 [platform](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12) 的变量。

```js
assert.exists(platform);
```

你的 [platform](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12) 变量应该是一个字符串。

```js
assert.isString(platform);
```

你应该使用 [message.slice(11, 23)](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12)。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

你应该将 [message.slice(11, 23)](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12) 的结果赋值给你的 [platform](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12) 变量。

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

你的 [platform](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-string-inspector/68e6172975008fea742a63ce.md#L12-L12) 变量应该等于 `"freeCodeCamp"`。

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
