---
id: 68ef6e42fa6e8a9a18a401fa
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

使用 [console.log()](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/node/console.d.ts#L82-L82) 向控制台输出 `"Here are some examples of the includes() method:"`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

你应该向控制台输出 `"Here are some examples of the includes() method:"`。

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
