---
id: 6735a7370e0ae93a4577c689
title: Design a Newspaper Layout
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

在这个练习中，你将使用 CSS Grid 设计一个报纸布局，练习网格行、网格列与区域排布等概念。

**目标：**按照下面的要求完成页面并通过所有测试。

**用户需求：**

1. 页面应包含一个 `main` 元素，并为其添加 `class="newspaper-layout"`。
2. `.newspaper-layout` 内应包含一个 `header` 元素，类名为 `title`，其中包含显示报纸名称的 `h1`。
3. `.newspaper-layout` 内应包含一个用于主篇幅的 `article`，类名为 `feature-article`。
4. `.feature-article` 应包含一个 `h2`（文章标题）和后续的 `p`（文章内容）。
5. 另外再添加三个较小的文章（`article`），分别使用类名 `small-article1`、`small-article2` 和 `small-article3`。
6. 每个较小文章应包含一个 `h3`（标题）和一个 `p`（内容）。
7. `.newspaper-layout` 内还应包含一个类名为 `secondary-article` 的 `article`，用于附加新闻内容。
8. `.secondary-article` 同样应包含 `h2` 和 `p`。
9. 页面中应包含一个 `figure` 元素，类名为 `cover-image`，用于显示代表报纸内容的图片。
10. 上述各部分（`title`、`feature-article`、`secondary-article`、`cover-image`、`small-article1`、`small-article2`、`small-article3`）应分别设置对应的 `grid-area` 名称，以便在网格中进行布局。
11. 使用 `grid-template-areas` 定义布局：
    - `title` 跨越顶行的三列；
    - 第二行为 `feature-article`（跨两列）和 `cover-image`（第三列）；
    - 第三行为 `secondary-article`（跨两列）和 `cover-image`（第三列）；
    - 第四行放置三个小文章（依次为 `small-article1`、`small-article2`、`small-article3`）。
12. `.newspaper-layout` 的列模版应分为三等份（例如 `1fr 1fr 1fr` 或 `repeat(3, 1fr)`）。
13. `.newspaper-layout` 的行模版应设置为 `auto 1fr 1fr 1fr`（也可以用 `auto repeat(3, 1fr)`）。
14. 在网格项目之间添加适当的间隙（`gap`）。
15. 确保 `.cover-image` 内的图片通过设置 `max-width: 100%` 能适应其容器。

**注意：** 请在 HTML 中正确链接样式表并应用 CSS。

# --hints--

下面是一些实现提示，帮助你完成布局：

请确保页面存在 `main.newspaper-layout`。

```js
assert.exists(document.querySelector('main'));
```

`.newspaper-layout` 内应包含 `header.title`。

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

`header` 内应包含一个非空的 `h1`。

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

请在 `.newspaper-layout` 中添加五个 `article` 元素，分别对应 `feature-article`、`small-article1`、`small-article2`、`small-article3` 和 `secondary-article`，并确保它们存在。

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`feature-article` 应包含 `h2` 与后续 `p`。

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

每个小文章应包含 `h3` 与后续 `p`。

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`secondary-article` 应包含 `h2` 与 `p`。

```js
const h2sec = document.querySelector(".secondary-article h2");
const psec = document.querySelector(".secondary-article h2 + p");
assert.exists(h2sec);
assert.exists(psec);
```

请包含 `.newspaper-layout figure.cover-image`，并在其中放置一张 `img`，且该图片应设置 `max-width: 100%` 以适应容器。

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
assert.exists(document.querySelector(".cover-image img"));
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

`.newspaper-layout` 的 `display` 应为 `grid`，并使用 `grid-template-areas`、`grid-template-columns`（三等分）以及 `grid-template-rows`（`auto 1fr 1fr 1fr`）来实现页面结构，同时设置合适的 `gap`。

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);

assert.match(getComputedStyle(document.querySelector('.newspaper-layout')).gridTemplateAreas, /^("|')title title title\1/);

const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.MasterPuti.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```

