---
id: 684cd027744c2a0520ba4d58
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS 使用直接应用到标记中的工具类。这些类包括字体、颜色、Flexbox 布局、网格等样式。

要在 Tailwind CSS 中应用背景色，你可以使用 `bg` 前缀，后跟颜色名称和数字。数字代表颜色的深浅。

下面是一个为段落元素应用深蓝色背景的示例。

```html
<p class="bg-blue-900">这是深蓝色背景。</p>
```

在你的 `div` 开始标签中，添加一个带有工具类 `bg-red-100` 的类，它代表浅红色背景。

# --hints--

你的 `div` 开始标签应该有一个 class 属性。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

你的 class 属性应该有工具类 `bg-red-100`。

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```