---
id: 684d0d369e2b3b0f288d2c20
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

如果能水平居中 `div` 元素就更好了。要同时对两边应用 `margin`，你可以使用以下模式：`mx-number` 或 `my-number`。

`mx` 设置左右边距，而 `my` 设置上下边距。

下面是为元素应用上下边距的示例：

```html
<p class="my-8">这是外边距。</p>
```

在你的 `div` 开始标签中，使用正确的工具类将左右边距设置为 `auto`。

# --hints--

你的 `div` 开始标签应该有 `mx-auto` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```