---
id: 684d11ffd15a531225d081e9
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

现在 `div` 元素的样式已经应用了，是时候为 `p` 元素应用一些样式了。

目前，段落元素的文本颜色是黑色。但如果颜色是深红色会更好。

在前面的步骤中，你学会了如何使用背景色，如：`bg-red-100`。你可以用类似的方法处理文本颜色，遵循以下模式：`text-color-shade`。

下面是为元素应用蓝色文本颜色的示例：

```html
<p class="text-blue-500">这是蓝色文本。</p>
```

首先为你的 `p` 元素添加一个 class 属性。然后，应用正确的工具类来设置红色文本颜色，深浅度设置为 `700`。

# --hints--

你的 `p` 元素应该有一个 class 属性。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

你的 `p` 开始标签应该有 `text-red-700` 工具类。

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```