---
id: 66ed41f912d0bb1dc62da5dd
title: Build a JavaScript Trivia Bot
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该将`"Hello! I'm your coding fun fact guide!"`记录到控制台，作为向用户的问候消息。
1. 你应该创建三个变量：`botName`、`botLocation`和`favoriteLanguage`，分别存储机器人的名字、它来自哪里以及它最喜欢的编程语言。
1. 你应该将`"My name is (botName) and I live on (botLocation)."`记录到控制台。
1. 你应该将`"My favorite programming language is (favoriteLanguage)."`记录到控制台。
1. 你应该使用`let`创建一个`codingFact`变量，并为其分配一个关于你的机器人最喜欢编程语言的有趣事实的字符串，并包含使用`favoriteLanguage`变量。
1. 你应该将`codingFact`记录到控制台。
1. 你应该使用`favoriteLanguage`变量再次重新分配`codingFact`变量为关于机器人最喜欢语言的新事实。
1. 你应该再次将`codingFact`记录到控制台。
1. 你应该再次重新分配`codingFact`变量为关于机器人最喜欢语言的另一个新事实，再次使用`favoriteLanguage`变量。
1. 你应该第三次将`codingFact`记录到控制台。
1. 你应该将`"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."`记录到控制台，作为机器人的告别语。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将`"Hello! I'm your coding fun fact guide!"`记录到控制台。

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

你应该声明一个`botName`变量。仔细检查拼写或大小写错误。

```js
assert.exists(botName);
```

你的`botName`变量应该是一个字符串。

```js
assert.isString(botName);
```

你应该声明一个`botLocation`变量。仔细检查拼写或大小写错误。

```js
assert.exists(botLocation);
```

你的`botLocation`变量应该是一个字符串。

```js
assert.isString(botLocation);
```

你应该声明一个`favoriteLanguage`变量。仔细检查拼写或大小写错误。

```js
assert.exists(favoriteLanguage);
```

你的`favoriteLanguage`变量应该是一个字符串。

```js
assert.isString(favoriteLanguage);
```

你应该将`"My name is (botName) and I live on (botLocation)."`记录到控制台，并将变量添加到字符串中。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

你应该将`"My favorite programming language is (favoriteLanguage)."`记录到控制台，并将变量添加到字符串中。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

你应该使用`let`声明一个新变量`codingFact`。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

你应该给`codingFact`一个包含`favoriteLanguage`的值。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

你应该将`codingFact`记录到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

你应该给`codingFact`分配一个新值，该值也包含`favoriteLanguage`，并将其记录到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

你应该第三次给`codingFact`赋值，该值也包含`favoriteLanguage`，并将其记录到控制台。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

你应该将`"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."`记录到控制台，并添加变量的值。

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js
// 向用户问好
console.log("Hello! I'm your coding fun fact guide!");

// 声明机器人信息变量
const botName = "TriviaBot";
const botLocation = "Internet";
const favoriteLanguage = "JavaScript";

// 介绍机器人
console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

// 分享有趣的编程事实
let codingFact = "你知道吗？" + favoriteLanguage + "最初被称为Mocha！";

console.log(codingFact);

// 更新编程事实
codingFact = "另一个有趣的事实：" + favoriteLanguage + "是在仅仅10天内创建的！";

console.log(codingFact);

// 再次更新编程事实
codingFact = "还有，" + favoriteLanguage + "是世界上最流行的编程语言。";

console.log(codingFact);

// 告别语
console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");
```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");
```