---
id: 6823d6244511f252c8300eed
title: æ­¥éª¤ 6
challengeType: 0
dashedName: step-6
---

# --description--

å½“ä½ éœ€è¦å°†ä¸€ä¸ªæ®µè½æ·»åŠ åˆ°ç½‘é¡µä¸Šæ—¶ï¼Œä½ å¯ä»¥åƒè¿™æ ·ä½¿ç”?`p` å…ƒç´ ï¼?

```html
<p>This is a paragraph element.</p>
```

æŠ?`Learn the skills to become a full stack developer` è½¬åŒ–æˆä¸€ä¸ªæ®µè½å…ƒç´ ã€?

# --hints--

ä½ çš„ `p` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector("p"));
```

ä½ çš„ `p` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`</elementName>`ã€?

```js
assert.match(code, /<\/p\s*\>/);
```

ä½ çš„ `p` å…ƒç´ çš„æ–‡æœ¬åº”è¯¥æ˜¯ `Learn the skills to become a full stack developer`ã€?

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer\s*\<\/p\s*\>/i);
```

ä½ çš„ `p` å…ƒç´ åº”è¯¥åœ?`h2` å…ƒç´ çš„ä¸‹æ–¹ã€?

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>Full Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full stack developer
--fcc-editable-region--
```

