---
id: 6823e036ea4b71553558c01b
title: æ­¥éª¤ 8
challengeType: 0
dashedName: step-8
---

# --description--

ä½ ä¼šæ³¨æ„åˆ°ï¼Œæ¯ä¸ªæ ‡é¢˜åœ¨é¢„è§ˆä¸­çœ‹èµ·æ¥éƒ½æœ‰ç‚¹ä¸åŒã€?æ®µè½ä¹Ÿæ˜¯å¦‚æ­¤ã€?

æ¥ä¸‹æ¥ï¼Œæ·»åŠ å¦ä¸€ä¸?`p` å…ƒç´ æ¥æ˜¾ç¤ºä»¥ä¸‹æ–‡æœ¬ï¼š

```md
HTML represents the content and structure of a webpage
```

# --hints--

ä½ åº”è¯¥æœ‰äº†ç¬¬äºŒä¸ª `p` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

ä½ çš„ `p` å…ƒç´ åº”è¯¥åŒ…å«æ–‡æœ¬ `HTML represents the content and structure of a webpage` ã€?ä½ å¯ä»¥ä»è¯´æ˜ä¸­å¤åˆ¶æ–‡æœ¬ã€?

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```

