---
id: 6823c1a0bcada44f32bf0bdc
title: æ­¥éª¤ 4
challengeType: 0
dashedName: step-4
---

# --description--

`h1` å…ƒç´ æ˜¯ç½‘é¡µçš„ä¸»æ ‡é¢˜ï¼Œæ¯ä¸ªé¡µé¢åªèƒ½æœ‰ä¸€ä¸ªã€?`h2` å…ƒç´ ä»£è¡¨å­æ ‡é¢˜ã€?æ¯ä¸ªé¡µé¢å¯ä»¥æœ‰å¤šä¸ªï¼Œå®ƒä»¬çœ‹èµ·æ¥å°±åƒè¿™æ ·ï¼š

```html
<h2>This is a subheading.</h2>
```

ç”?`h2` å¼€å§‹æ ‡ç­¾å’Œç»“æŸæ ‡ç­¾æŠ?`Full Stack Curriculum` æ‹¬èµ·æ¥ï¼Œå°†å…¶è½¬åŒ–ä¸ºä¸€ä¸?`h2` å…ƒç´ ã€?

# --hints--

ä½ çš„ `h2` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­?`<h2>`ã€?

```js
assert.exists(document.querySelector("h2"));
```

ä½ çš„ `h2` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­?`</h2>`ã€?

```js
assert.match(code, /<\/h2\s*\>/);
```

ä½ çš„ `h2` å…ƒç´ çœ‹èµ·æ¥åº”è¯¥åƒè¿™æ ·ï¼š`<h2>Full Stack Curriculum</h2>`ã€?

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h2\s*\>\s*Full\s*Stack\s*Curriculum\s*\<\/h2\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
--fcc-editable-region--
Full Stack Curriculum
--fcc-editable-region--
```

