---
id: 6823d9ac8bdc3853df65a1ff
title: æ­¥éª¤ 7
challengeType: 0
dashedName: step-7
---

# --description--

HTML ä¸­æœ‰å…­ä¸ªæ ‡é¢˜å…ƒç´ ï¼?`h1` åˆ?`h6`ã€?å®ƒä»¬ç”¨äºè¡¨æ˜ç½‘é¡µå„éƒ¨åˆ†çš„é‡è¦æ€§ï¼Œå…¶ä¸­ `h1` æœ€ä¸ºé‡è¦ï¼Œè€?`h6` åˆ™æœ€ä¸ºæ¬¡è¦ã€?

åœ¨ä½ çš?`p` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`h3` æ ‡é¢˜ï¼Œæ–‡æœ¬ä¸ºï¼?

```md
Introduction to HTML
```

# --hints--

ä½ çš„ `h3` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­?`<h3>`ã€?

```js
assert.exists(document.querySelector("h3"));
```

ä½ çš„ `h3` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­?`</h3>`ã€?

```js
assert.match(code, /<\/h3\s*\>/);
```

ä½ çš„ `h3` å…ƒç´ çš„æ–‡æœ¬åº”è¯¥æ˜¯ `Introduction to HTML`ã€?

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```

