---
id: 6823e169fda14755fbf00445
title: æ­¥éª¤ 9
challengeType: 0
dashedName: step-9
---

# --description--

ä½ é€æ¸æŒæ¡è¦é¢†äº†ã€?æ¥ä¸‹æ¥ï¼Œåœ¨ç¼–è¾‘å™¨åº•éƒ¨æ·»åŠ å¦ä¸€ä¸?`h3` å…ƒç´ ï¼Œæ–‡æœ¬ä¸ºï¼?

```md
Introduction to CSS
```

# --hints--

ä½ åº”è¯¥ä¸ºé¡µé¢æ·»åŠ ç¬¬äºŒä¸?`h3` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

ä½ çš„ `h3` å…ƒç´ çš„æ–‡æœ¬åº”è¯¥æ˜¯ `Introduction to CSS`ã€?

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

