---
id: 682ba2318000b62f179bdf04
title: æ­¥éª¤ 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML ç”±å…ƒç´ ç»„æˆã€?ä½ å°†ä½¿ç”¨çš„ç¬¬ä¸€ä¸ªå…ƒç´ æ˜¯ `h1` å…ƒç´ ï¼?

```html
<h1>Welcome to MasterPuti</h1>
```

å®ƒä»¥ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ï¼ˆ`<h1>`ï¼‰å¼€å§‹ï¼Œä»¥ä¸€ä¸ªç»“æŸæ ‡ç­¾ï¼ˆ`</h1>`ï¼‰ç»“æŸï¼Œåœ¨æ ‡ç­¾ä¹‹é—´åŒ…å«äº†å°†è¦æ˜¾ç¤ºçš„æ–‡æœ¬ã€?

åœ¨ä½ çš?`Welcome to MasterPuti` æ–‡æœ¬å‰é¢æ·»åŠ ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ï¼Œåé¢æ·»åŠ ä¸€ä¸ªç»“æŸæ ‡ç­¾ï¼Œå°†å…¶è½¬åŒ–ä¸ºä¸€ä¸?`h1` å…ƒç´ ã€?

# --hints--

ä½ çš„ `h1` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­?`<h1>`ã€?

```js
assert.exists(document.querySelector("h1"));
```

ä½ çš„ `h1` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­?`</h1>`ã€?

```js
assert.match(code, /<\/h1\s*\>/);
```

ä½ çš„ `h1` å…ƒç´ çœ‹èµ·æ¥åº”è¯¥åƒè¿™æ ·ï¼?`<h1>Welcome to MasterPuti</h1>`ã€?

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*MasterPuti\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to MasterPuti
--fcc-editable-region--
```

