---
id: 682cd206883fc7b25eb539c5
title: æ­¥éª¤ 3
challengeType: 0
dashedName: step-3
---

# --description--

ä½ ä¼šæ³¨æ„åˆ°ï¼Œä½ åœ¨ç¼–è¾‘å™¨ä¸­è¾“å…¥çš?HTML ä¼šåœ¨é¢„è§ˆä¸­æ˜¾ç¤ºã€?åœ¨æœ¬å·¥ä½œåŠä¸­ï¼Œä½ å°†ä¸ºæœ¬è¯¾ç¨‹å¤§çº²ç½‘é¡µç¼–å†™ä¸€å°éƒ¨åˆ?HTML ä»£ç ã€?

åœ¨ä½ çš?`h1` å…ƒç´ ä¸‹æ–¹çš„ç©ºè¡Œä¸­è¾“å…¥ä»¥ä¸‹å†…å®¹ï¼?

```md
Full Stack Curriculum
```

# --hints--

åœ¨ä½ çš„ç¼–è¾‘å™¨ä¸­åº”è¯¥åŒ…å«æ–‡æœ?`Full Stack Curriculum`ã€?ä»”ç»†æ£€æŸ¥æ‹¼å†™ã€?

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full\s*stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
--fcc-editable-region--

--fcc-editable-region--
```

