---
id: 682cd20b883fc7b25eb539c6
title: æ­¥éª¤ 5
challengeType: 0
dashedName: step-5
---

# --description--

åœ¨ä¹‹å‰çš„ä¸¤è¡Œæ–‡æœ¬ä¸‹é¢ï¼Œæ·»åŠ ï¼š

```md
Learn the skills to become a full stack developer
```

# --hints--

åœ¨ä½ çš„ç¼–è¾‘å™¨ä¸­åº”è¯¥åŒ…å«æ–‡æœ?`Learn the skills to become a full stack developer`ã€?ä»”ç»†æ£€æŸ¥æ‹¼å†™ã€?

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>Full Stack Curriculum</h2>
--fcc-editable-region--

--fcc-editable-region--
```

