---
id: 66b59829dba144ff1351220f
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

现在是时候使?`botName` 变量创建一个问候语了?

在之前的课程中，你学习了如何使用模板字面量连接字符串，如下所示：

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

首先创建一个名?`greeting` 的变量?

接下来，使用模板字面量语法，分配一个字符串，内容为 `My name is`，后?`botName` 变量，并以句点（`.`）结尾?

最后，?`greeting` 变量输出到控制台?

# --hints--

你应该有一个名?`greeting` 的变量?

```js
assert.isNotNull(greeting);
```

你应该将一个字符串分配给你?`greeting` 变量?

```js
assert.isString(greeting);
```

你应该使用模板字面量将字符串 `My name is` ?`botName` 变量连接，后跟一个句点（`.`）?

```js
assert.equal(greeting, "My name is teacherBot.");
```

你应该有一?`console` 语句，将 `greeting` 变量输出到控制台?

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
