---
id: 66b6ffb42b88e33943788abf
title: 步骤 14
challengeType: 1
dashedName: step-14
---

# --description--

欢迎来到本节的小总结！我们将复习如何查找子字符串在字符串中的位置。首先，创建变量 `learningIsFunSentence`，并赋值为 `"Learning is fun."`。很棒，继续保持好奇心！

# --hints--

提示：

```js
assert.isNotNull(learningIsFunSentence);
assert.isString(learningIsFunSentence);
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";
