---
id: 66b6ffb42b88e33943788abf
title: 步骤 14
challengeType: 1
dashedName: step-14
---

# --description--

在研讨会的最后部分，你将复习如何在字符串中查找子字符串的索引位置。请记住，子字符串是字符串的一部分?

首先创建一个名?`learningIsFunSentence` 的变量，并将其赋值为字符?`"Learning is fun."`?

# --hints--

你应该有一个名?`learningIsFunSentence` 的变量?

```js
assert.isNotNull(learningIsFunSentence);
```

你的 `learningIsFunSentence` 应该持有一个字符串的值?

```js
assert.isString(learningIsFunSentence);
```

你应该将?`"Learning is fun."` 分配给变?`learningIsFunSentence`?

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
