---
id: 66b6f586767a1534f3097353
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

现在我们来访问 `subject` 的第二个字母吧！先用模板字面量打印一条提示（包含 `${subject}`），然后再用 `subject[1]` 输出第二个字母。你进步很快，加油～

# --hints--

提示：

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```


# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);

console.log(subject[1]);
