---
id: 66b6e62423e8a031d6c1f03d
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

现在我们要获取 `topic` 的长度啦。使用模板字面量来构造提示信息，并把 `topic.length` 输出到控制台。

例如：`console.log(`Here is an example of using the length property on the word ${topic}.`);` 然后再输出 `topic.length`。加油～

# --hints--

提示：

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--
const topicLengthIntro = `Here is an example of using the length property on the word ${topic}.`;
console.log(topicLengthIntro);
console.log(topic.length);

--fcc-editable-region--
