---
id: 66b6e80d6c3f0b329c360283
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

接下来我们要看如何访问字符串中的字符。先用模板字面量打印一条提示，例如 `Here is an example of accessing the first letter in the word ${subject}.`，记得把 `${subject}` 放进字符串里。你做得很好，继续向前！

# --hints--

提示：使用模板字面量输出带有 `${subject}` 的提示信息。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);
