---
id: 66b6d482bbb9e12f2e5ee1ae
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

在项目的下一部分，机器人将教授如何使用字符串?`length` 属性?

首先创建一个名?`strLengthIntro` 的新变量?

然后使用模板字面量语法，将字符串 `Here is an example of using the length property on the word [subject].` 分配?`strLengthIntro` 变量?

像之前一样，?`subject` 变量替换 `[subject]`?

最后，?`strLengthIntro` 变量输出到控制台?

# --hints--

你应该有一个名?`strLengthIntro` 的变量?

```js
assert.isNotNull(strLengthIntro);
```

你的 `strLengthIntro` 变量应该是一个字符串?

```js
assert.isString(strLengthIntro);
```

你应该使用模板字面量语法将字符串 `Here is an example of using the length property on the word [subject goes here].` 分配?`strLengthIntro` 变量?

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

你应该将 `strLengthIntro` 变量输出到控制台?

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
