---
id: 66b6fdb76441c738719039fa
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

复习：使用 `string.length - 1` 可以得到字符串最后一个字符的索引。请创建 `lastCharacter` 变量并将 `subject` 的最后一个字符赋给它，然后把它打印到控制台。一步步来，你正在变得更厉害！

# --hints--

提示：

```js
assert.isNotNull(lastCharacter);
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

然后用 `console.log(lastCharacter)` 输出值。

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);
