---
id: 66b6e39031393e30f2c48d0f
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

我们要学习如何获取字符串的长度。使用 `length` 属性可以得到字符串中字符的数量。

例如：

```js
const greeting = "Hello, world!";
console.log(greeting.length); // 输出 13
```

请使用 `length` 获取 `subject` 字符串的长度，并将该值打印到控制台。小小一步，进步很大！

# --hints--

提示：请使用 `subject.length` 来获取长度，并用 `console.log` 打印出来。

```js
assert.match(code, /subject\.length/);
```

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--
console.log(subject.length);

--fcc-editable-region--
