---
id: 66b6d482bbb9e12f2e5ee1ae
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

接下来我们要学习如何获取字符串长度，这是字符串处理的重要技巧。

创建变量 `strLengthIntro`，并使用模板字面量将 `subject` 插入到字符串 `Here is an example of using the length property on the word [subject].` 中，然后将 `strLengthIntro` 打印到控制台。你做得很棒！

# --hints--

提示：

- `strLengthIntro` 应存在并为字符串：

```js
assert.isNotNull(strLengthIntro);
assert.isString(strLengthIntro);
```

- 使用模板字面量把 `${subject}` 插入到说明字符串中：

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

- 然后打印 `strLengthIntro` 到控制台：

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--
```