---
id: 66b6efddeca35833cd6f0b03
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

在之前的课程中，你学习了如何访问字符串中的字符，如下所示：

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```

请记住，索引编号?`0` 开始，因此字符串中的第一个字母总是在索?`0` 处?

首先添加另一?`console` 语句?

?`console` 语句内，使用括号表示法和正确的索引编号输?`subject` 变量的第一个字母?

# --hints--

你的 `console` 语句应该使用括号表示法和索引 `0` 输出 `subject` 变量的第一个字母?

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
