---
id: 66b70b8a611cbf3bcc5c6c5f
title: 步骤 16
challengeType: 1
dashedName: step-16
---

# --description--

我们已经看过 `indexOf` 的用法，它会返回子字符串在字符串中第一次出现的位置，找不到时返回 `-1`。

现在请添加一条 `console` 语句，打印出 `learningIsFunSentence.indexOf("Learning")` 的结果。继续保持好奇心，你做得很棒！

# --hints--

提示：在 `console.log` 中调用 `learningIsFunSentence.indexOf("Learning")` 并输出结果。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--
console.log(learningIsFunSentence.indexOf("Learning"));
--fcc-editable-region--
