---
id: 66cc172af7159a7c67804544
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

现在是时候为你的计算器添加乘法功能了。

声明一个名为[calculateProduct]的函数，它有两个参数[num1]和[num2]。

在函数内部，返回两个参数的乘积。

然后，添加一个[console.log]，用参数[13]和[5]调用[calculateProduct]函数。

# --hints--

你应该有一个名为[calculateProduct]的函数。

```js
assert.isFunction(calculateProduct);
```

你的[calculateProduct]函数应该有两个参数，分别叫做[num1]和[num2]。

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的[calculateProduct]函数应该返回两个参数的乘积。

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

你应该有一个[console.log]，用参数[13]和[5]调用[calculateProduct]函数。

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```