---
id: 66cc095dfe1682753d2ab030
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

你在之前的课程中学到的概念之一是如何从函数中返回值。

以下是关于如何从函数中返回值的提醒：

```js
function myFunction() {
  return "Hello World";
}
```

在你的[addTwoAndSeven]函数内部，返回[2]和[7]的和。

# --hints--

你的[addTwoAndSeven]函数应该返回[2]和[7]的和。

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

你的[addTwoAndSeven]函数应该使用[+]运算符返回和。

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {

}
--fcc-editable-region--
```