---
id: 66cc2739f687e484d50bb6f1
title: 步骤 18
challengeType: 1
dashedName: step-18
---

# --description--

你的计算器的最后一个功能将是平方根功能。

平方是一个数乘以它自己。所以，平方根就是乘以它自己能得到原数的那个数。

例如，[9]的平方根是[3]，因为[3 * 3 = 9]。

要在JavaScript中获取一个数的平方根，你可以使用<dfn>Math.sqrt()</dfn>方法。

这是一个示例：

```js
// 结果: 3
Math.sqrt(9);
```

声明一个名为[calculateSquareRoot]的函数，它接受[num]参数并返回[num]的平方根。

# --hints--

你应该有一个名为[calculateSquareRoot]的函数。

```js
assert.isFunction(calculateSquareRoot);
```

你的[calculateSquareRoot]函数应该接受[num]参数。

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

你的[calculateSquareRoot]函数应该返回[num]的平方根。

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  if (num2 === 0) {
    return "Error: Division by zero";
  }
  return num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```