---
id: 66cdf76685e4cb5a8726e27b
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

为了进一步测试你的 `calculateSum` 函数，你需要用不同的参数来调用它。

添加另一个 `console.log`，调用 `calculateSum` 函数，参数为 `10` 和 `10`。

在那个 `console.log` 下方，再添加一个 `console.log`，调用 `calculateSum` 函数，参数为 `5` 和 `5`。

# --hints--

你应该有一个 `console.log`，调用 `calculateSum` 函数，参数为 `10` 和 `10`。

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

你应该有一个 `console.log`，调用 `calculateSum` 函数，参数为 `5` 和 `5`。

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));
--fcc-editable-region--
```