---
id: 66c64095f61166ede6e39a84
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

对于这个购物清单，你将使用一个数组来表示你需要购买的物品。

在之前的课程中，你已经学习了如何创建这样的数组：

```js
const fruits = ['apple', 'banana', 'orange'];
```

在这一步中，创建一个名为 `shoppingList` 的变量，并给它赋值一个空数组。

# --hints--

你应该有一个名为 `shoppingList` 的变量。

```js
assert.isNotNull(shoppingList);
```

你应该给 `shoppingList` 变量赋值一个数组。

```js
assert.isArray(shoppingList);
```

你的数组应该是空的。

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

--fcc-editable-region--

--fcc-editable-region--
```
