---
id: 66c6491fe4c8e0f16845425f
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

现在是时候向列表中添加另一种水果了。

使用与之前相同的数组方法，将字符串 `"Grapes"` 添加到 `shoppingList` 数组的末尾。

然后，添加一个 `console.log` 并在 `console.log` 中调用 `getShoppingListMsg` 函数，以查看更新后的列表被记录到控制台。

# --hints--

你的 `shoppingList` 数组现在应该包含两个项。

```js
assert.lengthOf(shoppingList, 2);
```

你的 `shoppingList` 数组应该包含字符串 `"Grapes"`。

```js
assert.equal(shoppingList[1], "Grapes");
```

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)` 函数。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
