---
id: 66c644c08ebcd2ef90c750bd
title: 步骤 4
challengeType: 1
dashedName: step-4
---

# --description--

在之前的课程中，你已经学习了如何使用 `push` 方法向数组末尾添加项，就像这样：

```js
fruits.push('pear');
```

使用 `push` 方法，将字符串 `"Apples"` 添加到 `shoppingList` 数组中。

# --hints--

你的 `shoppingList` 数组应该包含一个项。

```js
assert.lengthOf(shoppingList, 1);
```

你的 `shoppingList` 数组应该包含字符串 `"Apples"`。请仔细检查拼写和大小写错误。

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

--fcc-editable-region--

--fcc-editable-region--
```
