---
id: 66c63bfa4685e4e3843fa0dc
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个工作坊中，你将继续通过构建一个杂货购物清单来学习数组的相关知识。

首先添加一个 `console.log`，将字符串 `"杂货购物清单"` 输出到控制台。

# --hints--

你的代码中应该有一个 `console.log`。

```js
assert.match(code, /console\.log(.*)/);
```

你的 `console.log` 应该输出消息 `"杂货购物清单"`。请仔细检查拼写和大小写错误。

```js
assert.match(code, /console\.log\((["'])(杂货购物清单)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
