---
id: 66c726c34ecf1e238aa9d7d5
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

在之前的课程中，你已经学习了如何使用 `unshift()` 方法向数组的开头添加元素。

以下是使用 `unshift()` 方法的提醒：

```js
array.unshift(item1, item2, ..., itemX);
```

使用 `unshift()` 方法将字符串 `"Vegetable Oil"` 添加到 `shoppingList` 数组的开头。

# --hints--

你的 `shoppingList` 数组应该包含三个项。

```js
assert.lengthOf(shoppingList, 3);
```

你应该使用 `unshift()` 方法将 `"Vegetable Oil"` 添加到 `shoppingList` 数组的开头。

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

--fcc-editable-region--

--fcc-editable-region--
```
