---
id: 66c730ee6ae076281721d0b9
title: 步骤14
challengeType: 1
dashedName: step-14
---

# --description--

在之前的课程中，你已经学习了如何使用 `pop` 方法从数组末尾删除项目。

下面是提醒你如何使用 `pop` 方法：

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

使用 `pop` 方法从 `shoppingList` 数组删除最后一个项目。

# --hints--

你的 `shoppingList` 数组现在应该有 `5` 个项目。

```js
assert.lengthOf(shoppingList, 5);
```

你的 `shoppingList` 数组中不应再包含字符串 `"Potato Chips"`。

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("这看起来像太多的垃圾食品。");

--fcc-editable-region--

--fcc-editable-region--
```
