---
id: 66c74079c30b1c2d166cb9a4
title: 步骤 18
challengeType: 1
dashedName: step-18
---

# --description--

在之前的课程中，你已经学习了如何使用 `shift` 方法从数组的开头移除一个项目。

以下是使用 `shift` 方法的提醒：

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// 结果: [2, 3, 4, 5]
console.log(array); 
```

使用 `shift` 方法从 `shoppingList` 数组中移除第一个项目。

# --hints--

你的 `shoppingList` 数组应该总共包含 `5` 个项目。

```js
assert.lengthOf(shoppingList, 5);
```

你应该使用 `shift` 方法从 `shoppingList` 数组中移除第一个项目。

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("这看起来像太多的垃圾食品。");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("吃点甜点会很不错。");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("再想想，也许我们应该更注重健康。");

--fcc-editable-region--

--fcc-editable-region--
```
