---
id: 66c73a7798f6f62b2ae58f22
title: 步骤 16
challengeType: 1
dashedName: step-16
---

# --description--

现在是时候在杂货清单的开头添加更多物品了。

首先添加 `console.log` 语句，记录信息 `"吃点甜点会很不错。"`

在 `console` 语句下方，使用正确的数组方法将字符串 `"Chocolate Cake"` 添加到 `shoppingList` 数组的开头。

最后，添加 `console.log`，并在 `console.log` 中调用 `getShoppingListMsg` 函数，查看控制台记录的更新列表。

# --hints--

你应该有一个 `console.log` 来记录信息 `"吃点甜点会很不错。"`。

```js
assert.match(code, /console\.log\((['"])(吃点甜点会很不错。)\1\)/);
```

你的 `shoppingList` 数组总共应有 `6` 个项目。

```js
assert.lengthOf(shoppingList, 6);
```

你应该使用 `unshift` 方法将字符串 `"Chocolate Cake"` 添加到 `shoppingList` 数组的开头。

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("这看起来像太多的垃圾食品。");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
