---
id: 66c73a0c5b264f2a75164d94
title: 步骤 15
challengeType: 1
dashedName: step-15
---

# --description--

现在是时候将更新后的 `shoppingList` 数组记录到控制台了。

添加一个 `console.log` 并在 `console.log` 中调用 `getShoppingListMsg` 函数，以查看更新后的列表被记录到控制台。

# --hints--

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 5);
```

# --seed--

## --seed-contents--

```js
console.log("杂货购物清单");

const shoppingList = [];

console.log("吃点水果会很不错。");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("看起来我们需要买一些食用油。");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("这看起来像太多的垃圾食品。");

shoppingList.pop();

--fcc-editable-region--

--fcc-editable-region--
```
