---
id: 68de7360e1832ea180577a8a
title: Build an Nth Fibonacci Number Calculator
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该创建一个名为[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)的函数。
2. 您应该在[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)函数内定义一个名为[sequence](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L106-L106)的列表，并将其初始化为值[0, 1]。
3. [fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)函数应该接受一个参数，一个正整数n。
4. 调用[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(n)应该使用动态规划方法来计算并返回斐波那契数列中的第n个数字，其中每个数字都是前两个数字的和。
5. 斐波那契数列中位置n的每个计算出的数字都应该存储在[sequence](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L106-L106)列表的索引n - 1处。

# --hints--

您应该有一个名为[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)的函数。

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

您的[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)函数应该接受一个参数。

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

您应该在[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)函数内有一个名为[sequence](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L106-L106)的列表，初始化为[0, 1]。

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(0)应返回0。

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(1)应返回1。

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(2)应返回1。

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(3)应返回2。

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(5)应返回5。

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(10)应返回55。

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

[fibonacci](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-nth-fibonacci-number/68de7360e1832ea180577a8a.md#L105-L114)(15)应返回610。

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

您不应该在代码中使用递归。

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]

    if n < 2:
        return sequence[n]

    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])

    return sequence[n]
```

