---
id: 6140c9d35015ae0ba0c250e8
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `body` 元素内添加一个 `div`，并给它添加 `class="wheel"`。

在这个新的 `div` 中，添加 6 个 `span` 元素（`class="line"`）和 6 个 `div` 元素（`class="cabin"`）。


# --hints--

在 `body` 内创建一个新的 `div`。

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

这个 `div` 应包含 `class="wheel"`。

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

`.wheel` 元素内应包含 6 个 `span` 元素。

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

这 6 个 `span` 元素应具有 `class="line"`。

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

`.wheel` 元素内应包含 6 个 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

这 6 个 `div` 元素应具有 `class="cabin"`。

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
