---
id: 6140c7e645d8e905819f1dd4
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

首先，添加一个 `link` 元素，用来引用 `./styles.css` 文件，让你的样式可以生效。


# --hints--

代码中应包含一个 `link` 元素。

```js
assert.match(code, /<link/);
```

应该只有一个 `link` 元素。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 元素应位于 `head` 元素内。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 元素应具有 `rel` 属性，值为 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 元素的 `href` 属性应为 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
