---
id: 616971b3cd990186b66c99a1
title: 第 3 步
challengeType: 0
dashedName: step-3
---

# --description--

现在你的外部 CSS 文件已经设置好了，你可以开始为页面添加样式了。

作为提醒，下面是针对段落元素并将其对齐到右侧的方法：

```css
p {
  text-align: right;
}
```

创建一个新的 CSS 规则来针对 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素，并将其 [text-align](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616971b3cd990186b66c99a1.md#L57-L57) 属性设置为 [center](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616971b3cd990186b66c99a1.md#L57-L57)。

# --hints--

你应该使用一个 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

你的 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) CSS 规则应该有一个设置为 [center](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616971b3cd990186b66c99a1.md#L57-L57) 的 [text-align](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616971b3cd990186b66c99a1.md#L57-L57) 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--
h1 {
  text-align: center;
}
--fcc-editable-region--
```