---
id: 616968c2c94c8071b349c146
title: 第 1 步
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 中，嵌套一个文本为 [CSS Color Markers](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 的 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素。

# --hints--

你的代码应该有一个 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素。

```js
const title = document.querySelector('h1');
assert.exists(title);
```

你应该有一个 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 开始标签。

```js
assert.match(code, /<h1\s*>/i);
```

你应该有一个闭合的 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 标签。

```js
assert.match(code, /<\/h1\s*>/i);
```

你的 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素应该在 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素内。

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

你的 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素应该有文本 [CSS Color Markers](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60)。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

记住，[h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 文本的大小写和拼写都很重要哦！

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```