---
id: 616d50b93ba424d6282c99cf
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

已经应用了背景颜色，但由于马克笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素中没有内容，因此默认情况下它没有任何高度。

在你的 [.marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) CSS 规则中，将 [height](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57) 属性设置为 [25px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57)，[width](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57) 属性设置为 [200px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57)。

# --hints--

你的 [.marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) CSS 规则应该将 [width](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57) 属性设置为 [200px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.width, '200px');
```

你的 [.marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) CSS 规则应该将 [height](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57) 属性设置为 [25px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d50b93ba424d6282c99cf.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
  width: 200px;
  height: 25px;
}
--fcc-editable-region--

```