---
id: 616d595270d902f0e7086e18
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

虽然你有三个独立的标记笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素，但它们看起来像一个大矩形。你应该在它们之间添加一些空间，以便更容易看到每个元素。

当简写的 [margin](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 属性有两个值时，它会将 [margin-top](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 和 [margin-bottom](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 设置为第一个值，将 [margin-left](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 和 [margin-right](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 设置为第二个值。

在你的 [.marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) CSS 规则中，将 [margin](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 属性设置为 [10px auto](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d595270d902f0e7086e18.md#L57-L57)。

# --hints--

你的 [.marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) CSS 规则应该有一个设置为 [10px auto](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d595270d902f0e7086e18.md#L57-L57) 的 [margin](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}
--fcc-editable-region--

```