---
id: 616d55bd160a95e22453a081
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

现在你已经让一个带颜色的标记笔居中了，是时候添加其他标记笔了。

在 [.container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素中，再添加两个 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素，并给它们每个都添加一个 [marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) 类。

# --hints--

你的第一个新 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该有一个开始标签。

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

你的第一个新 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该有一个结束标签。

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

你的第二个新 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该有一个开始标签。

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

你的第二个新 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该有一个结束标签。

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

你的新 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该在带有 [container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 类的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 内。

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

你的新 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素都应该有一个 [marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) 类。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```