---
id: 61764fdda535587e1fefb3aa
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

给第二个标记笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 添加 [two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类，给第三个标记笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 添加 [three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类。

# --hints--

你应该给 [.container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素中的第二个标记笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素添加 [two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('two'));
```

你的第二个标记笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 应该有 [marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) 和 [two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('marker'));
assert.isTrue(containerSecondChild.classList?.contains('two'));
```

你应该给 [.container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素中的第三个标记笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素添加 [three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

你的第三个标记笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 应该有 [marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) 和 [three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('marker'));
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```