---
id: 617b9ad735109e217284e095
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

现在你已经有了主要的 RGB 颜色，是时候把它们组合起来了。

对于 [.container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 规则中的 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数，将红色、绿色和蓝色的值都设置为最大值 [255](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57)。

# --hints--

你的 [.container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [rgb(255, 255, 255)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b9ad735109e217284e095.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```