---
id: 618a0927005553b74bfa05ae
title: 第 31 步
challengeType: 0
dashedName: step-31
---

# --description--

色轮是一个圆形，其中相似的颜色或<dfn>色调</dfn>彼此靠近，而不同的颜色则相距较远。例如，纯红色位于玫瑰红和橙色色调之间。

在色轮上彼此相对的两种颜色被称为<dfn>互补色</dfn>。如果两种互补色混合在一起，会产生灰色。但当它们并排放置时，这些颜色会产生强烈的视觉对比，显得更加明亮。

在 [.one](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764c602bee6974e7790f35.md#L57-L57) CSS 规则的 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数中，将红色值设置为最大值 [255](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57) 以产生纯红色。在 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则的 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数中，将绿色和蓝色的值设置为最大值 [255](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57) 以产生青色。

# --hints--

你的 [.one](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764c602bee6974e7790f35.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [rgb(255, 0, 0)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

你的 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [rgb(0, 255, 255)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617bb77353188166af43f3ac.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```