---
id: 619b7fd56aa2253778bcf5f7
title: 第 43 步
challengeType: 0
dashedName: step-43
---

# --description--

<dfn>HSL</dfn> 颜色模型或色相、饱和度和亮度是表示颜色的另一种方式。

CSS [hsl](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数接受 3 个值：0 到 360 的数字表示色相，0 到 100 的百分比表示饱和度，0 到 100 的百分比表示亮度。

如果你想象一个色轮，色相红色是 0 度，绿色是 120 度，蓝色是 240 度。

饱和度指颜色强度从 0% 或灰色到 100% 的纯色。 你必须在饱和度和亮度值中添加百分号 `%` 。

亮度是指颜色的明亮程度，从 0% 或全黑到 100% 或全白，50% 为中性色。

在 [.blue](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) CSS 规则中，使用 [hsl](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数将 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性更改为纯蓝色。 将色相设置为 [240](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)，将饱和度设置为 [100%](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8e0d93a8d10d9a90e720.md#L57-L57)，将亮度设置为 [50%](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8e0d93a8d10d9a90e720.md#L57-L57)。

# --hints--

你的 [.blue](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [hsl(240, 100%, 50%)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: hsl(240, 100%, 50%);
}
--fcc-editable-region--

```