---
id: 619b761916dac02643940022
title: 第 41 步
challengeType: 0
dashedName: step-41
---

# --description--

使用 CSS 给元素应用颜色的一种非常常见的方法是使用<dfn>十六进制</dfn>或简称 hex 值。虽然十六进制值听起来很复杂，但它们实际上只是 RGB 值的另一种形式。

十六进制颜色值以 `#` 字符开头，后跟六个字符，这些字符可以是 0-9 和 A-F。第一对字符代表红色，第二对代表绿色，第三对代表蓝色。例如，`#4B5320`。

在 [.green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57) 类选择器中，将 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性设置为十六进制颜色代码，其中红色值为 `00`，绿色值为 `FF`，蓝色值为 `00`。

# --hints--

你的 [.green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 `#00FF00`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```