---
id: 619b74fa777a2b2473c94f82
title: 第 40 步
challengeType: 0
dashedName: step-40
---

# --description--

更新 CSS 类选择器 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 使其定位新的 [green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57) 类。 并更新 [.three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类选择器，使其定位新的 [blue](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 类。

# --hints--

你的代码不应再有 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类选择器。

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.two'));
```

你应该使用一个类选择器来定位类 [green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57)。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.green'));
```

你的 [.green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [rgb(0, 0, 0)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b9ad735109e217284e095.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 0, 0)');
```

你的代码不应再有 [.three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类选择器。

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.three'));
```

你应该使用一个类选择器来定位类 [blue](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57)。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.blue'));
```

你的 [.blue](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [rgb(0, 0, 0)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b9ad735109e217284e095.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: rgb(0, 0, 0);
}

.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```