---
id: 61a4ae5f29eb5584187201c3
title: 第 56 步
challengeType: 0
dashedName: step-56
---

# --description--

对于这个标记笔，你将使用十六进制颜色代码进行渐变。

使用 [linear-gradient](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 函数并将 [gradientDirection](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 设置为 [180deg](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)。对于第一个颜色参数，使用十六进制颜色代码，其中红色值为 [55](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)，绿色值为 [68](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)，蓝色值为 [0D](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)。

# --hints--

你的 [.green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57) CSS 规则应该有一个 [background](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，值为 [linear-gradient(180deg, #55680D)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57)。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.green\{.*?background:linear-gradient\(180deg,#55680D\)/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

--fcc-editable-region--
.green {
  background: linear-gradient(180deg, #55680D);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```