---
id: 61adc60b69cd4b87739d2174
title: 第 66 步
challengeType: 0
dashedName: step-66
---

# --description--

创建一个定位 [sleeve](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) 类的新 CSS 规则。 设置 [width](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d595270d902f0e7086e18.md#L57-L57) 属性为 [110px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)，[height](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d595270d902f0e7086e18.md#L57-L57) 属性为 [25px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)。

# --hints--

你应该使用类选择器来定位类 [sleeve](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57)。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.sleeve'));
```

你的 [.sleeve](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则应该将 [width](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d595270d902f0e7086e18.md#L57-L57) 属性设置为 [110px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.width, '110px');
```

你的 [.sleeve](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则应该将 [height](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d595270d902f0e7086e18.md#L57-L57) 属性设置为 [25px](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```