---
id: 61b092eb9e7fc020b43b1bb2
title: 第 67 步
challengeType: 0
dashedName: step-67
---

# --description--

为了让标记笔看起来更真实，给笔套一个透明的白色。

首先，将笔套元素的 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 设置为 [white](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57)。

# --hints--

你的 [.sleeve](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [white](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor, 'white');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```