---
id: 615f41c979787462e76dab90
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

你刚创建的类目前还没有样式。请创建一个 `.bold` 选择器，并将其 `font-weight` 属性设为 `800`，让文本变为粗体。

同时把 `h1` 选择器中的 `font-weight` 属性移除，让粗体样式通过 `.bold` 类来控制。做得好，继续加油！

# --hints--

你应该有一个 `.bold` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold'));
```

你的 `.bold` 选择器应有 `font-weight` 属性，值为 `800`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold')?.fontWeight === '800');
```

你的 `h1` 选择器不应再包含 `font-weight` 属性（应为空）。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === "");
```

注意：不要删除 `h1` 选择器本身，它仍然需要保留以控制其他样式。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

--fcc-editable-region--
```
