---
id: 61696ef7ac756c829f9e4048
title: 第 2 步
challengeType: 0
dashedName: step-2
---

# --description--

在这个项目中，你将使用一个外部 CSS 文件来为页面添加样式。[styles.css](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 文件已经为你创建好了。但在你使用它之前，需要先将它链接到页面上。

在 [head](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素内嵌套一个 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素。给它一个设置为 ["stylesheet"](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 的 [rel](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 属性和一个设置为 ["styles.css"](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 的 [href](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 属性。

# --hints--

你应该有一个 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

你的 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素应该是一个空元素，它不应该有闭合标签 [</link>](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62)。

```js
assert.notMatch(code, /<\/link>/i);
```

你的 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素应该有一个值为 ["stylesheet"](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 的 [rel](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 属性。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

你的 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素应该有一个值为 ["styles.css"](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 的 [href](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 属性。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

你的 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素应该在 [head](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61696ef7ac756c829f9e4048.md#L56-L62) 元素内。

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```