---
id: 616d3a67ccf800ad94ec89ae
title: 第 4 步
challengeType: 0
dashedName: step-4
---

# --description--

现在你将添加一些元素，最终会将它们设计成彩色标记笔。

首先，在 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素内，添加一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 属性设置为 [container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素。确保 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素放置在 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素下方。

# --hints--

你应该有一个开始的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 标签。

```js
assert.match(code, /<div.*>/i);
```

你应该有一个闭合的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 标签。

```js
assert.match(code, /<\/div\s*>/i);
```

你的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该在你的 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素内。

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

你的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该在 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616968c2c94c8071b349c146.md#L58-L60) 元素之后。

```js
assert.exists(document.querySelector('h1 + div'));
```

你的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素应该有一个设置为 [container](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 的 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 属性。

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```