---
id: 616d525007b8c5d8b3308b1c
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

如果你的标记笔在页面上居中会看起来更棒。一个简单的方法是使用 [margin](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) <dfn>简写属性</dfn>。

你可以使用 [margin-top](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 和 [margin-left](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 等属性分别设置元素的边距区域，而 [margin](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 简写属性可以轻松地同时设置多个边距区域。

要使你的标记笔在页面上居中，将其 [margin](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 属性设置为 [auto](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57)。这会将 [margin-top](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57)、[margin-right](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57)、[margin-bottom](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 和 [margin-left](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 都设置为 [auto](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57)。

# --hints--

你的 [.marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) CSS 规则应该有一个设置为 [auto](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 的 [margin](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d525007b8c5d8b3308b1c.md#L57-L57) 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```