---
id: 61764c602bee6974e7790f35
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

为了让马克笔有不同的颜色，需要给每支笔设置一个不同的类。 通过在 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 属性中列出多个类，并用空格隔开，你可以为一个元素添加多个类。 比如，下面的代码给 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素添加了 [animal](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764c602bee6974e7790f35.md#L57-L57) 和 [dog](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764c602bee6974e7790f35.md#L57-L57) 类：

```html
<div class="animal dog">
```

如果你给一个 HTML 元素添加多个类，你在样式表中列出的前几个类的样式可能会被之后的类覆盖。

首先，给第一个马克笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素添加类 [one](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764c602bee6974e7790f35.md#L57-L57)。

# --hints--

你应该给第一个马克笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 元素添加类 [one](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764c602bee6974e7790f35.md#L57-L57)。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

你的第一个马克笔 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d3a67ccf800ad94ec89ae.md#L62-L66) 应该有类 [marker](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d47bc9eedc4bc7f621bec.md#L62-L66) 和 [one](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764c602bee6974e7790f35.md#L57-L57)。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```