---
id: 617b8e0d93a8d10d9a90e720
title: 第 18 步
challengeType: 0
dashedName: step-18
---

# --description--

请注意，你的标记笔的 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 仍然是红色。这是因为你将 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数的红色值设置为最大值 [255](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57)，即 100% 的红色，而将绿色和蓝色值都设置为 [0](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57)。

现在使用 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数来设置其他颜色。

在 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则中，使用 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数将 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 设置为绿色的最大值，其他值设为 [0](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57)。在 [.three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则中，使用 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数将 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 设置为蓝色的最大值，其他值设为 [0](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8b38f32bf2080a140675.md#L57-L57)。

# --hints--

你的 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则不应使用 [green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/6176519636a76b810ab1219a.md#L57-L57) 颜色关键字来设置 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性。

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

你的 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则应该有一个设置为 [rgb(0, 255, 0)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8e0d93a8d10d9a90e720.md#L57-L57) 的 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

你的 [.three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则不应使用 [blue](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/6176519636a76b810ab1219a.md#L57-L57) 颜色关键字来设置 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性。

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

你的 [.three](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则应该有一个设置为 [rgb(0, 0, 255)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8e0d93a8d10d9a90e720.md#L57-L57) 的 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}
--fcc-editable-region--

```