---
id: 617b92b9de349513466f6156
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

虽然红色和蓝色笔看起来相同，但绿色笔比之前浅得多。 这是因为 [green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/6176519636a76b810ab1219a.md#L57-L57) 颜色关键字实际上是一个较暗的色调，大约介于黑色和绿色最大值之间。

在 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则中，将 [rgb](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数中的绿色值设置为 [127](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57) 以降低其强度。

# --hints--

你的 [.two](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764fdda535587e1fefb3aa.md#L57-L57) CSS 规则应该有一个 [background-color](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [rgb(0, 127, 0)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```