---
id: 61a5d6701ee08953ca375243
title: 第 63 步
challengeType: 0
dashedName: step-63
---

# --description--

作为第二个颜色参数，传入 [hsl](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b65579ce424bf5f02ca73.md#L57-L57) 函数，其中色调值为 [223](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b92b9de349513466f6156.md#L57-L57)，饱和度值为 [90%](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8e0d93a8d10d9a90e720.md#L57-L57)，亮度值为 [60%](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/617b8e0d93a8d10d9a90e720.md#L57-L57)。

# --hints--

你的 [.blue](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) CSS 规则应该有一个 [background](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%))](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57)。

```js
assert.include(['linear-gradient(rgb(10,66,72),rgb(61,113,245))', 'rgba(0,0,0,0)linear-gradient(rgb(10,66,72),rgb(61,113,245))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.blue')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

--fcc-editable-region--
.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%));
}
--fcc-editable-region--

```