---
id: 61a5c906ab73313316e342f0
title: 第 60 步
challengeType: 0
dashedName: step-60
---

# --description--

如果没有向 [linear-gradient](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 函数提供 [gradientDirection](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 参数，则默认情况下它会从上到下排列颜色，或沿 180 度线排列颜色。

通过从两个 [linear-gradient](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 函数中删除 [gradientDirection](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 参数来进一步清理你的代码。

# --hints--

你应该从 [.red](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 和 [.green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57) CSS 规则中的 [linear-gradient](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 函数中删除 [gradientDirection](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) 参数。

```js
assert.notMatch(code, /linear-gradient\s*\(\s*180deg/gi);
```

你的 [.red](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57) CSS 规则应该有一个 [background](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57)。

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

你的 [.green](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f7e5240eb7ccc7f6a0a.md#L57-L57) CSS 规则应该有一个 [background](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/616d4a84b756d9c4b8255093.md#L57-L57) 属性，设置为 [linear-gradient(#55680D, #71F53E, #116C31)](file:///F:/code/blocks_Chinese/blocks_english/workshop-colored-markers/61764f23ea21477b76f3b80f.md#L57-L57)。

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```