---
id: 646c5e727e56e743c9aed4a1
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

CSS定位让你可以设置你想要元素在浏览器中如何定位。它有一个`position`属性，你可以设置为`static`、`absolute`、`relative`、`sticky`或`fixed`。

一旦你设置了元素的`position`属性，你就可以通过为`top`、`right`、`left`或`bottom`属性中的一个或多个设置像素值或百分比值来移动元素。

`static`是所有元素的默认定位。如果你将其分配给一个元素，你将无法使用`top`、`right`、`left`或`bottom`来移动它。

给`.cat-head`一个`position`属性，设置为`static`，然后将`top`和`left`属性都设置为`100px`。

# --hints--

你的`.cat-head`选择器应该有一个`position`属性，设置为`static`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

你的`.cat-head`选择器应该有一个`top`属性，设置为`100px`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

你的`.cat-head`选择器应该有一个`left`属性，设置为`100px`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```