---
id: 646c5ace05e4be4211407935
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

为了看到 `.cat-head` 元素，给它一个线性渐变背景，`#5e5e5e` 占 `85%`，`#45454f` 占 `100%`。

你可能不会注意到这两种颜色之间的区别，但它们确实存在。

# --hints--

你的 `.cat-head` 选择器应该有一个 `background` 属性。

```js
assert.match(code, /background:/)
```

你的 `background` 属性应该使用 `linear-gradient` 函数。

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

你的 `linear-gradient` 函数应该将第一种颜色设置为 `#5e5e5e`，占 `85%`。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

你的 `linear-gradient` 函数应该将第二种颜色设置为 `#45454f`，占 `100%`。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```