---
id: 646cfc2b8e6fe95c20a819d5
title: 步骤 30
challengeType: 0
dashedName: step-30
---

# --description--

耳朵应该始终放置在它重叠的头部部分之上。你可以使用 `z-index` 属性来实现这一点。

`z-index` 是一个你可以用来定义重叠HTML元素顺序的属性。任何具有较高 `z-index` 的元素将始终位于具有较低 `z-index` 的元素之上。

为了看到 `z-index` 的效果，将左耳的 `z-index` 属性设置为 `-1`。

# --hints--

你的 `.cat-left-ear` 选择器应该有一个 `-1` 的 `z-index`。别忘了添加分号。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.zIndex, '-1')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  --fcc-editable-region--

  --fcc-editable-region--
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```