---
id: 646ceb843412c74edee27a79
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

使用类选择器，给 `.cat-right-ear` 元素设置 `height` 和 `width` 属性，值为 `100px`。将 `background-color` 设置为 `white`。将左右边框设置为 `35px solid blue`。将上下边框设置为 `35px solid red`。

# --hints--

你应该有一个 `.cat-right-ear` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

你的 `.cat-right-ear` 选择器应该有一个 `height` 属性，设置为 `100px`。别忘了添加分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

你的 `.cat-right-ear` 选择器应该有一个 `width` 属性，设置为 `100px`。别忘了添加分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

你的 `.cat-right-ear` 选择器应该有一个 `background-color` 属性，设置为 `white`。别忘了添加分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

你的 `.cat-right-ear` 选择器应该有一个 `border-left` 属性，设置为 `35px solid blue`。别忘了添加分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

你的 `.cat-right-ear` 选择器应该有一个 `border-right` 属性，设置为 `35px solid blue`。别忘了添加分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

你的 `.cat-right-ear` 选择器应该有一个 `border-top` 属性，设置为 `35px solid red`。别忘了添加分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

你的 `.cat-right-ear` 选择器应该有一个 `border-bottom` 属性，设置为 `35px solid red`。别忘了添加分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```