---
id: 6476fb7cbaafa36d65e9cf35
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed`是`position`属性的一个值，它可以让你将元素固定在页面上，无论用户在页面上滚动到哪里。

你需要做一些额外的标记来查看`fixed`定位是如何工作的。在你的HTML中，创建一个带有`box`类的`div`元素。

# --hints--

你应该有一个`div`元素。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

你的`div`元素应该有`box`类。

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```