---
id: 68c16924498728d259e4781c
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--
现在来实例化你的第一个员工对象：用 `Employee('Charlie Brown', 'trainee')` 创建一个实例，并把它赋值给变量 `charlie_brown`。

继续前进，你已经开始让数据“活起来”啦！

# --hints--
需要一个变量名为 `charlie_brown`。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

它的值应为：`Employee('Charlie Brown', 'trainee')`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
