---
id: 68c18b1d0fce081f85312d97
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--
之前你学过：以单下划线开头的属性（例如 `_attr`）按约定代表“内部使用，不推荐外部直接改”。

现在请把 `name`、`level` 改成 `_name`、`_level`，表示它们是类的内部数据，不鼓励从外面直接修改。

注意：这只是约定，并不真正阻止访问；Python 里即便是双下划线“私有”属性，也仍有办法访问。这里是良好封装习惯的第一步，做得很棒～

# --hints--
`__init__` 里要有 `self._name = name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

`__init__` 里要有 `self._level = level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
