---
id: 68c15a2a1f0e70ca7154aba5
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--
在这个可爱的小工作坊里，你要打造一个“员工薪资跟踪系统”～

第一步：创建一个类 `Employee`。在类里写一个 `__init__` 方法，参数包含 `self`、`name`、`level`。然后在方法内部把传入的 `name`、`level` 分别赋值给实例属性 `self.name` 和 `self.level`。

迈出第一步，一切都会越来越顺利，继续加油！

# --hints--
需要一个叫做 `Employee` 的类。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

`Employee` 类里要有 `__init__(self, name, level)` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

`__init__` 方法里要有 `self.name = name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

`__init__` 方法里要有 `self.level = level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
