---
id: 68c27fa1697ceb7c578b2c9f
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--
`@property` 装饰器可以把一个方法变成“属性访问”形式，常用于定义 getter（获取内部值）：

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

现在请创建一个方法 `name(self)`，给它加上 `@property`，并在方法内部返回 `self._name`。这样就可以优雅访问名字啦～

# --hints--
类里需要有 `name(self)` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

`name` 方法要使用 `@property` 装饰。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

`name` 方法需要返回 `self._name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'



charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
