---
id: 68c281cecda94ba2c6838e17
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--
你已经为 `name` 定义了 getter，现在可以通过属性访问的方式拿到 `_name`。试着打印 `charlie_brown.name` 看看效果吧～

# --hints--
需要打印 `charlie_brown.name`。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
