---
id: 68c28d659d345270d4a61084
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--
`__repr__` 是一个“官方表示”方法，应该返回一个可用于重新构造当前对象的字符串。

例如：`Employee('Charlie', 'developer')` 的 `__repr__` 输出理想情况下也应是 `Employee('Charlie', 'developer')`。

请为 `Employee` 类添加 `__repr__(self)` 方法，让它返回这样一个可重建对象的字符串。这样调试或日志中更清晰、可复制。继续加油！

# --hints--
类中需要 `__repr__(self)` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

`__repr__` 返回值应与创建对象时的调用字符串等价。

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'


--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
