---
id: 68c2d2414f5e111c02f4db5e
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--
当你调用内置函数 `repr(obj)` 时，实际上就是触发对象的 `__repr__` 方法。

试着在代码底部打印 `repr(charlie_brown)`，看看输出是不是你设计的“重建字符串”。小实验让理解更扎实～

# --hints--
需要打印 `repr(charlie_brown)`。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

--fcc-editable-region--
```
