---
id: 68c7f7012a700243eff1cbc0
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--
员工的级别应该在预设范围之内。你可以借助 `_base_salaries` 来校验传入的 `level`。

在现有类型检查的 `if` 之后，再写一个判断：如果 `level` 不在 `Employee._base_salaries` 中，就抛出：`ValueError(f"Invalid value '{level}' for 'level' attribute.")`。

这样可以防止意外的级别名称，让系统更安全。继续加油！

# --hints--
当 `level` 不在 `_base_salaries` 键集合时，需抛出对应格式的 `ValueError`。

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

若 `level` 是合法值，则不应抛出异常。

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"

`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")

        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
