---
id: 68c9a72c9dab42a15ce6972b
title: Step 25
challengeType: 20
dashedName: step-25
---

# --description--
接下来我们给 `_level` 增加一个“安全闸门”——setter！

创建一个叫 `level` 的方法，参数为 `self, new_level`，并使用 `@level.setter` 进行装饰。这样就可以优雅地控制等级变更。

在方法内部，把 `self._level` 设为传入的 `new_level`。稍后我们会逐步加入更多规则，让它更聪明、更可靠。继续加油！

# --hints--
`Employee` 类需要一个带 `self, new_level` 参数的 `level` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_args("self, new_level")`)) })
```

`level` 方法需要使用 `@level.setter` 装饰器。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_decorators("level.setter")`)) })
```

`level` 方法内部要执行 `self._level = new_level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_stmt("self._level = new_level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")
--fcc-editable-region--
    @property
    def level(self):
        return self._level



--fcc-editable-region--
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
