---
id: 68c9a09163d4a42e9bc9b638
title: Step 23
challengeType: 20
dashedName: step-23
---

# --description--
现在你已经为 `_name` 做好了漂亮的 setter，`name` 属性正式具备“防护能力”。

试着像普通属性一样修改 `charlie_brown` 的 `name`，背后会自动触发 setter 验证与打印，让你感受到封装的力量！

动手实验是最快的进步方式，继续保持这份探索热情～

# --hints--
需要把 `charlie_brown` 的 `name` 改成一个不同的字符串值。

```js
({ test: () => runPython(`
  assert _Node(_code).has_variable("charlie_brown.name")
  assert charlie_brown.name != 'Charlie Brown'
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
--fcc-editable-region--

--fcc-editable-region--
```
