---
id: 68c9a624fd54288f694ebdbf
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--
可以大胆做一个“小实验”哟：尝试把 `name` 属性设成不是字符串的值，观察验证机制如何温柔但坚定地保护你的类。

体验过之后，请把修改 `charlie_brown` 名字的那行删掉，让它恢复成原来的 `'Charlie Brown'`。良好的测试 + 恢复习惯，会让你写代码越来越稳健！

# --hints--
你需要把 `charlie_brown` 的 `name` 恢复成原始值。

```js
({ test: () => runPython(`
  assert charlie_brown.name == 'Charlie Brown'
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
--fcc-editable-region--
charlie_brown.name = 'Snoopy'
--fcc-editable-region--
```
