---
id: 68c9a2c015ed7b54cf81e789
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--
现在来给你的程序增加一点“成果展示”吧！

在代码底部用一个 f-string 打印出 `Base salary: $`，后面跟上 `charlie_brown` 的薪资金额。让输出开心地告诉你：属性访问器已经生效啦！

继续保持这种细致与耐心，你正在构建一个既规范又易读的类。❤️

# --hints--
你需要打印一个 f-string：前缀是 `Base salary: $`，后面跟上 `charlie_brown.salary` 的数值。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(f'Base salary: \${charlie_brown.salary}')")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

--fcc-editable-region--
```
