---
id: 68c9a1c241206043c81293d5
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--
现在你已经有了一个崭新的小属性，我们要给它一扇优雅的“入口”啦——可爱的访问器（getter）！

请创建一个名为 `salary` 的方法，只需要一个 `self` 参数，并用 `@property` 装饰它。这样写之后，你就能像访问普通属性一样轻松获取薪资。方法体里返回 `self._salary` 就好。

加油，一步一步构建出专业又贴心的类，你已经走得很棒了！

# --hints--
你的 `Employee` 类需要有一个带 `self` 参数的 `salary` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("salary").has_args("self")`)) })
```

`salary` 方法应该使用 `@property` 装饰器。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("salary").has_decorators("property")`)) })
```

`salary` 方法应该返回 `self._salary`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("salary").has_return("self._salary")`)) })
```


# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--

--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
