---
id: 68cab02fd80a91042c0165b8
title: Step 34
challengeType: 20
dashedName: step-34
---

# --description--
继续强化 `salary` setter：防止把薪资改得比当前等级的基础薪资还低。

在已有类型检查后再加一个 `if`：若新薪资 < 当前等级的基础薪资，就抛出 `ValueError`，信息：`Salary must be higher than minimum salary $<基础薪资>.`。

这样就保证不会出现“降到不合理”的情况，逻辑更严谨！

# --hints--
`salary` setter 中需要第二个 `if` 用于最低薪资约束。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("salary")[1].find_ifs()[1]`)) })
```

当新薪资低于基础薪资时应抛出对应 `ValueError`（带数字提示）。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  minimum_salary = Employee._base_salaries['trainee']
  try:
    emp.salary = minimum_salary - 1
  except ValueError as e:
    print = built_in_print
    assert str(e) == f"Salary must be higher than minimum salary \${minimum_salary}."
  else:
    assert False, "Expected to raise ValueError with invalid new_salary"
`) })
```

当新薪资高于最低限制时不应抛异常。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  built_in_print = print
  print = lambda x: x
  minimum_salary = Employee._base_salaries['trainee']
  try:
    emp.salary = minimum_salary + 1
  except Exception:
    assert False, "Expected not to raise exception with valid new_salary"
  try:
    emp.salary = minimum_salary + 1000
  except Exception:
    assert False, "Expected not to raise exception with valid new_salary"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
        if new_level == self.level:
            raise ValueError(f"'{self.level}' is already the selected level.")
        if Employee._base_salaries[new_level] < Employee._base_salaries[self.level]:
            raise ValueError(f"Cannot change to lower level.")
        print(f"'{self.name}' promoted to '{new_level}'.")
        self._salary = Employee._base_salaries[new_level]
        self._level = new_level

    @property
    def salary(self):
        return self._salary
--fcc-editable-region--
    @salary.setter
    def salary(self, new_salary):
        if not isinstance(new_salary, (int, float)):
            raise TypeError("'salary' must be a number.")


        self._salary = new_salary
        print(f'Salary updated to ${self.salary}.')
--fcc-editable-region--

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

charlie_brown.level = 'junior'

```
