---
id: 674757668684f9deaa6ac3c0
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

太棒了！你的歌曲数组已经有八首歌曲了。让我们继续添加第九首歌曲。

在数组中第八首歌曲对象的下方，添加第九个歌曲对象，包含以下属性：

```js
id: 8,
title: "The Surest Way Out is Through",
artist: "Quincy Larson",
duration: "3:10",
src: "https://cdn.MasterPuti.org/curriculum/js-music-player/the-surest-way-out-is-through.mp3",
```

# --hints--

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有九个对象。

```js
assert.lengthOf(allSongs, 9);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第九个对象应该有一个 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `8`。

```js
assert.strictEqual(allSongs[8].id, 8);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第九个对象应该有一个 [title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"The Surest Way Out is Through"`。

```js
assert.strictEqual(allSongs[8].title, "The Surest Way Out is Through");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第九个对象应该有一个 [artist](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"Quincy Larson"`。

```js
assert.strictEqual(allSongs[8].artist, "Quincy Larson");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第九个对象应该有一个 [duration](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"3:10"`。

```js
assert.strictEqual(allSongs[8].duration, "3:10");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第九个对象应该有一个 [src](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"https://cdn.MasterPuti.org/curriculum/js-music-player/the-surest-way-out-is-through.mp3"`。

```js
assert.strictEqual(allSongs[8].src, "https://cdn.MasterPuti.org/curriculum/js-music-player/the-surest-way-out-is-through.mp3");
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [
  {
    id: 0,
    title: "Scratching The Surface",
    artist: "Quincy Larson",
    duration: "4:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3",
  },
  {
    id: 1,
    title: "Can't Stay Down",
    artist: "Quincy Larson",
    duration: "4:15",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/can't-stay-down.mp3",
  },
  {
    id: 2,
    title: "Still Learning",
    artist: "Quincy Larson",
    duration: "3:51",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/still-learning.mp3",
  },
  {
    id: 3,
    title: "Cruising for a Musing",
    artist: "Quincy Larson",
    duration: "3:34",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/cruising-for-a-musing.mp3",
  },
  {
    id: 4,
    title: "Never Not Favored",
    artist: "Quincy Larson",
    duration: "3:35",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/never-not-favored.mp3",
  },
  {
    id: 5,
    title: "From the Ground Up",
    artist: "Quincy Larson",
    duration: "3:12",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/from-the-ground-up.mp3",
  },
  {
    id: 6,
    title: "Walking on Air",
    artist: "Quincy Larson",
    duration: "3:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/walking-on-air.mp3",
  },
  {
    id: 7,
    title: "Can't Stop Me. Can't Even Slow Me Down.",
    artist: "Quincy Larson",
    duration: "3:52",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/cant-stop-me-cant-even-slow-me-down.mp3",
  },
  {
    id: 8,
    title: "The Surest Way Out is Through",
    artist: "Quincy Larson",
    duration: "3:10",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/the-surest-way-out-is-through.mp3",
  }
]
```
