---
id: 67471d60ce8e6268fe9e12f6
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

在这个工作坊中，你将使用 HTML、CSS 和 JavaScript 构建一个音乐播放器。你将学习如何操作 DOM、处理事件以及使用 HTML5 音频 API。

首先，你需要创建一个包含所有歌曲信息的数组。这将是你的播放器数据结构的基础。

在 `script.js` 文件中，创建一个名为 `allSongs` 的空数组。

# --hints--

你应该声明一个名为 `allSongs` 的变量。

```js
assert.isDefined(allSongs);
```

`allSongs` 变量应该是一个数组。

```js
assert.isArray(allSongs);
```

`allSongs` 数组应该为空。

```js
assert.lengthOf(allSongs, 0);
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
```
