---
id: 674724ff8af92f7733e6d980
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

接下来，你需要创建一个 `Audio` 对象来处理音频播放。这将允许你播放、暂停和控制音频。

在 `userData` 对象下方，创建一个名为 `audio` 的新变量，并将其设置为一个新的 `Audio` 实例。

# --hints--

你应该声明一个名为 `audio` 的变量。

```js
assert.isDefined(audio);
```

`audio` 变量应该是一个 `Audio` 实例。

```js
assert.instanceOf(audio, Audio);
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [];

const userData = {
  songs: allSongs,
  currentSong: null,
  songCurrentTime: 0,
};
```