---
id: 67473042e8751eac62fafc82
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

继续加油！你的歌曲数组已经包含了六首歌曲。让我们继续添加第七首歌曲。

在数组中的第六首歌曲对象下方，添加一个具有以下属性的第七首歌曲对象：

```js
id: 6,
title: "Walking on Air",
artist: "Quincy Larson",
duration: "3:25",
src: "https://cdn.freecodecamp.org/curriculum/js-music-player/walking-on-air.mp3",
```

# --hints--

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有七个对象。

```js
assert.lengthOf(allSongs, 7);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第七个对象应该有一个值为 `6` 的 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性。

```js
assert.strictEqual(allSongs[6].id, 6);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第七个对象应该有一个值为 `"Walking on Air"` 的 [title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性。

```js
assert.strictEqual(allSongs[6].title, "Walking on Air");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第七个对象应该有一个值为 `"Quincy Larson"` 的 [artist](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性。

```js
assert.strictEqual(allSongs[6].artist, "Quincy Larson");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第七个对象应该有一个值为 `"3:25"` 的 [duration](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性。

```js
assert.strictEqual(allSongs[6].duration, "3:25");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第七个对象应该有一个值为 `"https://cdn.freecodecamp.org/curriculum/js-music-player/walking-on-air.mp3"` 的 [src](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性。

```js
assert.strictEqual(allSongs[6].src, "https://cdn.freecodecamp.org/curriculum/js-music-player/walking-on-air.mp3");
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [
  {
    id: 0,
    title: "Scratching The Surface",
    artist: "Quincy Larson",
    duration: "4:25",
    src: "https://cdn.freecodecamp.org/curriculum/js-music-player/scratching-the-surface.mp3",
  },
  {
    id: 1,
    title: "Can't Stay Down",
    artist: "Quincy Larson",
    duration: "4:15",
    src: "https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3",
  },
  {
    id: 2,
    title: "Still Learning",
    artist: "Quincy Larson",
    duration: "3:51",
    src: "https://cdn.freecodecamp.org/curriculum/js-music-player/still-learning.mp3",
  },
  {
    id: 3,
    title: "Cruising for a Musing",
    artist: "Quincy Larson",
    duration: "3:34",
    src: "https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3",
  },
  {
    id: 4,
    title: "Never Not Favored",
    artist: "Quincy Larson",
    duration: "3:35",
    src: "https://cdn.freecodecamp.org/curriculum/js-music-player/never-not-favored.mp3",
  },
  {
    id: 5,
    title: "From the Ground Up",
    artist: "Quincy Larson",
    duration: "3:12",
    src: "https://cdn.freecodecamp.org/curriculum/js-music-player/from-the-ground-up.mp3",
  },
  {
    id: 6,
    title: "Walking on Air",
    artist: "Quincy Larson",
    duration: "3:25",
    src: "https://cdn.freecodecamp.org/curriculum/js-music-player/walking-on-air.mp3",
  }
]
```