---
id: 6747257cdf5412781e138bb2
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

在 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组内部，创建一个具有以下属性和值的对象：

```js
id: 0,
title: "Scratching The Surface",
artist: "Quincy Larson",
duration: "4:25",
src: "https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3",
```

太棒了！你已经学会了如何创建包含歌曲信息的对象数组。这是构建音乐播放器数据结构的重要一步！

# --hints--

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有一个值。

```js
assert.equal(allSongs.length, 1);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组的第一个值应该是一个对象。

```js
assert.isObject(allSongs[0]);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有一个对象，其 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性设置为数字 `0`。

```js
assert.strictEqual(allSongs[0].id, 0);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有一个对象，其 [title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性设置为字符串 `"Scratching The Surface"`。

```js
assert.equal(allSongs[0].title, "Scratching The Surface");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有一个对象，其 [artist](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性设置为字符串 `"Quincy Larson"`。

```js
assert.equal(allSongs[0].artist, "Quincy Larson");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有一个对象，其 [duration](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性设置为字符串 `"4:25"`。

```js
assert.equal(allSongs[0].duration, "4:25");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有一个对象，其 [src](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性设置为字符串 `"https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3"`。

```js
assert.equal(allSongs[0].src, "https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3");
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [
  {
    id: 0,
    title: "Scratching The Surface",
    artist: "Quincy Larson",
    duration: "4:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3",
  }
]
```

