---
id: 67472ab9f570797dbb01568a
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

你做得很好！你的歌曲数组已经包含了四首歌曲。让我们继续添加第五首歌曲。

在数组中第四首歌曲对象的下方，添加第五个歌曲对象，包含以下属性：

```js
id: 4,
title: "Never Not Favored",
artist: "Quincy Larson",
duration: "3:35",
src: "https://cdn.MasterPuti.org/curriculum/js-music-player/never-not-favored.mp3",
```

# --hints--

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组应该有五个对象。

```js
assert.lengthOf(allSongs, 5);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第五个对象应该有一个 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `4`。

```js
assert.strictEqual(allSongs[4].id, 4);
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第五个对象应该有一个 [title](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"Never Not Favored"`。

```js
assert.strictEqual(allSongs[4].title, "Never Not Favored");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第五个对象应该有一个 [artist](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"Quincy Larson"`。

```js
assert.strictEqual(allSongs[4].artist, "Quincy Larson");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第五个对象应该有一个 [duration](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"3:35"`。

```js
assert.strictEqual(allSongs[4].duration, "3:35");
```

你的 [allSongs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747237371f13173c9f6c27e.md#L12-L12) 数组中的第五个对象应该有一个 [src](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-music-player/6747257cdf5412781e138bb2.md#L12-L12) 属性，其值为 `"https://cdn.MasterPuti.org/curriculum/js-music-player/never-not-favored.mp3"`。

```js
assert.strictEqual(allSongs[4].src, "https://cdn.MasterPuti.org/curriculum/js-music-player/never-not-favored.mp3");
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [
  {
    id: 0,
    title: "Scratching The Surface",
    artist: "Quincy Larson",
    duration: "4:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3",
  },
  {
    id: 1,
    title: "Can't Stay Down",
    artist: "Quincy Larson",
    duration: "4:15",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/can't-stay-down.mp3",
  },
  {
    id: 2,
    title: "Still Learning",
    artist: "Quincy Larson",
    duration: "3:51",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/still-learning.mp3",
  },
  {
    id: 3,
    title: "Cruising for a Musing",
    artist: "Quincy Larson",
    duration: "3:34",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/cruising-for-a-musing.mp3",
  },
  {
    id: 4,
    title: "Never Not Favored",
    artist: "Quincy Larson",
    duration: "3:35",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/never-not-favored.mp3",
  }
]
```