---
id: 6747237371f13173c9f6c27e
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

太棒了！你已经创建了一个空数组来存储你的歌曲。现在你需要创建一个对象来存储用户数据，比如当前播放的歌曲和歌曲的当前时间。

在 `allSongs` 数组下方，创建一个名为 `userData` 的对象，包含以下属性：

1. `songs`：设置为 `allSongs` 数组
2. `currentSong`：设置为 `null`
3. `songCurrentTime`：设置为 `0`

# --hints--

你应该声明一个名为 `userData` 的变量。

```js
assert.isDefined(userData);
```

`userData` 变量应该是一个对象。

```js
assert.isObject(userData);
```

`userData` 对象应该有一个 `songs` 属性，其值为 `allSongs` 数组。

```js
assert.strictEqual(userData.songs, allSongs);
```

`userData` 对象应该有一个 `currentSong` 属性，其值为 `null`。

```js
assert.isNull(userData.currentSong);
```

`userData` 对象应该有一个 `songCurrentTime` 属性，其值为 `0`。

```js
assert.strictEqual(userData.songCurrentTime, 0);
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [];
```