---
id: 6748bab9827a356c39559214
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

很好！现在让我们处理音频播放的当前时间。你需要检查是否已经有当前播放的歌曲，如果没有，将音频时间设置为0。

在设置 `audio.title` 之后，添加一个 `if/else` 语句：

1. 如果 `userData.currentSong` 为 `null`，则将 `audio.currentTime` 设置为 `0`
2. 否则，将 `audio.currentTime` 设置为 `userData.songCurrentTime`

这将确保当你播放一首新歌时从头开始播放，而当你恢复暂停的歌曲时从之前的位置继续播放。

# --hints--

你应该有一个 `if/else` 语句来检查 `userData.currentSong === null`。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*userData\s*\.\s*currentSong\s*===\s*null\s*\)/);
```

当 `userData.currentSong` 为 `null` 时，你应该将 `audio.currentTime` 设置为 `0`。

```js
assert.match(__helpers.removeJSComments(code), /audio\s*\.\s*currentTime\s*=\s*0/);
```

在 `else` 块中，你应该将 `audio.currentTime` 设置为 `userData.songCurrentTime`。

```js
assert.match(__helpers.removeJSComments(code), /audio\s*\.\s*currentTime\s*=\s*userData\s*\.\s*songCurrentTime/);
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [
  {
    id: 0,
    title: "Scratching The Surface",
    artist: "Quincy Larson",
    duration: "4:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3",
  },
  {
    id: 1,
    title: "Can't Stay Down",
    artist: "Quincy Larson",
    duration: "4:15",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/can't-stay-down.mp3",
  },
  {
    id: 2,
    title: "Still Learning",
    artist: "Quincy Larson",
    duration: "3:51",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/still-learning.mp3",
  },
  {
    id: 3,
    title: "Cruising for a Musing",
    artist: "Quincy Larson",
    duration: "3:34",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/cruising-for-a-musing.mp3",
  },
  {
    id: 4,
    title: "Never Not Favored",
    artist: "Quincy Larson",
    duration: "3:35",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/never-not-favored.mp3",
  },
  {
    id: 5,
    title: "From the Ground Up",
    artist: "Quincy Larson",
    duration: "3:12",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/from-the-ground-up.mp3",
  },
  {
    id: 6,
    title: "Walking on Air",
    artist: "Quincy Larson",
    duration: "3:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/walking-on-air.mp3",
  },
  {
    id: 7,
    title: "Can't Stop Me. Can't Even Slow Me Down.",
    artist: "Quincy Larson",
    duration: "3:52",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/cant-stop-me-cant-even-slow-me-down.mp3",
  },
  {
    id: 8,
    title: "The Surest Way Out is Through",
    artist: "Quincy Larson",
    duration: "3:10",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/the-surest-way-out-is-through.mp3",
  },
  {
    id: 9,
    title: "Chasing That Feeling",
    artist: "Quincy Larson",
    duration: "2:43",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/chasing-that-feeling.mp3",
  },
];

const audio = new Audio();

const userData = {
  songs: allSongs,
  currentSong: null,
  songCurrentTime: 0,
};

const playSong = (id) => {
  const song = userData.songs.find((song) => song.id === id);
  audio.src = song.src;
  audio.title = song.title;
  if (userData.currentSong === null) {
    audio.currentTime = 0;
  } else {
    audio.currentTime = userData.songCurrentTime;
  }
};
```