---
id: 674e305b7bda7c1f4f37a930
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

很好！现在让我们为播放按钮添加事件监听器，以便用户可以播放歌曲。

在 `pauseSong` 函数下方，添加一个事件监听器到 `playButton` 元素上，监听 `click` 事件。在事件处理函数内部，添加以下代码：

1. 如果 `userData.currentSong` 为 `null`，则调用 `playSong` 函数并传入 `userData.songs[0].id` 作为参数
2. 否则，调用 `playSong` 函数并传入 `userData.currentSong.id` 作为参数

这将确保如果当前没有播放歌曲，则播放列表中的第一首歌曲；如果已有歌曲暂停，则恢复播放。

# --hints--

你应该为 `playButton` 元素添加一个事件监听器。

```js
assert.match(__helpers.removeJSComments(code), /playButton\s*\.\s*addEventListener\s*\(\s*['"`]click['"`]/);
```

当 `userData.currentSong` 为 `null` 时，你应该调用 `playSong(userData.songs[0].id)`。

```js
assert.match(__helpers.removeJSComments(code), /playSong\s*\(\s*userData\s*\.\s*songs\s*\[\s*0\s*\]\s*\.\s*id\s*\)/);
```

当 `userData.currentSong` 不为 `null` 时，你应该调用 `playSong(userData.currentSong.id)`。

```js
assert.match(__helpers.removeJSComments(code), /playSong\s*\(\s*userData\s*\.\s*currentSong\s*\.\s*id\s*\)/);
```

# --seed--

## --seed-contents--

```js
const playlistSongs = document.getElementById("playlist-songs");
const playButton = document.getElementById("play");
const pauseButton = document.getElementById("pause");
const nextButton = document.getElementById("next");
const previousButton = document.getElementById("previous");
const allSongs = [
  {
    id: 0,
    title: "Scratching The Surface",
    artist: "Quincy Larson",
    duration: "4:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/scratching-the-surface.mp3",
  },
  {
    id: 1,
    title: "Can't Stay Down",
    artist: "Quincy Larson",
    duration: "4:15",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/can't-stay-down.mp3",
  },
  {
    id: 2,
    title: "Still Learning",
    artist: "Quincy Larson",
    duration: "3:51",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/still-learning.mp3",
  },
  {
    id: 3,
    title: "Cruising for a Musing",
    artist: "Quincy Larson",
    duration: "3:34",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/cruising-for-a-musing.mp3",
  },
  {
    id: 4,
    title: "Never Not Favored",
    artist: "Quincy Larson",
    duration: "3:35",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/never-not-favored.mp3",
  },
  {
    id: 5,
    title: "From the Ground Up",
    artist: "Quincy Larson",
    duration: "3:12",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/from-the-ground-up.mp3",
  },
  {
    id: 6,
    title: "Walking on Air",
    artist: "Quincy Larson",
    duration: "3:25",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/walking-on-air.mp3",
  },
  {
    id: 7,
    title: "Can't Stop Me. Can't Even Slow Me Down.",
    artist: "Quincy Larson",
    duration: "3:52",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/cant-stop-me-cant-even-slow-me-down.mp3",
  },
  {
    id: 8,
    title: "The Surest Way Out is Through",
    artist: "Quincy Larson",
    duration: "3:10",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/the-surest-way-out-is-through.mp3",
  },
  {
    id: 9,
    title: "Chasing That Feeling",
    artist: "Quincy Larson",
    duration: "2:43",
    src: "https://cdn.MasterPuti.org/curriculum/js-music-player/chasing-that-feeling.mp3",
  },
];

const audio = new Audio();

const userData = {
  songs: allSongs,
  currentSong: null,
  songCurrentTime: 0,
};

const playSong = (id) => {
  const song = userData.songs.find((song) => song.id === id);
  audio.src = song.src;
  audio.title = song.title;
  if (userData.currentSong === null) {
    audio.currentTime = 0;
  } else {
    audio.currentTime = userData.songCurrentTime;
  }
  userData.currentSong = song;
  playButton.classList.add("playing");
  audio.play();
};

const pauseSong = () => {
  userData.songCurrentTime = audio.currentTime;
  playButton.classList.remove("playing");
  audio.pause();
};

playButton.addEventListener("click", () => {
  if (userData.currentSong === null) {
    playSong(userData.songs[0].id);
  } else {
    playSong(userData.currentSong.id);
  }
});
```