---
id: 683ecb010f4c6bafd428802c
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本练习中，你将对一组医疗数据进行验证，确保它们符合一系列校验规则。

请声明一个名为 `medical_records` 的变量，并将其赋值为一个空列表。在接下来的步骤中，你会使用它来存放你的医疗数据，继续加油！

# --hints--

你需要声明一个名为 `medical_records` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

将 `medical_records` 赋值为一个空列表即可。

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

