---
id: 6842d659a4ae7973dabab3b0
title: 步骤 4
challengeType: 20
dashedName: step-4
---

# --description--

完善该字典，添加以下三对键值：

- 键 `diagnosis`，值为字符串 `Hypertension`
- 键 `medications`，值为列表 `['Lisinopril']`
- 键 `last_visit_id`，值为字符串 `V2301`

完成这些字段后，你就补全了该患者的基本记录，棒极了！

# --hints--

你的字典应包含 `diagnosis` 键，值为字符串 `Hypertension`。

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

应包含 `medications` 键，值为列表 `['Lisinopril']`。

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

应包含 `last_visit_id` 键，值为字符串 `V2301`。

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

此外，请确保此前要求的 `patient_id`、`age` 和 `gender` 字段仍然存在并符合要求。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female'
    }    
]
--fcc-editable-region--
```
