---
id: 6842aee5bcaf4ab8883f8191
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

正如你在之前的课程中学到的，字典是一种存储键值对的数据结构，其中键通常是不可变的数据类型（例如字符串），并且在同一个字典中保持唯一：

```py
person = {
    'name': 'John',
    'age': 33
}
```

`medical_records` 列表将用来存放字典，每个字典代表一位患者。请向 `medical_records` 列表中添加一个字典，包含键 `patient_id`，其值为字符串 `P1001`。加油，开始吧！

# --hints--

你的 `medical_records` 列表中应包含一个字典。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

该字典应包含键 `patient_id`，其值为字符串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
medical_records = [

]
--fcc-editable-region--
```
