---
id: 6842d5b537c64266420faf62
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

向你的字典中添加键 `age`（值为整数 `34`）和键 `gender`（值为字符串 `Female`）。别忘了在键值对之间加逗号哦。

# --hints--

你的 `medical_records` 列表中应包含一个字典。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

该字典应包含键 `age`，其值为整数 `34`。

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

该字典应包含键 `gender`，其值为字符串 `Female`。

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

此外，该字典还应包含先前要求的 `patient_id` 键，其值为字符串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
medical_records = [
    {
        'patient_id': 'P1001'
    }    
]
--fcc-editable-region--
```
