---
id: 68e6a99a0fd7386fb66e156f
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将通过构建一个游戏设置面板来练习样式化复选框。

所有的HTML模板代码（`DOCTYPE`、`html`、`head`和`body`）都已经为你提供好了。

在`body`中，创建一个`div`元素，内部嵌套一个`h1`元素。给`h1`元素添加文本`Game Settings`

# --hints--

你应该有一个`div`元素。

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

你应该有一个`h1`元素。

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

你的`h1`元素应该在你的`div`元素内部。

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

你的`h1`元素应该有文本`Game Settings`。

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
