# --description--

紧接在你的`h1`元素之后，创建四个`label`元素。

在每个`label`元素内部，创建一个`input`元素，将`type`属性设置为`"checkbox"`。

在每个`input`元素之后，按顺序给你的`label`元素添加以下文本：`Sound Effects`、`Background Music`、`Hard Mode`和`Haptic Feedback`。

# --hints--

你应该有四个`label`元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

你的第一个`label`元素应该有一个`input`元素。

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

你的第一个`input`元素应该将`type`属性设置为`checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

你的第一个`label`元素应该有文本`Sound Effects`。

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

`Sound Effects`文本应该在`input`元素之后。

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

你的第二个`label`元素应该有一个`input`元素。

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

你的第二个`input`元素应该将`type`属性设置为`checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

你的第二个`label`元素应该有文本`Background Music`。

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

`Background Music`文本应该在`input`元素之后。

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

你的第三个`label`元素应该有一个`input`元素。

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

你的第三个`input`元素应该将`type`属性设置为`checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

你的第三个`label`元素应该有文本`Hard Mode`。

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

`Hard Mode`文本应该在`input`元素之后。

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

你的第四个`label`元素应该有一个`input`元素。

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

你的第四个`input`元素应该将`type`属性设置为`checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

你的第四个`label`元素应该有文本`Haptic Feedback`。

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

`Haptic Feedback`文本应该在`input`元素之后。

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```