---
id: 68e9a4a47d8a9802cba886b7
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

现在body已经设置了样式，你将为游戏设置面板创建一个容器卡片。首先，为`settings-card`创建一个类选择器。

这里将放置你容器的所有格式化设置。将`max-width`设置为`250px`来定义容器的整体大小。

接下来，将`padding`设置为`20px`，这样你的内容与容器边框之间就有空间了。

之后，通过将`border-radius`设置为`10px`来创建圆角边缘。

然后设置`box-shadow`值为`0 2px 6px rgba(0,0,0,0.2)`。这将创建一个微妙的"悬浮"外观，为容器增加深度感。

# --hints--

你的CSS中应该有一个`settings-card`的类选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

你应该为`.settings-card`类选择器添加`max-width`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

你应该有一个值为`250px`的`max-width`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

你应该为`.settings-card`类选择器添加`padding`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

你应该有一个值为`20px`的`padding`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

你应该为`.settings-card`类选择器添加`border-radius`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

你应该有一个值为`10px`的`border-radius`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

你应该为`.settings-card`类选择器添加`box-shadow`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

你应该有一个值为`0 2px 6px rgba(0,0,0,0.2)`的`box-shadow`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--