---
id: 68e9a6201e2efa142d62bfe1
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

接下来，你将放大复选框以提高可见性。

首先设置一个`input`选择器，但特别针对你的`[type="checkbox"]`。

在你的选择器中，将`width`和`height`都设置为`20px`。这使它比以前更大。

为了符合在标签上设置的`cursor`，添加`cursor`并将值设置为`pointer`。这样，当你悬停在标签上时，它会显示一个指针。

# --hints--

你的CSS中应该有一个`input[type="checkbox"]`选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

你应该为`input[type="checkbox]`选择器添加`width`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

你应该有一个值为`20px`的`width`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

你应该为`input[type="checkbox]`选择器添加`height`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

你应该有一个值为`20px`的`height`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

你应该为`input[type="checkbox]`选择器添加`cursor`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

你应该有一个值为`pointer`的`cursor`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
