---
id: 68f29dbf9125476b4808d6fb
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

太棒了！现在在你的`.settings-card`中设置`margin`属性为`auto`。

将`margin`属性设置为`auto`会自动调整元素的边距，使其在容器中均匀分布剩余空间，这通常用于水平居中块级元素。

最后，设置一个值为`left`的`text-align`属性。这会将内联内容（如文本）对齐到其包含元素的左侧。

# --hints--

你应该在`settings-card`类选择器中添加一个`margin`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

你应该有一个值为`auto`的`margin`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

你应该在`settings-card`类选择器中添加一个`text-align`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

你应该有一个值为`left`的`text-align`属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--

--fcc-editable-region--
}
```