---
id: 5d822fd413a79914d39e98d3
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `.bb1` 容器中嵌套四个 `div` 元素。按顺序给它们添加 `bb1a`、`bb1b`、`bb1c` 和 `bb1d` 类。这座建筑物将有四个部分。

# --hints--

你应该创建四个新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

你应该给其中一个新 `div` 元素添加 `bb1a` 类。

```js
assert.exists(document.querySelector('div.bb1a'));
```

你应该给其中一个新 `div` 元素添加 `bb1b` 类。

```js
assert.exists(document.querySelector('div.bb1b'));
```

你应该给其中一个新 `div` 元素添加 `bb1c` 类。

```js
assert.exists(document.querySelector('div.bb1c'));
```

你应该给其中一个新 `div` 元素添加 `bb1d` 类。

```js
assert.exists(document.querySelector('div.bb1d'));
```

你应该按正确的顺序放置新的 `div` 元素。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

你应该将新的 `div` 元素放置在 `.bb1` 元素内。

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```