---
id: 5d822fd413a79914d39e98ce
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

还要为所有元素添加 `box-sizing` 属性，值为 `border-box`。这样你添加的边框就不会增加元素的大小了。

# --hints--

你应该使用 `box-sizing` 属性。

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

你应该使用现有的 `*` 选择器。

```js
// Two selectors create two CSSStyleRule objects
assert.lengthOf(new __helpers.CSSHelp(document).getStyleDeclarations('*'), 1);
```

所有元素都应该有 `box-sizing` 属性，值为 `border-box`。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--

```