---
id: 5d822fd413a79914d39e98cd
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 CSS 中，你可以使用星号来选择所有元素。使用 `*` 选择器为所有元素添加边框，并设置 `border` 为 `1px solid black`。这是一个有助于可视化元素位置和大小的技巧。你稍后会删除这个边框。

# --hints--

你应该使用 `*` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

你应该使用 `border` 属性来设置所有元素的样式。

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

所有元素都应该有 `1px solid black` 的边框。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}
--fcc-editable-region--

```