---
id: 686386e76197458b7dfd52ed
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

在刚才新增的那个价格说明 `p` 下面，再放一个新的 `p`，文本写 `Start exploring millions of songs with basic features and ads.`。紧接着在它下面添加一个空的 `ul`。稍后我们会往这个列表里塞进可爱的小功能条目～

# --hints--

需要再创建一个文本为 `Start exploring millions of songs with basic features and ads.` 的 `p`。

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)
assert.exists(targetP)
```

在该 `p` 的下面要有一个空的 `ul` 元素。

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)

const ul = targetP.nextElementSibling
assert.equal(ul?.tagName, "UL")
assert.isEmpty(ul?.innerHTML.trim())
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            --fcc-editable-region--

            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
