---
id: 686386e76197458b7dfd52ef
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

现在来做一个 CTA 按钮。CTA（Call to Action）指的是引导用户执行某个特定操作的按钮或链接，这里就是鼓励用户“开始试听”。

在最内层的那个 `div` 外面，添加一个 `a`，`href` 设为 `#`，文字写 `Start listening`。

这样第一个定价卡片的 HTML 就圆满完成啦，撒花～

# --hints--

需要创建一个 `a` 元素。

```js
assert.exists(document.querySelector("a"))
```

该 `a` 的 `href` 应该是 `#`。

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

该 `a` 的文本应为 `Start listening`。

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--

          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
