---
id: 68eac57fadcc662ce917e75e
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

接下来，添加一个 `button` 元素，将 `type` 属性设置为 `button`。

`button` 元素的文本应该是 `Play`。

# --hints--

你应该有一个 `button` 元素。

```js
assert.exists(document.querySelector('button'));
```

你的 `button` 元素应该有一个 `type` 属性。

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

你的 `button` 元素应该将 `type` 属性设置为 `button`。

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

你的 `button` 元素应该有文本 `Play`。

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
