# --description--

接下来，你需要添加一个 `input` 元素，将 `type` 属性设置为 `range` 来创建音量滑块。将 `min` 属性设置为 `0`，`max` 属性设置为 `100`，`value` 属性设置为 `50` 来定义默认音量级别。

# --hints--

你应该在第二个 `span` 元素下方有一个 `input` 元素。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

你的 `input` 元素应该有一个 `type` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

你的 `input` 元素应该有一个 `type` 属性，值为 `range`。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

你的 `input` 元素应该有一个 `min` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

你的 `input` 元素应该有一个 `min` 属性，值为 `0`。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

你的 `input` 元素应该有一个 `max` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

你的 `input` 元素应该有一个 `max` 属性，值为 `100`。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

你的 `input` 元素应该有一个 `value` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

你的 `input` 元素应该有一个 `value` 属性，值为 `50`。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```