---
id: 6733b5e6f210c1bd1afdf8b3
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

其余的植物对象已经为你添加好了。现在是时候创建你的目录了。

就像普通的JavaScript对象一样，`Map`对象包含一组键值对。要创建一个`Map`对象，你可以使用`Map`构造函数：

```js
const myMap = new Map();
```

声明一个`catalog`变量并为其赋值一个`Map`对象。

# --hints--

你应该有一个名为`catalog`的变量。

```js
assert.exists(catalog);
```

你应该为`catalog`变量赋值一个`Map`对象。

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```