---
id: 6733b5632fb854bc6c58dc0d
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

现在创建另一个对象，具有相同的三个属性，用来表示名为`Pretty Polly`的西班牙薰衣草的另一个品种，并将其赋值给名为`prettyPolly`的变量。

# --hints--

你应该有一个名为`prettyPolly`的变量。

```js
assert.exists(prettyPolly);
```

你应该为`prettyPolly`变量赋值一个对象。

```js
assert.isObject(prettyPolly);
```

你的`prettyPolly`对象应该有一个`commonName`属性，值为`"西班牙薰衣草"`。

```js
assert.strictEqual(prettyPolly.commonName, "西班牙薰衣草");
```

你的`prettyPolly`对象应该有一个`scientificName`属性，值为`"Lavandula stoechas"`。

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

你的`prettyPolly`对象应该有一个`cultivar`属性，值为`"Pretty Polly"`。

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

你的`prettyPolly`对象应该有三个键。

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```