---
id: 6734fa225c6667121e589f7e
title: 步骤 17
challengeType: 1
dashedName: step-17
---

# --description--

你的`sellPlants`函数正在引发一个`TypeError`，因为它试图读取一个已从目录中删除的对象的属性。因此，你需要处理在`catalog`map中找不到键的情况。为此，`has`方法非常有用。它返回一个布尔值，指示指定的键是否存在于map中。

修改你的`sellPlants`函数，当目录不包含传递给函数的植物时，返回`Item not found.`。

# --hints--

你的`sellPlants`函数应该在`catalog`上调用`has`方法。

```js
let flag = false;
const temp = catalog.has;
try {
    catalog.has = (p) => {flag = true}
    sellPlants("test", "", 0);
    sellPlants(imperialGem, "medium", 200);
    assert.isTrue(flag);
} finally {
    catalog.has = temp;
}
```

你不应该直接修改插入到`catalog`map中的键值对。

```js
const testCatalog = new Map();
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

`sellPlants(ballerina, "small", 10)`应该返回`Item not found.`。

```js
assert.equal(sellPlants(ballerina, "small", 10), "Item not found.")
```

`sellPlants(hidcote, "large", 95)`应该返回`Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.`。

```js
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")
assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

`sellPlants(hidcote, "large", 5)`应该返回`Catalog successfully updated.`并从目录中与`hidcote`键关联的对象的`large`属性中减去`5`。

```js
try {
    assert.equal(sellPlants(hidcote, "large", 5), "Catalog successfully updated.");
    assert.strictEqual(catalog.get(hidcote).large, 13);

    assert.equal(sellPlants(royalCrown, "medium", 5), "Catalog successfully updated.");
    assert.strictEqual(catalog.get(royalCrown).medium, 17);
} finally {
    catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
}
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const removePlant = plant => catalog.delete(plant);

console.log(removePlant(ballerina));
--fcc-editable-region--
function sellPlants(plant, potSize, numPots) {
    if (!catalog.has(plant)) {
        return "Item not found.";
    }
    
    const plantInfo = catalog.get(plant);
    if (numPots > plantInfo[potSize]) {
        return `Not enough ${potSize} size pots for ${plant.scientificName} '${plant.cultivar}'. Only ${plantInfo[potSize]} left.`;
    } else {
        plantInfo[potSize] -= numPots;
        return "Catalog successfully updated.";
    }
}
--fcc-editable-region--
console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```