---
id: 6734f0ae9d1cf4dfd43a7805
title: 步骤15
challengeType: 1
dashedName: step-15
---

# --description--

要从map对象中删除一个元素，你可以使用`delete`方法，它接受一个键作为参数并删除与该键关联的条目。它返回一个布尔值，指示删除是否成功。

创建一个名为`removePlant`的函数，它接受一个植物对象作为参数并将其从目录中删除。根据指定的植物是否已从目录中删除，让它返回`true`或`false`。

之后，将使用字符串`"ballerina"`作为参数调用你的新函数的结果输出到控制台。

# --hints--

你应该有一个名为`removePlant`的函数。

```js
assert.isFunction(removePlant);
```

你的`removePlant`函数应该接受一个参数。

```js
assert.lengthOf(removePlant, 1);
```

你的`removePlant`函数应该在`catalog`上调用`delete`方法。

```js
let flag = false;
const temp = catalog.delete;
try {
    catalog.delete = (p) => {flag = true}
    removePlant("");
    assert.isTrue(flag);
} finally {
    catalog.delete = temp;
}
```

`removePlant(item)`应该从你的目录中删除`item`。

```js
const testArr = [];
const temp = catalog.delete;
try {
    catalog.delete = (p) => {testArr.push(p)}
    removePlant("test");
    assert.deepEqual(testArr, ["test"]);
} finally {
    catalog.delete = temp;
}
```

你的`removePlant`函数应该返回一个布尔值，指示删除是否成功。

```js
try {
    assert.isFalse(removePlant("test"));
    assert.isTrue(removePlant(royalCrown));
} finally {
    catalog.set(royalCrown, { small: 40, medium: 22, large: 9 })
}
```

你应该将`removePlant("ballerina")`输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*removePlant\s*\(\s*("|')ballerina\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

--fcc-editable-region--
function removePlant(plant) {
    return catalog.delete(plant);
}

console.log(removePlant("ballerina"));

--fcc-editable-region--

const sellPlants = (plant, size, potsNo) => {
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

console.log(sellPlants(ballerina, "small", 10));

console.log(catalog)
```