---
id: 6734a606e5968a5a19b77171
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

继续像前面步骤中那样将剩余的植物添加到你的目录中。

# --hints--

你应该将你的`willowVale`对象设置为`catalog`map的键。

```js
assert.isTrue(catalog.has(willowVale));
```

`catalog`map中`willowVale`键的值应该是一个具有`small`、`medium`和`large`属性的对象，每个属性的值都是一个正整数。

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你应该将你的`hidcote`对象设置为`catalog`map的键。

```js
assert.isTrue(catalog.has(hidcote));
```

`catalog`map中`hidcote`键的值应该是一个具有`small`、`medium`和`large`属性的对象，每个属性的值都是一个正整数。

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你应该将你的`imperialGem`对象设置为`catalog`map的键。

```js
assert.isTrue(catalog.has(imperialGem));
```

`catalog`map中`imperialGem`键的值应该是一个具有`small`、`medium`和`large`属性的对象，每个属性的值都是一个正整数。

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你应该将你的`royalCrown`对象设置为`catalog`map的键。

```js
assert.isTrue(catalog.has(royalCrown));
```

`catalog`map中`royalCrown`键的值应该是一个具有`small`、`medium`和`large`属性的对象，每个属性的值都是一个正整数。

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你的`catalog`map中应该仍然有`ballerina`键。

```js
assert.isTrue(catalog.has(ballerina));
```

你的`catalog`map中应该仍然有`prettyPolly`键。

```js
assert.isTrue(catalog.has(prettyPolly));
```

你的`catalog`map应该包含六个元素。

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 15, medium: 10, large: 8 });
catalog.set(hidcote, { small: 25, medium: 20, large: 15 });
catalog.set(imperialGem, { small: 18, medium: 12, large: 10 });
catalog.set(royalCrown, { small: 12, medium: 8, large: 5 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```