---
id: 6734a13a29acdf47f2b8e161
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

与普通对象不同，普通对象的键必须始终是字符串，而map对象允许你使用任何数据类型作为键。在你的目录中，你将使用对象-对象键值对，将每个植物对象与一个包含不同尺寸可用花盆数量的对象关联起来。

修改你的`set`调用，使用你的`ballerina`对象作为键。对于值，使用一个包含三个属性的对象：`small`、`medium`和`large`，每个属性的值都是一个正整数。

# --hints--

你应该将你的`ballerina`对象设置为`catalog`map的键。

```js
assert.isTrue(catalog.has(ballerina));
```

`catalog`map中`ballerina`键的值应该是一个具有`small`、`medium`和`large`属性的对象。

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`catalog`map中与`ballerina`键关联的对象的`small`属性应该是一个正整数。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`catalog`map中与`ballerina`键关联的对象的`medium`属性应该是一个正整数。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`catalog`map中与`ballerina`键关联的对象的`large`属性应该是一个正整数。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你的`catalog`map应该包含一个键值对。

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, {
    small: 5,
    medium: 10,
    large: 3
});
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```