---
id: 67349ba230cff237c0f2bdf2
title: 步骤 4
challengeType: 1
dashedName: step-4
---

# --description--

Map对象会记住键插入到map中的顺序。要在你的map对象中设置一个键值对，你可以使用`set`方法。以下代码在`myMap`map对象中将键`"name"`设置为值`"Jerry"`。

```js
myMap.set("name", "Jerry");
```

在你的`catalog`map中将`"ballerina"`键设置为你的`ballerina`对象的值。然后，将你的catalog输出到控制台。

注意，要正确查看你的map对象，你需要打开浏览器控制台。

# --hints--

你应该在`catalog`map中将`"ballerina"`键设置为你的`ballerina`对象的值。

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

你的`catalog`map应该包含一个键值对。

```js
assert.lengthOf(catalog, 1);
```

你应该将你的`catalog`map输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```