---
id: 66a37f37ef5823a313de8c26
title: Build a Moon Orbit
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

在这个实验中，您将使用HTML和CSS创建一个月亮绕地球运行的简单动画。地球将位于系统的中心，月亮将围绕它运行。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该通过将高度设置为视口的[100%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)并设置body元素中的flexbox布局来将所有元素对齐到页面中心。
2. 您应该有一个带有[space](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L332-L336)类的div。
3. 在.space元素内，应该还有两个div元素，分别带有[earth](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L338-L348)和[orbit](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L350-L356)类，按顺序排列。
4. 在.orbit元素内，应该有另一个带有[moon](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L358-L368)类的div。
5. 带有[space](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L332-L336)类的div元素应该在页面上居中，并且宽度和高度为200px。
6. 带有[space](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L332-L336)类的div元素应该使用相对定位。
7. .earth元素应该使用绝对定位。将.earth元素的中心定位在其父元素垂直（top）和水平（left）轴的中点。之后，将.earth元素向后移动其自身宽度和高度的一半，使其在父元素.space元素内居中。
8. .earth元素应该有100px的宽度和高度。
9. .orbit元素应该有200px的宽度和高度。
10. .orbit元素应该使用绝对定位。其右下角应该使用transform属性在垂直和水平轴上都偏移-50%来定位在.space元素的中心。
11. 月亮绕地球运行的轨道路径应该是一个圆形。
12. .moon元素应该使用绝对定位，宽度和高度为30px。.moon元素应该定位在.orbit元素的顶部并水平居中。
13. 您应该通过将元素向左移动其宽度的一半来进一步调整.moon元素的水平定位。
14. .earth选择器应该有背景颜色和50%的border-radius。
15. .moon选择器应该有背景颜色和50%的border-radius。
16. 您应该定义一个@keyframes orbit动画，使.orbit元素绕其中心旋转360度。您应该将此动画应用于.orbit元素，持续时间为5秒，线性时间函数和无限迭代。

**注意：** 确保在HTML中链接样式表并应用CSS。

# --hints--

body元素的内容应该通过将高度设置为视口高度的[100%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)、display设置为flex、justify-content和align-items设置为center来在页面上居中。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

您应该有一个带有[space](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L332-L336)类的主div。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

.space元素应该有200px的宽度和高度。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

.space元素应该有relative的position属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

您应该在带有[space](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L332-L336)类的div内有一个带有[earth](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L338-L348)类的div。

```js
assert.exists(document.querySelector('div.space div.earth'));
```

.earth元素应该是.space元素的第一个子元素。

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

.earth元素应该有100px的宽度和高度。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

.earth元素应该使用绝对定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

.earth元素应该有[top](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-colored-boxes/66c071cd3475cd92421b9ac4.md#L45-L63)和[left](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-colored-boxes/66c071cd3475cd92421b9ac4.md#L45-L63)位置为[50%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

.earth元素应该通过将transform属性设置为translate(-50%, -50%)来向后和向上移动[50%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

您应该在.space元素内有一个带有[orbit](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L350-L356)类的div。

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

.orbit元素应该在.earth元素之后。

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

.orbit元素应该有200px的宽度和高度。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

.orbit元素应该使用绝对定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

.orbit元素应该使用transform属性在垂直和水平轴上都偏移-50%来定位在.space元素的中心。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

您应该在.orbit元素内有一个带有[moon](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-moon-orbit/66a37f37ef5823a313de8c26.md#L358-L368)类的div。

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

.moon元素应该有30px的宽度和高度。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

.moon元素应该使用绝对定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

.moon元素的顶部边缘应该与其包含块的顶部边缘对齐。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

您应该将.moon元素的左边缘移动到包含块的中心（从左侧[50%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)）。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

.moon元素应该通过在x轴上平移-50%来调整水平位置。

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

.earth选择器应该有背景颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

.earth选择器应该有[50%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)的border-radius。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

.moon选择器应该有背景颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

.moon选择器应该有[50%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)的border-radius。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

您应该有一个@keyframes规则。

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

您的新@keyframes规则应该命名为orbit。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

您的@keyframes orbit规则应该有一个[0%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)选择器。

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

您的@keyframes orbit规则应该有一个[100%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)选择器。

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

您的[0%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)选择器应该有一个设置为rotate(0deg) translate(-50%, -50%)的transform属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

您的[100%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)选择器应该在您的[0%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)选择器之后。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

您的[100%](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-celestial-bodies-database/66db559d34c7089b9b41bfd5.md#L47-L65)选择器应该有一个设置为rotate(360deg) translate(-50%, -50%)的transform属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

您的.orbit选择器应该有一个设置为orbit 5s linear infinite的animation属性。

```js
const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css
