---
id: 66ea89c25e1bd871688eba6a
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

到目前为止，你一直在处理浮点数和数字 `0`。

在工作坊的下一部分，我们将学习如何使用 `Math.floor()` 方法。

让我们先添加一个 `console.log()`，输出这句话："Math.floor() 方法会将数值向下舍入到最接近的整数。"

# --hints--

你应该输出这句话："Math.floor() 方法会将数值向下舍入到最接近的整数。"

```js
assert.match(code, /console\.log\((["'])Math\.floor\(\)\s+方法\s*会\s*将\s*数\s*值\s*向\s*下\s*舍\s*入\s*到\s*最\s*接\s*近\s*的\s*整\s*数\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);
--fcc-editable-region--

--fcc-editable-region--
```