---
id: 66ea7adae8053065a64f9002
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个精彩的编程工作坊中，我们将一起构建一个可爱的数学机器人，通过实践来复习 JavaScript 中各种 `Math` 对象的方法。这个数学机器人会在控制台中展示一些数学运算和有趣的消息，帮助你更好地理解 JavaScript 的 `Math` 对象！

在第一步中，请创建一个名为 `botName` 的变量，并将字符串值 `"MathBot"` 赋给它。

接下来，再创建一个名为 `greeting` 的变量，将这句话赋给它："你好呀！我的名字是[这里填写botName]，我将会教会你有关 Math 对象的知识！"。记得将 `[这里填写botName]` 这个占位符替换成实际的 `botName` 变量哦。

你可以自由选择使用模板字符串或者使用 `+` 运算符来进行字符串拼接。

最后，记得将 `greeting` 变量输出到控制台中。

# --hints--

你应该声明一个叫做 `botName` 的变量。

```js
assert.isNotNull(botName);
```

你的 `botName` 变量应该是一个字符串类型。

```js
assert.isString(botName);
```

你应该给 `botName` 变量赋值字符串 `"MathBot"`。

```js
assert.strictEqual(botName, "MathBot");
```

你应该声明一个叫做 `greeting` 的变量。

```js
assert.isNotNull(greeting);
```

你的 `greeting` 变量应该是一个字符串类型。

```js
assert.isString(greeting);
```

你应该将这句话赋值给 `greeting` 变量："你好呀！我的名字是[MathBot在这里]，我将会教会你有关 Math 对象的知识！"。记得要将 `[MathBot在这里]` 替换为 `botName` 变量。

```js
assert.strictEqual(greeting, "你好呀！我的名字是MathBot，我将会教会你有关 Math 对象的知识！");
```

你应该将 `greeting` 变量输出到控制台。

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```