---
id: 66ea845f54195c7020cc521b
title: 步骤 5
challengeType: 1
dashedName: step-5
---

# --description--

在两个值之间生成随机数的公式如下所示：

```js
Math.random() * (maximum - minimum) + minimum;
```

这将产生一个位于两个值之间的浮点数结果。

创建一个名为 `randomNum2` 的变量，并将生成介于 `min` 和 `max` 值之间的结果赋值给它。

然后，输出 `randomNum2` 变量以查看结果。试着刷新页面来看看不同的结果吧！

# --hints--

你应该有一个名为 `randomNum2` 的变量。

```js
assert.isNotNull(randomNum2);
```

你的 `randomNum2` 变量应该是一个数字。

```js
assert.isNumber(randomNum2);
```

你应该将生成介于 `min` 和 `max` 值之间的结果赋值给变量。确保不要硬编码值，要使用上面提供的公式哦。

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// 检查是否有人直接写整数
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

你应该将 `randomNum2` 变量输出到控制台。

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");

const randomNum = Math.random();
console.log(randomNum);

console.log("现在，让我们在两个数值之间生成一个随机数。");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```