---
id: 66ea7fc048a6e26e33c598fe
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

太棒了！现在是时候使用 `Math.random()` 方法来生成一个随机数啦！

创建一个名为 `randomNum` 的变量，并将调用 `Math.random()` 方法得到的结果赋值给它。

然后，将 `randomNum` 变量输出到控制台中。

试着在代码中添加一个空格来重新运行机器人，看看能生成哪些不同的随机数吧！

# --hints--

你应该声明一个叫做 `randomNum` 的变量。

```js
assert.isNotNull(randomNum);
```

你应该将调用 `Math.random()` 方法的结果赋值给 `randomNum` 变量。

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

你不应该直接给 `randomNum` 变量指定一个固定的随机数。请确保使用 `Math.random()` 方法来生成随机数。

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// 检查是否有人直接写整数
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

你应该将 `randomNum` 变量输出到控制台。

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `你好呀！我的名字是 ${botName}，我将会教会你有关 Math 对象的知识！`;

console.log(greeting);

console.log("Math.random() 方法会返回一个在 0 到小于 1 范围内的伪随机数。");
--fcc-editable-region--

--fcc-editable-region--
```