---
id: 66c8ba41a77db20a93f9d7a1
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个工作坊中，你将通过构建一个贷款资格检查器应用程序来复习JavaScript条件语句。太棒了！

该应用程序将根据用户的年收入和信用评分检查他们是否有资格获得联排别墅、汽车和公寓贷款。信用评分是一个数字，代表你管理借入资金的能力。

让我们开始吧，创建以下变量和值。继续加油！

| 变量名称             | 值   |
| ------------------------- | ------- |
| `minIncomeForDuplex`      | `60000` |
| `minCreditScoreForDuplex` | `700`   |
| `minIncomeForCondo`       | `45000` |
| `minCreditScoreForCondo`  | `680`   |
| `minIncomeForCar`         | `30000` |
| `minCreditScoreForCar`    | `650`   |

# --hints--

你应该创建一个设置为 `60000` 的 `minIncomeForDuplex` 变量。

```js
assert.equal(minIncomeForDuplex, 60000);
```

你应该创建一个设置为 `700` 的 `minCreditScoreForDuplex`。

```js
assert.equal(minCreditScoreForDuplex, 700);
```

你应该创建一个设置为 `30000` 的 `minIncomeForCar` 变量。

```js
assert.equal(minIncomeForCar, 30000);
```

你应该创建一个设置为 `650` 的 `minCreditScoreForCar` 变量。

```js
assert.equal(minCreditScoreForCar, 650);
```

你应该创建一个设置为 `45000` 的 `minIncomeForCondo` 变量。

```js
assert.equal(minIncomeForCondo, 45000);
```

你应该创建一个设置为 `680` 的 `minCreditScoreForCondo` 变量。

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```