---
id: 683ec95f2f8781355556ff83
title: 步骤 11
challengeType: 20
dashedName: step-11
---

# --description--

你可以看到第一次调用返回了 [3]。这是因为 [3] 是 [1, 2, 3, 4, 5] 中的中间值。

另一方面，第二次调用返回了一个空列表。这是因为检查中间值后，循环目前缺乏通过更新 [low] 或 [high] 变量来缩小搜索范围的逻辑。它只作为单次中点检查，然后就结束了，没有进一步搜索。

为了让二分查找能够继续缩小搜索范围，添加一个 [elif] 块，检查 [value] 是否大于 [value_at_middle]。现在在 [elif] 中添加 [pass] 关键字。

# --hints--

你应该为 [value] 大于 [value_at_middle] 的情况添加一个 [elif] 块。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```