---
id: 683ec95f2f8781355556ff7e
title: 步骤 6
challengeType: 20
dashedName: step-6
---

# --description--

现在你已经计算了 [mid] 索引，你需要实际从 [search_list] 中获取该索引处的值。这个值就是你将与参数 [value]（即你要搜索的目标）进行比较的值。

创建一个 [value_at_middle] 变量，并将它赋值为 [search_list] 中位于 [mid] 索引处的元素。

# --hints--

你应该创建一个 [value_at_middle] 变量并将其设置为 [search_list[mid]]。

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2


--fcc-editable-region--
```