---
id: 683ec95f2f8781355556ff7d
title: 步骤 5
challengeType: 20
dashedName: step-5
---

# --description--

要开始循环，你需要找到当前搜索空间的中间索引。

你可以通过计算 [low] 和 [high] 变量的平均值来获得中间索引。在将两个值相加后，你必须使用地板除法来获得平均值，这样答案总是向下取整的整数，因为索引必须是整数。

从你的 [while] 循环中删除 [pass] 关键字，并创建一个设置为 [low] 和 [high] 平均值的 [mid] 变量。

# --hints--

你应该删除 while 循环中的 [pass] 关键字，并创建一个设置为 [(low + high) // 2] 的 [mid] 变量。

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```