---
id: 683ec95f2f8781355556ff7c
title: 步骤 4
challengeType: 20
dashedName: step-4
---

# --description--

现在你已经定义了搜索的边界，是时候创建执行二分查找算法的循环了。

二分查找通过不断缩小搜索空间来工作。只要还有有效的元素范围可供检查，这个过程就会继续。

你可以用一个 [while] 循环来表达这一点，当 [low] 指针小于或等于 [high] 指针时循环会继续。

如果 [low] 变得大于 [high]，这意味着搜索空间已经变空了，列表中不存在 [value]。

在 [binary_search] 函数内部，创建一个条件为检查 [low] 是否小于或等于 [high] 的 [while] 循环。在循环内部，暂时添加 [pass] 关键字作为占位符。

# --hints--

你应该创建一个 [while] 循环，当 [low] 小于或等于 [high] 时运行。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--

--fcc-editable-region--
```