---
id: 684169b00a073c9b9f0a703c
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

你有没有注意到第二次函数调用现在在列表中有值 [3, 5, 4]？这表明二分查找正如预期的那样工作。

这就是它发生的经过：算法首先检查初始列表的中间值 [3]。由于 [4] 大于 [3]，搜索转移到右半部分。然后它检查 [5] 作为新的中间值。因为 [4] 小于 [5]，搜索向左移动，最终在最终范围内识别出 [4] 作为中间值。

为了再次测试函数，调用 [binary_search([1, 3, 5, 9, 14, 22], 10)] 并立即打印调用结果。这是一种找不到值的情况。

# --hints--

你应该调用函数 [binary_search([1, 3, 5, 9, 14, 22], 10)] 并立即打印它。

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```