---
id: 63ec14d1c216aa063f0be4af
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

在这个项目中，你将使用你已经学到的关于 HTML、CSS 和 JavaScript 的知识来构建一个购物车应用程序。你将实现添加商品到购物车、调整数量和删除商品等功能。

首先，你需要创建基本的 HTML 结构。在 `body` 元素内，添加一个 `header` 元素和一个 `main` 元素。

# --hints--

你应该有一个 `header` 元素。

```js
assert.exists(document.querySelector('header'));
```

你应该有一个 `main` 元素。

```js
assert.exists(document.querySelector('main'));
```

你的 `header` 元素应该在 `main` 元素之前。

```js
const header = document.querySelector('header');
const main = document.querySelector('main');
assert.isTrue(header.nextElementSibling === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    <script src="./script.js"></script>
  </body>
</html>
```