---
id: 63ee5d38a5d29d0696f8d820
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

现在你需要添加一些商品到页面上。在 `.products-container` 元素内，添加一个 `div` 元素，并给它一个 `class` 属性，值为 `product`。

# --hints--

你应该在 `.products-container` 元素内有一个 `div` 元素。

```js
assert.exists(document.querySelector('.products-container div'));
```

你的 `div` 元素应该有一个 `class` 属性，值为 `product`。

```js
assert.equal(document.querySelector('.products-container div')?.className, 'product');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        --fcc-editable-region--
        <div class="products-container"></div>
        --fcc-editable-region--
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        <div class="cart-total">
          <h3>总计: $0</h3>
        </div>
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```