---
id: 63ee5d8f9e7168076e932fe2
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

在 `.product` 元素内，添加一个 `h3` 元素，文本为 `西瓜`。

# --hints--

你应该在 `.product` 元素内有一个 `h3` 元素。

```js
assert.exists(document.querySelector('.product h3'));
```

你的 `h3` 元素应该有文本 `西瓜`。

```js
assert.equal(document.querySelector('.product h3')?.textContent, '西瓜');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container">
          --fcc-editable-region--
          <div class="product">

          </div>
          --fcc-editable-region--
        </div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        <div class="cart-total">
          <h3>总计: $0</h3>
        </div>
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```