---
id: 63f01c9791a0aa1751c73760
title: 步骤 29
challengeType: 0
dashedName: step-29
---

# --description--

在 `addItem` 函数内，声明一个 `product` 变量，并将其设置为 `products` 数组中 `id` 与传入的 `id` 参数匹配的第一个元素。

# --hints--

你应该声明一个 `product` 变量。

```js
assert.match(code, /var\s+product|let\s+product|const\s+product/);
```

你应该使用 `products.find()` 方法来查找匹配的商品。

```js
assert.match(code, /products\.find\(/);
```

你应该将 `product` 变量设置为 `products.find()` 的结果，查找 `id` 与传入的 `id` 参数匹配的商品。

```js
assert.match(code, /var\s+product\s*=\s*products\.find\(\s*\(?\s*item\s*\)?\s*=>\s*item\.id\s*===\s*id\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {

}
```