---
id: 63f0289df84a581bbdbd29b7
title: 步骤 33
challengeType: 0
dashedName: step-33
---

# --description--

在第二个按钮的事件监听器下方，为 `add-3` 按钮添加一个点击事件监听器，当按钮被点击时调用 `addItem(3)`。

# --hints--

你应该使用 `document.getElementById()` 获取 `add-3` 按钮。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-3\1\s*\)/);
```

你应该为按钮添加一个点击事件监听器。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-3\1\s*\)\.addEventListener\(\s*('|"|`)click\2/);
```

当按钮被点击时，你应该调用 `addItem(3)`。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-3\1\s*\)\.addEventListener\(\s*('|"|`)click\2\s*,\s*\(\s*\)\s*=>\s*{\s*addItem\(\s*3\s*\)\s*}\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

document.getElementById("add-1").addEventListener("click", () => {
  addItem(1);
});

document.getElementById("add-2").addEventListener("click", () => {
  addItem(2);
});

```