---
id: 63f01861f813e01564c95315
title: 步骤 27
challengeType: 0
dashedName: step-27
---

# --description--

在 `cart` 变量下方，声明一个 `products` 变量，并将其初始化为一个包含三个对象的数组。每个对象应该有 `id`、`name` 和 `price` 属性。第一个对象的 `id` 应该是 `1`，`name` 应该是 `西瓜`，`price` 应该是 `12.99`。第二个对象的 `id` 应该是 `2`，`name` 应该是 `香蕉`，`price` 应该是 `0.99`。第三个对象的 `id` 应该是 `3`，`name` 应该是 `橙子`，`price` 应该是 `0.89`。

# --hints--

你应该声明一个 `products` 变量。

```js
assert.match(code, /var\s+products|let\s+products|const\s+products/);
```

你的 `products` 变量应该初始化为一个包含三个对象的数组。

```js
assert.equal(products?.length, 3);
```

第一个对象应该有 `id` 为 `1`，`name` 为 `西瓜`，`price` 为 `12.99`。

```js
assert.equal(products?.[0]?.id, 1);
assert.equal(products?.[0]?.name, '西瓜');
assert.equal(products?.[0]?.price, 12.99);
```

第二个对象应该有 `id` 为 `2`，`name` 为 `香蕉`，`price` 为 `0.99`。

```js
assert.equal(products?.[1]?.id, 2);
assert.equal(products?.[1]?.name, '香蕉');
assert.equal(products?.[1]?.price, 0.99);
```

第三个对象应该有 `id` 为 `3`，`name` 为 `橙子`，`price` 为 `0.89`。

```js
assert.equal(products?.[2]?.id, 3);
assert.equal(products?.[2]?.name, '橙子');
assert.equal(products?.[2]?.price, 0.89);
```

# --seed--

## --seed-contents--

```js
const cart = [];

```