---
id: 63f034d012f74627ce538d3a
title: 步骤 44
challengeType: 0
dashedName: step-44
---

# --description--

在设置 `li` 元素的 `textContent` 的代码下方，使用 `document.getElementById()` 获取 `cart-items` 元素，并将其 `appendChild` 设置为 `li` 元素。这将把新的 `li` 元素添加到购物车中。

# --hints--

你应该使用 `document.getElementById()` 获取 `cart-items` 元素。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)cart-items\1\s*\)/);
```

你应该将 `li` 元素添加到 `cart-items` 元素中。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)cart-items\1\s*\)\.appendChild\(\s*li\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

  for (let i = 0; i < cart.length; i++) {
    const item = cart[i];
    const li = document.createElement("li");
    li.textContent = `${item.name} - $${item.price}`;

  }

}

```