---
id: 63f0284532742c1b26c7a052
title: 步骤 32
challengeType: 0
dashedName: step-32
---

# --description--

在第一个按钮的事件监听器下方，为 `add-2` 按钮添加一个点击事件监听器，当按钮被点击时调用 `addItem(2)`。

# --hints--

你应该使用 `document.getElementById()` 获取 `add-2` 按钮。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-2\1\s*\)/);
```

你应该为按钮添加一个点击事件监听器。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-2\1\s*\)\.addEventListener\(\s*('|"|`)click\2/);
```

当按钮被点击时，你应该调用 `addItem(2)`。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-2\1\s*\)\.addEventListener\(\s*('|"|`)click\2\s*,\s*\(\s*\)\s*=>\s*{\s*addItem\(\s*2\s*\)\s*}\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

document.getElementById("add-1").addEventListener("click", () => {
  addItem(1);
});

```