---
id: 63f036ec91fdf238c90665f5
title: 步骤 46
challengeType: 0
dashedName: step-46
---

# --description--

现在你需要为购物车中的每个商品添加一个删除按钮。在 `displayCart` 函数中，在设置 `li` 元素的 `textContent` 的代码下方，声明一个 `button` 变量，并将其设置为一个新的 `button` 元素。

# --hints--

你应该声明一个 `button` 变量。

```js
assert.match(code, /var\s+button|let\s+button|const\s+button/);
```

你应该将 `button` 变量设置为一个新的 `button` 元素。

```js
assert.match(code, /var\s+button\s*=\s*document\.createElement\(\s*('|"|`)button\1\s*\)|let\s+button\s*=\s*document\.createElement\(\s*('|"|`)button\1\s*\)|const\s+button\s*=\s*document\.createElement\(\s*('|"|`)button\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
  displayCart();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

  for (let i = 0; i < cart.length; i++) {
    const item = cart[i];
    const li = document.createElement("li");
    li.textContent = `${item.name} - $${item.price}`;

  }

}

```