---
id: 63f026d041bc6c1a3d5cba0f
title: 步骤 31
challengeType: 0
dashedName: step-31
---

# --description--

现在你需要添加一些代码来处理按钮点击事件。在 `addItem` 函数下方，使用 `document.getElementById()` 获取 `add-1` 按钮，并为其添加一个点击事件监听器，当按钮被点击时调用 `addItem(1)`。

# --hints--

你应该使用 `document.getElementById()` 获取 `add-1` 按钮。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-1\1\s*\)/);
```

你应该为按钮添加一个点击事件监听器。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-1\1\s*\)\.addEventListener\(\s*('|"|`)click\2/);
```

当按钮被点击时，你应该调用 `addItem(1)`。

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)add-1\1\s*\)\.addEventListener\(\s*('|"|`)click\2\s*,\s*\(\s*\)\s*=>\s*{\s*addItem\(\s*1\s*\)\s*}\s*\)/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
}

```