---
id: 669fd6fd12918e3de87854d4
title: æ­¥éª¤ 8
challengeType: 0
dashedName: step-8
---

# --description--

ç”¨æˆ·åº”è¯¥æœ‰åŠæ³•æµè§ˆä½ çš„åšå®¢é¡µé¢ä¸Šçš„ä¸åŒéƒ¨åˆ†ã€?

`nav` å…ƒç´ ç”¨äºæä¾›æŒ‡å‘æ–‡æ¡£ä¸­å…¶ä»–éƒ¨åˆ†æˆ–ç½‘ç«™ä¸­å…¶ä»–éƒ¨åˆ†çš„å¯¼èˆªé“¾æ¥ã€?ä½ ç»å¸¸ä¼šçœ‹åˆ°ç”¨äºèœå•æˆ–ç›®å½•çš„ `nav` å…ƒç´ ã€?

ä¸‹é¢æ˜¯ä½¿ç”?`nav` å…ƒç´ çš„ç¤ºä¾‹ï¼š

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

åœ?`figure` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`nav` å…ƒç´ ï¼Œåœ¨å…¶ä¸­åµŒå¥—ä¸€ä¸?`ul` å…ƒç´ ã€?

åœ?`ul` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸‰ä¸ª `li` å…ƒç´ ã€?

# --hints--

ä½ çš„ `nav` å…ƒç´ åº”æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<nav\>/);
```

ä½ çš„ `nav` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/nav\>/);
```

ä½ çš„ `nav` å…ƒç´ åº”è¯¥åœ?`header` å…ƒç´ ä¸­ã€?

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

ä½ çš„ `ul` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<ul\>/);
```

ä½ çš„ `ul` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/ul\>/);
```

ä½ çš„ `ul` å…ƒç´ åº”åœ¨ `nav` å…ƒç´ ä¸­ã€?

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

ä½ çš„ `ul` å…ƒç´ åº”åŒ…å«ä¸‰ä¸?`li` å…ƒç´ ã€?

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```

