---
id: 669fde2081f65141ad703fe4
title: æ­¥éª¤ 11
challengeType: 0
dashedName: step-11
---

# --description--

é¡µé¢çš„ç¬¬ä¸€ä¸ªéƒ¨åˆ†å°†æ˜¯å…³äºéƒ¨åˆ†ã€?è¯¥éƒ¨åˆ†å°†ä»‹ç» Whiskers å…ˆç”Ÿï¼Œè®©ç”¨æˆ·äº†è§£æœ¬åšå®¢çš„å†…å®¹ã€?

åœ¨ä½ çš?`main` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`section` å…ƒç´ ï¼Œå…¶ `id` å±æ€§è®¾ç½®ä¸º `"about"`ã€?

åœ?`section` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`h2`ï¼Œå…¶æ–‡æœ¬ä¸?`About`ã€?

# --hints--

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<section.*>/);
```

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/section\>/);
```

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`id`ï¼Œè®¾ç½®ä¸º `"about"`ã€?

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

ä½ çš„ `section` å…ƒç´ åº”è¯¥åœ?`main` å…ƒç´ ä¸­ã€?

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

ä½ çš„ `h2` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<h2\>/);
```

ä½ çš„ `h2` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/h2\>/);
```

ä½ çš„ `h2` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`About`ã€?ä»”ç»†æ£€æŸ¥ä½ çš„æ‹¼å†™ã€?

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

ä½ çš„ `h2` å…ƒç´ åº”è¯¥åœ?`section` å…ƒç´ ä¸­ã€?

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>
```

