---
id: 669fd003cb89ee3c2402e041
title: æ­¥éª¤ 7
challengeType: 0
dashedName: step-7
---

# --description--

åœ¨è¿™ä¸ªä»‹ç»æ€§å†…å®¹ä¸­ï¼Œéœ€è¦å±•ç¤ºä¸€å¼?Whiskers å…ˆç”Ÿçš„å›¾ç‰‡ï¼Œå¹¶é…ä¸Šè¯´æ˜ã€?

åœ?`h1` å…ƒç´ ä¸‹æ–¹ï¼Œé¦–å…ˆæ·»åŠ ä¸€ä¸?`figure` å…ƒç´ ã€?

åœ?`figure` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`img` å…ƒç´ ã€?

`img` çš?`src` å±æ€§å€¼åº”ä¸?`"https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"`ï¼Œ`alt` æ–‡æœ¬åº”ä¸º `"a cat in the garden"`ã€?

åœ¨ä½ çš?`img` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`figcaption</code`ï¼Œæ–‡æœ¬ä¸º `Mr. Whiskers in the Garden`ã€?

# --hints--

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<figure\>/);
```

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/figure\>/);
```

ä½ çš„ `figure` å…ƒç´ åº”è¯¥åœ?`header` å…ƒç´ ä¸­ã€?

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`img` å…ƒç´ ã€?

```js
assert.exists(document.querySelector("img"));
```

ä½ çš„ `img` å…ƒç´ åº”åœ¨ `figure` å…ƒç´ ä¸­ã€?

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

ä½ çš„ `img` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`src` å±æ€§ï¼Œè®¾ç½®ä¸?`"https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"`ã€?

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg');
```

ä½ çš„ `img` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`alt` å±æ€§ï¼Œè®¾ç½®ä¸?`"a cat in the garden"`ã€?

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

ä½ çš„ `figcaption` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<figcaption\>/);
```

ä½ çš„ `figcaption` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/figcaption\>/);
```

ä½ çš„ `figcaption` å…ƒç´ åº”åœ¨ `figure` å…ƒç´ ä¸­ã€?

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

ä½ çš„ `figcaption` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Mr. Whiskers in the Garden`ã€?

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

ä½ çš„ `figcaption` åº”åœ¨ `img` å…ƒç´ ä¸‹æ–¹ã€?ä½ æŠŠå®ƒä»¬çš„é¡ºåºå†™é”™äº†ã€?

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>

    </header>
  </body>
--fcc-editable-region--
</html>
```

