---
id: 669fdd3965f36f40af9615e5
title: æ­¥éª¤ 10
challengeType: 0
dashedName: step-10
---

# --description--

ç°åœ¨ä½ å·²ç»å®Œæˆäº†é¡µå¤´çš„åˆ¶ä½œï¼Œéœ€è¦å¼€å§‹æ·»åŠ ä¸»ä½“å†…å®¹ã€?

åœ¨ä½ çš?`header` ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`main` å…ƒç´ ã€?

# --hints--

ä½ çš„ `main` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<main\>/);
```

ä½ çš„ `main` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/main\>/);
```

ä½ çš„ `main` å…ƒç´ åº”åœ¨ `header` å…ƒç´ ä¸‹æ–¹ã€?

```js
assert.match(code, /<\/header>[.\n\s]*<main>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```

