---
id: 669fdc11c9b0263fe0814a7a
title: æ­¥éª¤ 9
challengeType: 0
dashedName: step-9
---

# --description--

åœ¨æ¯ä¸?`li` å…ƒç´ ä¸­ï¼Œéƒ½éœ€è¦æœ‰ä¸€ä¸ªé”šå…ƒç´ ã€?

ç¬¬ä¸€ä¸ªé”šå…ƒç´ çš„æ–‡æœ¬åº”ä¸?`About`ï¼Œ`href` å±æ€§å€¼åº”ä¸?`"#about"`ã€?`about` å‰é¢çš„å“ˆå¸Œç¬¦å·ä»£è¡¨ä¸€ä¸?`id` åç§°ï¼Œç¨åå°†æ·»åŠ åˆ°é¡¹ç›®ä¸­ã€?

ç¬¬äºŒä¸ªé”šå…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Posts`ï¼Œ`href` å±æ€§å€¼åº”ä¸?`"#posts"`ã€?

ç¬¬ä¸‰ä¸ªé”šå…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Contact`ï¼Œ`href` å±æ€§å€¼åº”ä¸?`"#contact"`ã€?

# --hints--

ä½ çš„ç¬¬ä¸€ä¸ªé”šå…ƒç´ åº”åœ¨ç¬¬ä¸€ä¸?`li` å…ƒç´ ä¸­åŒ…å?`About` æ–‡æœ¬ã€?

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

ä½ çš„ç¬¬ä¸€ä¸ªé”šå…ƒç´ åº”å…·æœ?`href` å±æ€§ï¼Œè®¾ç½®ä¸?`"#about"`ã€?

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

ä½ çš„ç¬¬äºŒä¸ªé”šå…ƒç´ åº”åœ¨ç¬¬äºŒä¸?`li` å…ƒç´ ä¸­åŒ…å?`Posts` æ–‡æœ¬ã€?

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

ä½ çš„ç¬¬äºŒä¸ªé”šå…ƒç´ åº”å…·æœ?`href` å±æ€§ï¼Œè®¾ç½®ä¸?`"#posts"`ã€?

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

ä½ çš„ç¬¬ä¸‰ä¸ªé”šå…ƒç´ åº”åœ¨ç¬¬ä¸‰ä¸?`li` å…ƒç´ ä¸­åŒ…å?`Contact` æ–‡æœ¬ã€?

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

ä½ çš„ç¬¬ä¸‰ä¸ªé”šå…ƒç´ åº”å…·æœ?`href` å±æ€§ï¼Œè®¾ç½®ä¸?`"#contact"`ã€?

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--
      <nav>
        <ul>
          <li></li>
          <li></li>
          <li></li>
        </ul>
      </nav>
--fcc-editable-region--
    </header>
  </body>
</html>
```

