---
id: 669fdfc9a5e3da42d2376609
title: æ­¥éª¤ 13
challengeType: 0
dashedName: step-13
---

# --description--

ç°åœ¨ä½ å·²æ·»åŠ äº†å…³äºéƒ¨åˆ†ï¼Œè¯·å°è¯•ç‚¹å‡?`About` é“¾æ¥ï¼ŒæŸ¥çœ‹é¡µé¢æ˜¯å¦èƒ½è·³è½¬åˆ°è¯¥éƒ¨åˆ†ã€?

åšå®¢é¡µé¢çš„ä¸‹ä¸€éƒ¨åˆ†å°†æ˜¯æœ‰å…³ Whiskers å…ˆç”Ÿçš„åšæ–‡åˆ—è¡¨ã€?

æ·»åŠ å¦ä¸€ä¸?`section` å…ƒç´ ï¼Œå…¶ `id` è®¾ç½®ä¸?`"posts"`ã€?

åœ?`section` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`h2` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Posts`ã€?

# --hints--

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`section` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<section.*>/i);
```

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`section` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/section>/i);
```

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`id`ï¼Œè®¾ç½®ä¸º `"posts"`ã€?

```js
assert.match(code, /<section.*id="posts".*>/i);
```

ä½ çš„ `section` å…ƒç´ åº”è¯¥åœ?`main` å…ƒç´ ä¸­ã€?

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

ä½ çš„ `h2` å…ƒç´ åº”è¯¥åœ?`section` å…ƒç´ ä¸­ã€?

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

ä½ çš„ `h2` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Posts`ã€?

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
  </body>
</html>
```

