---
id: 66a338be7f1dee383a0e0ecb
title: æ­¥éª¤ 16
challengeType: 0
dashedName: step-16
---

# --description--

å¯¹äºŽç¬¬äºŒç¯‡åšæ–‡ï¼Œéœ€è¦æ·»åŠ å¦ä¸€ä¸?`article` å…ƒç´ ã€?

åœ?`article` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`h3` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Mr. Whiskers' First Bath`ã€?

åœ?`h3` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸¤æ®?lorem ipsum æ–‡æœ¬ã€?

# --hints--

åœ?`section` å…ƒç´ ä¸­ï¼Œä½ åº”è¯¥æœ‰ç¬¬äºŒä¸?`article` å…ƒç´ ï¼Œ`id` è®¾ç½®ä¸?`posts`ã€?

```js
const articleElement = document.querySelector('#posts article:nth-of-type(2)');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

åœ?`article` å…ƒç´ ä¸­ï¼Œä½ åº”è¯¥æœ‰ä¸€ä¸?`h3` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Mr. Whiskers' First Bath`ã€?

```js
const h3Element = document.querySelector('#posts article:nth-of-type(2) h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Bath");
```

ä½ çš„ `h3` å…ƒç´ åº”åœ¨ `article` å…ƒç´ ä¸­ã€?

```js
const h3Element = document.querySelector('#posts article:nth-of-type(2) h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

ä½ åº”è¯¥åœ¨ `article` å…ƒç´ ä¸­æ·»åŠ ä¸¤ä¸ªæ®µè½å…ƒç´ ã€?

```js
const paragraphElements = document.querySelectorAll('#posts article:nth-of-type(2) p');
assert.strictEqual(paragraphElements.length, 2);
```

ç¬¬ä¸€ä¸ªæ®µå…ƒç´ åº”ä½¿ç”?lorem ipsum æ–‡å­—ã€?ä»”ç»†æ£€æŸ¥æ‹¼å†™å’Œæ ‡ç‚¹ç¬¦å·æ˜¯å¦æœ‰è¯¯ã€?

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(2) p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

ç¬¬äºŒä¸ªæ®µå…ƒç´ åº”ä½¿ç”?lorem ipsum æ–‡å­—ã€?ä»”ç»†æ£€æŸ¥æ‹¼å†™å’Œæ ‡ç‚¹ç¬¦å·æ˜¯å¦æœ‰è¯¯ã€?

```js
const paragraphElement = document.querySelector('#posts article:nth-of-type(2) p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim(), loremIpsum);
```

ä½ çš„æ®µè½åº”ä½äº?`h3` å…ƒç´ ä¹‹åŽã€?

```js
let h3 = document.querySelectorAll('#posts article:nth-of-type(2) > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('#posts article:nth-of-type(2) > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>

        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        --fcc-editable-region--

        --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

