---
id: 66a33c449ffdf93b90c5821e
title: æ­¥éª¤ 20
challengeType: 0
dashedName: step-20
---

# --description--

åœ¨è”ç³»æ–¹å¼éƒ¨åˆ†ï¼Œéœ€è¦æ˜¾ç¤ºåšå®¢ä½œè€…çš„è”ç³»ä¿¡æ¯ã€?ä¸ºæ­¤ï¼Œä½ å°†ä½¿ç”?<dfn>address</dfn> å…ƒç´ ã€?

`address` å…ƒç´ ç”¨äºè¡¨ç¤ºä¸ªäººæˆ–ç»„ç»‡çš„è”ç³»ä¿¡æ¯ã€?

ä¸‹é¢æ˜¯ä¸€ä¸ªä½¿ç”?`address` å…ƒç´ è¡¨ç¤ºç‰©ç†åœ°å€çš„ç¤ºä¾‹ã€?æ­¤å¤„ä½¿ç”¨ `br` å…ƒç´ åœ¨æ–‡æœ¬ä¹‹é—´åˆ›å»ºæ¢è¡Œç¬¦ã€?

```html
<address>
  1234 Make Believe Lane <br />
  Pretend City, USA
</address>
```

åœ¨çš„ `h2` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`address` å…ƒç´ ã€?

# --hints--

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`address` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<address>/i);
```

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`address` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/address>/i);
```

ä½ çš„ `address` å…ƒç´ åº”è¯¥åœ?`section` å…ƒç´ ä¸­ã€?

```js
const addressElement = document.querySelector('footer #contact address');
assert.strictEqual(addressElement?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>

        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        <article>
          <h3>Mr. Whiskers' First Bath</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        <article>
          <h3>Mr. Whiskers' First Birthday Party</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
      </section>
    </main>
--fcc-editable-region--
    <footer>
      <section id="contact">
        <h2>Contact</h2>

      </section>
    </footer>
--fcc-editable-region--
  </body>
</html>
```

