---
id: 66a334a1a7cca6354999f9bf
title: æ­¥éª¤ 14
challengeType: 0
dashedName: step-14
---

# --description--

å¯¹äºç¬¬ä¸€ç¯‡åšæ–‡ï¼Œå°†ä½¿ç”?<dfn>article</dfn> å…ƒç´ ã€?

`article` å…ƒç´ è¡¨ç¤ºç½‘é¡µä¸­çš„è‡ªåŒ…å«å†…å®¹ã€?

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

åœ?`h2` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`article` å…ƒç´ ã€?

åœ?`article` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`h3` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Mr. Whiskers' First Day Home`ã€?

è¿™é‡Œä½¿ç”¨ `h3` çš„åŸå› æ˜¯ï¼Œä¸ºæ ‡é¢˜å…ƒç´ ä¿æŒé€‚å½“çš„ç»“æ„å±‚çº§éå¸¸é‡è¦ã€?ç”±äºåšæ–‡å‰¯æ ‡é¢˜æ˜¯ä¸€ä¸?`h2` å…ƒç´ ï¼Œå±‚çº§ç»“æ„ä¸­çš„ä¸‹ä¸€çº§æ˜¯ `h3` ã€?

# --hints--

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`article` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<article>/i);
```

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`article` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/article>/i);
```

ä½ çš„ `article` å…ƒç´ åº”åœ¨ `section` å…ƒç´ ä¸­ï¼Œ`id` è®¾ç½®ä¸?`posts`ã€?

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`h3` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<h3>/i);
```

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`h3` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/h3>/i);
```

ä½ çš„ `h3` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Mr. Whiskers' First Day Home`ã€?

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

ä½ çš„ `h3` å…ƒç´ åº”åœ¨ `article` å…ƒç´ ä¸­ã€?

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
--fcc-editable-region--
      <section id="posts">
        <h2>Posts</h2>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

