---
id: 66a335e5888ffb367633200a
title: æ­¥éª¤ 15
challengeType: 0
dashedName: step-15
---

# --description--

æœ¬åšæ–‡å°†åŒ…å«å‡ æ®µ <dfn>lorem ipsum</dfn> æ–‡æœ¬ã€?

Lorem ipsum é€šå¸¸åœ¨ç½‘ç»œå¼€å‘ç”¨ä½œå ä½ç¬¦æ–‡æœ¬ã€?å½“ä½ æƒ³é›†ä¸­ç²¾åŠ›æ„å»ºç½‘é¡µçš„åŸºæœ¬ç»“æ„ï¼Œè€Œä¸å¿…æ‹…å¿ƒå®é™…å†…å®¹æ—¶ï¼Œå®ƒå°±éå¸¸æœ‰ç”¨ã€?

ä¸‹é¢æ˜¯ä¸€ä¸ªä½¿ç”?lorem ipsum çš„ä¾‹å­ï¼š

```html
<p>
  Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
  quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
  quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
  quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
</p>
```

åœ?`h3` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸¤æ®?lorem ipsum æ–‡æœ¬ã€?

# --hints--

ä½ åº”è¯¥åœ¨ `article` å…ƒç´ ä¸­æ·»åŠ ä¸¤ä¸ªæ®µè½å…ƒç´ ã€?

```js
const paragraphElements = document.querySelectorAll('#posts article p');
assert.strictEqual(paragraphElements.length, 2);
```

ç¬¬ä¸€ä¸ªæ®µå…ƒç´ åº”ä½¿ç”?lorem ipsum æ–‡å­—ã€?ä»”ç»†æ£€æŸ¥æ‹¼å†™å’Œæ ‡ç‚¹ç¬¦å·æ˜¯å¦æœ‰è¯¯ã€?

```js
const paragraphElement = document.querySelector('article p:first-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

ç¬¬äºŒä¸ªæ®µå…ƒç´ åº”ä½¿ç”?lorem ipsum æ–‡å­—ã€?ä»”ç»†æ£€æŸ¥æ‹¼å†™å’Œæ ‡ç‚¹ç¬¦å·æ˜¯å¦æœ‰è¯¯ã€?

```js
const paragraphElement = document.querySelector('article p:last-of-type');
const loremIpsum = `Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas. Quisquam quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.`

assert.strictEqual(paragraphElement?.innerText.trim().replaceAll(/\s+/g, ' '), loremIpsum);
```

ä½ çš„æ®µè½åº”ä½äº?`h3` å…ƒç´ ä¹‹åã€?

```js
let h3 = document.querySelectorAll('article > p ~ h3');
assert.isEmpty(h3);
const paragraphs = document.querySelectorAll('article > h3 ~ p');
assert.lengthOf(paragraphs, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>
        --fcc-editable-region--
        <article>
          <h3>Mr. Whiskers' First Day Home</h3>

        </article>
        --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

