---
id: 66a7e72adf226c02626715a3
title: æ­¥éª¤ 12
challengeType: 0
dashedName: step-12
---

# --description--

åœ?`h2` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸ªæ®µè½å…ƒç´ ï¼Œæ–‡æœ¬ä¸?`Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.`ã€?

åœ¨æ®µè½å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ å¦ä¸€ä¸ªæ®µè½å…ƒç´ ï¼Œæ–‡æœ¬ä¸?`His playful nature and boundless energy keeps me on my toes. I love him so much.`

# --hints--

ç¬¬ä¸€ä¸ªæ®µè½çš„æ–‡æœ¬åº”ä¸º `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.` ä»”ç»†æ£€æŸ¥ä½ çš„æ‹¼å†™ã€?

```js
const pElement = document.querySelector('body p:first-of-type');
const text = "Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers."
assert.strictEqual(pElement?.innerText.trim(), text);
```

ä½ çš„ç¬¬ä¸€ä¸ªæ®µè½åº”åµŒå¥—åœ?`section` å…ƒç´ å†…ã€?

```js
const pElement = document.querySelector('#about p:first-of-type');
assert.strictEqual(pElement?.parentElement.tagName, 'SECTION');
```

ä½ åº”è¯¥åœ¨ `section` å…ƒç´ ä¸­æ·»åŠ ç¬¬äºŒä¸ªæ®µè½ã€?

```js
const pElement = document.querySelectorAll('#about p');
assert.strictEqual(pElement?.length, 2);
```

ç¬¬äºŒä¸ªæ®µè½çš„æ–‡æœ¬åº”ä¸º `His playful nature and boundless energy keeps me on my toes. I love him so much.` ä»”ç»†æ£€æŸ¥ä½ çš„æ‹¼å†™ã€?

```js
const pElement = document.querySelector('#about p:last-of-type');
const text = "His playful nature and boundless energy keeps me on my toes. I love him so much."
assert.strictEqual(pElement?.innerText.trim(), text);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      --fcc-editable-region--
      <section id="about">
        <h2>About</h2>

      </section>
      --fcc-editable-region--
    </main>
  </body>
</html>
```

