---
id: a97fd23d9b809dac9921074f
title: Build an Optional Arguments Sum Function
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

在这个实验中，您将构建一个最多接受两个参数并求和的函数，但如果只提供一个参数，则返回一个等待第二个数字求和的函数。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 创建一个[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)函数。
2. 如果函数接收到两个参数，函数应该返回两个参数的和。
3. 如果只提供一个参数，[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)应该返回一个函数。
   - 当返回的函数用单个参数调用时，它应该返回总和。

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. 如果任一参数不是有效数字，返回[undefined](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L104-L104)。

# --hints--

您应该有[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)函数。

```js
assert.isFunction(addTogether);
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)(2, 3)应该返回5。

```js
assert.deepEqual(addTogether(2, 3), 5);
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)(23.4, 30)应该返回53.4。

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)("2", 3)应该返回[undefined](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L104-L104)。

```js
assert.isUndefined(addTogether('2', 3));
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)(5, undefined)应该返回[undefined](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L104-L104)。

```js
assert.isUndefined(addTogether(5, undefined));
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)("https://www.youtube.com/watch?v=dQw4w9WgXcQ") 应该返回 [undefined](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L104-L104)。

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)(5)应该返回一个函数。

```js
assert.isFunction(addTogether(5));
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)(5)(7)应该返回12。

```js
assert.deepEqual(addTogether(5)(7), 12);
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)(2)([3]) 应返回 [undefined](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L104-L104)。

```js
assert.isUndefined(addTogether(2)([3]));
```

[addTogether](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L101-L119)(2, "3")应该返回[undefined](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-optional-arguments-sum-function/a97fd23d9b809dac9921074f.md#L104-L104)。

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}