---
id: 67116c175a77cb64c5c10d49
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在之前的课程中，你学会了如何使用高阶函数，例如 [map](#)、[filter](#) 和 [reduce](#)。

在这个工作坊中，你将构建一个图书馆管理应用程序，这将给你一个练习使用不同高阶函数的机会。

首先，创建一个名为 [library](#) 的变量并给它分配一个空数组。

# --hints--

你应该创建一个 [library](#) 数组。

```js
assert.isArray(library)
```

你的 [library](#) 数组应该是空的。

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```