---
id: 67116d7584d0b469b14579bf
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

在 [library](#) 数组内，创建一个具有以下属性和值的对象：

| 属性     | 值                                                            |
| -------- | ------------------------------------------------------------- |
| [title](#)  | `"Your Next Five Moves: Master the Art of Business Strategy"` |
| [author](#) | `"Patrick Bet-David and Greg Dinkin"`                         |
| [about](#)  | `"一本关于如何提前规划的书"`                                  |
| [pages](#)  | `320`                                                         |

# --hints--

你的 [library](#) 数组内应该有一个对象。

```js
assert.isObject(library[0])
```

你的对象应该有 [title](#) 属性。

```js
assert.property(library[0], "title")
```

你的 [title](#) 属性的值应该是 `"Your Next Five Moves: Master the Art of Business Strategy"`。

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

你的对象应该有 [author](#) 属性。

```js
assert.property(library[0], "author")
```

你的 [author](#) 属性应该有值 `"Patrick Bet-David and Greg Dinkin"`。

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

你的对象应该有 [about](#) 属性。

```js
assert.property(library[0], "about")
```

你的 [about](#) 属性应该有值 `"一本关于如何提前规划的书"`。

```js
assert.propertyVal(library[0], "about", "一本关于如何提前规划的书");
```

你的对象应该有 [pages](#) 属性。

```js
assert.property(library[0], "pages")
```

你的 [pages](#) 属性的值应该是一个数字。

```js
assert.isNumber(library[0]?.pages)
```

你的 [pages](#) 属性应该设置为 `320`。

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--

--fcc-editable-region--
];
```