---
id: 67116d7584d0b469b14579c0
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在 [library](#) 数组内创建另一个具有以下属性和值的对象：

| 属性     | 值                                                                      |
| -------- | ----------------------------------------------------------------------- |
| [title](#)  | `"Atomic Habits"`                                                       |
| [author](#) | `"James Clear"`                                                         |
| [about](#)  | `"一本实用的书，关于摒弃坏习惯并培养好习惯"`                            |
| [pages](#)  | `320`                                                                   |

# --hints--

你的 [library](#) 数组第二项应该是一个对象。

```js
assert.isObject(library[1]);
```

你的对象应该有 [title](#) 属性。

```js
assert.property(library[1], "title")
```

你的 [title](#) 属性应该有值 `"Atomic Habits"`。

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

你的对象应该有 [author](#) 属性。

```js
assert.property(library[1], "author")
```

你的 [author](#) 属性应该有值 `"James Clear"`。

```js
assert.propertyVal(library[1], "author", "James Clear");
```

你的对象应该有 [about](#) 属性。

```js
assert.property(library[1], "about")
```

你的 [about](#) 属性应该有值 `"一本实用的书，关于摒弃坏习惯并培养好习惯"`。

```js
assert.propertyVal(library[1], "about", "一本实用的书，关于摒弃坏习惯并培养好习惯");
```

你的对象应该有 [pages](#) 属性。

```js
assert.property(library[1], "pages")
```

你的 [pages](#) 属性的值应该是一个数字。

```js
assert.isNumber(library[1]?.pages)
```

你的 [pages](#) 属性应该设置为 `320`。

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--

--fcc-editable-region--
];
```