---
id: 67116d7584d0b469b14579c3
title: 步骤 14
challengeType: 1
dashedName: step-14
---

# --description--

在之前的课程中，你学会了如何使用 [filter()](#) 方法，该方法用于返回过滤结果的新数组。

这是一个示例：

```js
const developers = [
  { name: "Alice", city: "New York", age: 30 },
  { name: "Bob", city: "San Francisco", age: 25 },
  { name: "Charlie", city: "New York", age: 35 },
  { name: "Diana", city: "Chicago", age: 28 }
];

const newYorkCityDevelopers = developers.filter(dev => dev.city === "New York");

console.log(newYorkCityDevelopers);
/*
[
  { name: "Alice", city: "New York", age: 30 },
  { name: "Charlie", city: "New York", age: 35 }
]
*/
```

创建一个 [getBooksByAuthor](#) 函数，带有两个参数 - 包含图书对象的数组和包含作者的字符串。

该函数必须返回一个包含特定作者书籍的数组。

# --hints--

你应该创建一个 [getBooksByAuthor](#) 函数。

```js
assert.isFunction(getBooksByAuthor);
```

你的 [getBooksByAuthor](#) 函数应该有两个参数。

```js
assert.lengthOf(getBooksByAuthor, 2);
```

你的 [getBooksByAuthor](#) 函数应该返回一个数组。

```js
assert.isArray(getBooksByAuthor(library, ''));
```

你的 [getBooksByAuthor](#) 函数应该为任何作者返回正确数量的书籍。

```js
const arvidKahlBooks = getBooksByAuthor(library, 'Arvid Kahl');
const robertAndSharonBook = getBooksByAuthor(library, 'Robert Kiyosaki and Sharon Lechter');
const pbdAndGDBooks = getBooksByAuthor(library, 'Patrick Bet-David and Greg Dinkin');
const pbdBook = getBooksByAuthor(library, 'Patrick Bet-David');
const johnGordon = getBooksByAuthor(library, 'Jon Gordon');
const jamesClearBook = getBooksByAuthor(library, 'James Clear');
const JDAndSDBook = getBooksByAuthor(library, 'Jeff DeGraff and Staney DeGraff');

assert.lengthOf(arvidKahlBooks, 2);
assert.lengthOf(robertAndSharonBook, 1);
assert.lengthOf(pbdAndGDBooks, 1);
assert.lengthOf(pbdBook, 1);
assert.lengthOf(johnGordon, 1);
assert.lengthOf(jamesClearBook, 1);
assert.lengthOf(JDAndSDBook, 1);

const _differentLibrary = [
  {
    title: 'Title A',
    author: 'Author A',
    about: 'About A',
    pages: 320,
  },
  {
    title: 'Title B',
    author: 'Author B',
    about: 'About B',
    pages: 320,
  },
  {
    title: 'Title C',
    author: 'Author B',
    about: 'About C',
    pages: 304,
  },
];
assert.lengthOf(getBooksByAuthor(_differentLibrary, 'Author B'), 2);
assert.lengthOf(getBooksByAuthor(_differentLibrary, 'Author A'), 1);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: '一本关于如何提前规划的书',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: '一本实用的书，关于摒弃坏习惯并培养好习惯',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "一本书强调识别和理解对手的重要性，以在商业世界中取得成功",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: '一本专注于如何建立受众驱动型企业的书',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: '一本关于有效方法引领咖啡豆生活方式的书',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: '一本关于如何培养创造力和创新技能的书',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: '一本关于财务素养、财务独立和创造财富的书。',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: '一本关于如何启动企业的书',
    pages: 500,
  },
];

console.log("图书馆中的书籍：\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\n图书摘要列表：\n");

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about).join("\n");
}

console.log(getBookSummaries(library));

console.log("\nArvid Kahl的书籍列表：\n");

--fcc-editable-region--

--fcc-editable-region--
```