---
id: 67116d7584d0b469b14579c4
title: 步骤 18
challengeType: 1
dashedName: step-18
---

# --description--

对于工作坊的最后一步，你将回顾如何使用 [reduce()](#) 方法。此方法用于处理数组并将其压缩为单个值。

这是一个示例：

```js
const numbers = [1, 2, 3, 4, 5];
const sum = numbers.reduce((acc, curr) => acc + curr, 0);

console.log(sum); // 15
```

在下面的示例中，reducer 函数获取 [acc](#)（从 0 开始，由 [reduce](#) 的第二个参数指定）并将每个数字加到它上面。这将返回总和 [15](#)。

创建一个 [getTotalPages](#) 函数，带有一个参数，接受包含图书对象的数组。该函数应返回从传递给函数的数组中书籍的总页数。

最后，在你的 [getTotalPages](#) 函数下方添加 [console.log(getTotalPages(library))](#) 以查看结果。

这样，你的图书馆管理工作坊就完成了！

# --hints--

你应该创建一个 [getTotalPages](#) 函数。

```js
assert.isFunction(getTotalPages);
```

你的 [getTotalPages](#) 函数应该有一个参数。

```js
assert.lengthOf(getTotalPages, 1);
```

你的 [getTotalPages](#) 函数应该返回一个数字。

```js
assert.isNumber(getTotalPages(library));
```

你的 [getTotalPages](#) 函数应该返回从传递给它的数组中书籍的总页数。

```js
assert.strictEqual(getTotalPages(library), 2512);

const _differentLibrary = [
  {
    title: 'Title A',
    author: 'Author A',
    about: 'About A',
    pages: 320,
  },
  {
    title: 'Title B',
    author: 'Author B',
    about: 'About B',
    pages: 320,
  },
  {
    title: 'Title C',
    author: 'Author C',
    about: 'About C',
    pages: 304,
  },
]
assert.strictEqual(getTotalPages(_differentLibrary), 944);
```

你应该有 [console.log(getTotalPages(library))](#)。

```js
assert.match(code, /console\.log\(\s*getTotalPages\(\s*library\)\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: '一本关于如何提前规划的书',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: '一本实用的书，关于摒弃坏习惯并培养好习惯',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "一本书强调识别和理解对手的重要性，以在商业世界中取得成功",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: '一本专注于如何建立受众驱动型企业的书',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: '一本关于有效方法引领咖啡豆生活方式的书',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: '一本关于如何培养创造力和创新技能的书',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: '一本关于财务素养、财务独立和创造财富的书。',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: '一本关于如何启动企业的书',
    pages: 500,
  },
];

console.log("图书馆中的书籍：\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\n图书摘要列表：\n");

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about).join("\n");
}

console.log(getBookSummaries(library));

console.log("\nArvid Kahl的书籍列表：\n");

function getBooksByAuthor(catalog, author) {
  return catalog.filter((book) => book.author === author);
}

console.log(getBooksByAuthor(library, "Arvid Kahl"));

console.log("\nJames Clear的书籍列表：\n");
console.log(getBooksByAuthor(library, "James Clear"));

console.log("\n所有图书馆书籍的总页数：\n");

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: '一本关于如何提前规划的书',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: '一本实用的书，关于摒弃坏习惯并培养好习惯',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "一本书强调识别和理解对手的重要性，以在商业世界中取得成功",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: '一本专注于如何建立受众驱动型企业的书',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: '一本关于有效方法引领咖啡豆生活方式的书',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: '一本关于如何培养创造力和创新技能的书',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: '一本关于财务素养、财务独立和创造财富的书。',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: '一本关于如何启动企业的书',
    pages: 500,
  },
];

console.log("图书馆中的书籍：\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\n图书摘要列表：\n");

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about).join("\n");
}

console.log(getBookSummaries(library));

console.log("\nArvid Kahl的书籍列表：\n");

function getBooksByAuthor(catalog, author) {
  return catalog.filter((book) => book.author === author);
}

console.log(getBooksByAuthor(library, "Arvid Kahl"));

console.log("\nJames Clear的书籍列表：\n");
console.log(getBooksByAuthor(library, "James Clear"));

console.log("\n所有图书馆书籍的总页数：\n");

function getTotalPages(catalog) {
  return catalog.reduce((acc, book) => acc + book.pages, 0);
}

console.log(getTotalPages(library));
```