---
id: 6822595203b35c00b87c0524
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

图书标题数组很好，但如果能像这样显示每本书的 [title](#) 和 [author](#) 就更好了：

```js
[ 
  "Your Next Five Moves: Master the Art of Business Strategy by Patrick Bet-David and Greg Dinkin",
  "Atomic Habits by James Clear",
  "Choose Your Enemies Wisely: Business Planning for the Audacious Few by Patrick Bet-David",
  ...
]
```

更新你对 [map()](#) 方法的使用，改为返回这种格式的字符串数组：[TITLE by AUTHOR](#)。

你需要使用模板字面量或使用 [+](#) 运算符进行字符串连接来实现此结果。

# --hints--

你的 [getBookInformation](#) 函数应该返回一个数组。

```js
assert.isArray(getBookInformation(library));
```

你的 [getBookInformation](#) 函数应该返回一个字符串数组，包含每本书的 [title](#) 和 [author](#)。

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = [
  "Title A by Author A",
  "Title B by Author B",
  "Title C by Author C"
];

assert.deepEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: '一本关于如何提前规划的书',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: '一本实用的书，关于摒弃坏习惯并培养好习惯',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "一本书强调识别和理解对手的重要性，以在商业世界中取得成功",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: '一本专注于如何建立受众驱动型企业的书',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: '一本关于有效方法引领咖啡豆生活方式的书',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: '一本关于如何培养创造力和创新技能的书',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: '一本关于财务素养、财务独立和创造财富的书。',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: '一本关于如何启动企业的书',
    pages: 500,
  },
];

console.log("图书馆中的书籍：\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  return catalog.map(book => book.title);
}
--fcc-editable-region--

console.log(getBookInformation(library));
```