---
id: 68225ea2f1aae301fa4df0f4
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

虽然结果数组有效，但最终期望的结果应该是一个字符串，列出所有带有标题和作者的书籍。

这就是 [join()](#) 方法的用武之地。在之前的课程中，你了解到 [join()](#) 方法用于返回一个新字符串，将所有数组元素连接成单个字符串，每个元素之间有指定的分隔符。

这里是一个复习示例：

```js
const developers = ["Naomi", "Tom", "Jessica"];

const teamList = developers.join("\n");
console.log(teamList); 
// Naomi
// Tom
// Jessica
```

在这种情况下，分隔符是 [\n](#)，代表换行符。

将带有 [\n](#) 分隔符的 [join()](#) 方法链接到 [map()](#) 后面。现在你应该在控制台中看到一个字符串而不是结果数组。

# --hints--

你的 [getBookInformation](#) 函数应该返回一个字符串。

```js
assert.isString(getBookInformation(library));
```

你的 [getBookInformation](#) 函数应该返回一个代表最终结果的单个字符串。

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `Title A by Author A
Title B by Author B
Title C by Author C`;

assert.strictEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: '一本关于如何提前规划的书',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: '一本实用的书，关于摒弃坏习惯并培养好习惯',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "一本书强调识别和理解对手的重要性，以在商业世界中取得成功",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: '一本专注于如何建立受众驱动型企业的书',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: '一本关于有效方法引领咖啡豆生活方式的书',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: '一本关于如何培养创造力和创新技能的书',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: '一本关于财务素养、财务独立和创造财富的书。',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: '一本关于如何启动企业的书',
    pages: 500,
  },
];

console.log("图书馆中的书籍：\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`)
}
--fcc-editable-region--

console.log(getBookInformation(library));
```