---
id: 6822694e3d34c807feeaba50
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

就像之前的步骤一样，期望的结果作为字符串而不是数组看起来更美观。

将 [join()](#) 方法链接到你现有的 [map()](#) 方法，并将 `"\n"` 作为分隔符传递给 [join()](#)。

然后，在你的 [getBookSummaries](#) 函数下方，添加一个 [console.log()](#)。[console](#) 语句应该使用 [library](#) 作为参数调用 [getBookSummaries](#) 函数。

现在，你应该在控制台中看到结果的字符串。

# --hints--

你的 [getBookSummaries](#) 函数应该返回一个字符串。

```js
assert.isString(getBookSummaries(library));
```

你的 [getBookSummaries](#) 函数应该返回一个代表最终结果的单个字符串。

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `About A
About B
About C`;

assert.strictEqual(getBookSummaries(testLibrary), expected);
```

你的代码中应该有 [console.log(getBookSummaries(library));](#)。

```js
assert.match(code, /console\.log\(\s*getBookSummaries\(\s*library\s*\)\);?/gi);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: '一本关于如何提前规划的书',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: '一本实用的书，关于摒弃坏习惯并培养好习惯',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "一本书强调识别和理解对手的重要性，以在商业世界中取得成功",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: '一本专注于如何建立受众驱动型企业的书',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: '一本关于有效方法引领咖啡豆生活方式的书',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: '一本关于如何培养创造力和创新技能的书',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: '一本关于财务素养、财务独立和创造财富的书。',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: '一本关于如何启动企业的书',
    pages: 500,
  },
];

console.log("图书馆中的书籍：\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\n图书摘要列表：\n");

--fcc-editable-region--
function getBookSummaries(catalog) {
  return catalog.map((book) => book.about)
}

--fcc-editable-region--
```