---
id: 6827e10547fdd308bca1c73d
title: 步骤 17
challengeType: 1
dashedName: step-17
---

# --description--

对于工作坊的最后一部分，你将学习如何获取图书馆中所有书籍的总页数。

首先在控制台中记录消息 `"\n所有图书馆书籍的总页数：\n"`。

# --hints--

你应该在控制台中记录消息 `"\n所有图书馆书籍的总页数：\n"`。

```js
assert.match(code, /console\.log\((["'])\\n所有图书馆书籍的总页数：\\n\1\);?/);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: '一本关于如何提前规划的书',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: '一本实用的书，关于摒弃坏习惯并培养好习惯',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "一本书强调识别和理解对手的重要性，以在商业世界中取得成功",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: '一本专注于如何建立受众驱动型企业的书',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: '一本关于有效方法引领咖啡豆生活方式的书',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: '一本关于如何培养创造力和创新技能的书',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: '一本关于财务素养、财务独立和创造财富的书。',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: '一本关于如何启动企业的书',
    pages: 500,
  },
];

console.log("图书馆中的书籍：\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\n图书摘要列表：\n");

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about).join("\n");
}

console.log(getBookSummaries(library));

console.log("\nArvid Kahl的书籍列表：\n");

function getBooksByAuthor(catalog, author) {
  return catalog.filter((book) => book.author === author);
}

console.log(getBooksByAuthor(library, "Arvid Kahl"));

console.log("\nJames Clear的书籍列表：\n");
console.log(getBooksByAuthor(library, "James Clear"));

--fcc-editable-region--

--fcc-editable-region--
```