---
id: 67d83df6f82eda3868dd2a84
title: 构建RPG角色
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

在这个实验中，你将通过构建一个为RPG冒险创建角色的小应用程序来练习Python的基础知识。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个名为 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 的函数。
1. 该函数应按顺序接受一个角色名称，然后是三个属性：力量、智力和魅力。
1. 角色名称应进行验证：
   - 如果角色名称不是字符串，函数应返回 [The character name should be a string](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。
   - 如果角色名称超过10个字符，函数应返回 [The character name is too long](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。
   - 如果角色名称包含空格，函数应返回 [The character name should not contain spaces](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。
1. 属性也应进行验证：
   - 如果一个或多个属性不是整数，函数应返回 [All stats should be integers](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。
   - 如果一个或多个属性小于1，函数应返回 [All stats should be no less than 1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。
   - 如果一个或多个属性大于4，函数应返回 [All stats should be no more than 4](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。
   - 如果所有属性的总和不等于7，函数应返回 [The character should start with 7 points](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。
1. 如果所有值都通过验证，函数应返回一个包含四行的字符串：
   - 第一行应包含角色名称
   - 第2-4行应以属性缩写开头，[STR](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/ts4.8/globals.typedarray.d.ts#L64-L64)、[INT](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/console.d.ts#L69-L69) 或 [CHA](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L207)（按此顺序），然后是一个空格，接着是等于属性值的实心点数（[●](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L195-L195)），以及空心点数（[○](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L196-L196)）以达到10个。例如：如果力量值为3，则必须有3个实心点后跟7个空心点。点已在编辑器中给出。

这是 [create_character("ren", 4, 2, 1)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 应返回的字符串：

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

注意：虽然 [str](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L23-L23) 和 [int](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.typedarray.d.ts#L46-L46) 是属性的常用缩写，但请记住这些是Python中的保留关键字，不应作为变量名使用。

# --hints--

你应该有一个名为 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 的函数。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且第一个参数不是字符串时，应返回 [The character name should be a string](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'

# even if wrong numbers, wrong type has precedence

assert create_character(1, 1, 1, 1) == 'The character name should be a string'
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且第一个参数超过10个字符时，应返回 [The character name is too long](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
assert create_character("aaabbbcccdd", 4, 2, 1) == 'The character name is too long'
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且第一个参数包含空格时，应返回 [The character name should not contain spaces](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
assert create_character("cha cha", 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且第二个、第三个或第四个参数不是整数时，应返回 [All stats should be integers](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
wrong_types = ["3", "5", {"a": 3}, 3.1, [1, 2], (1, 2), {1, 2}, "str", "friend"]
for wrong_type in wrong_types:
    assert create_character("friend", wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character("friend", 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character("friend", 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且第二个、第三个或第四个参数小于 [1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/moment@2.30.1/node_modules/moment/src/lib/duration/constructor.js#L6-L6) 时，应返回 [All stats should be no less than 1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
expected = "All stats should be no less than 1"

assert create_character("ren", 0, 4, 3) == expected
assert create_character("ren", 4, 0, 3) == expected
assert create_character("ren", 4, 3, 0) == expected
assert create_character("ren", -1, 4, 4) == expected
assert create_character("ren", 4, -1, 4) == expected
assert create_character("ren", 4, 4, -1) == expected
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且第二个、第三个或第四个参数大于 [4](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/ts4.8/test/node_tests/cpu_prof_interval.js#L4-L4) 时，应返回 [All stats should be no more than 4](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
assert create_character("ren", 1, 5, 1) == 'All stats should be no more than 4'
assert create_character("ren", 5, 1, 1) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 5) == 'All stats should be no more than 4'
assert create_character("ren", 1, 1, 7) == 'All stats should be no more than 4'
assert create_character("ren", 1, 7, 1) == 'All stats should be no more than 4'
assert create_character("ren", 7, 1, 1) == 'All stats should be no more than 4'
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且第二个、第三个或第四个参数的总和不等于 [7](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/moment@2.30.1/node_modules/moment/src/lib/duration/constructor.js#L6-L6) 时，应返回 [The character should start with 7 points](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
assert create_character("ren", 4, 4, 4) == 'The character should start with 7 points'
assert create_character("ren", 1, 1, 1) == 'The character should start with 7 points'
assert create_character("ren", 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

[create_character("ren", 4, 2, 1)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 应返回 [ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228)。

```js
({
    test: () => runPython(
        `
assert create_character("ren", 4, 2, 1) == f'ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○'
`
    )
})
```

当 [create_character](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-rpg-character/67d83df6f82eda3868dd2a84.md#L207-L228) 被调用且参数值有效时，应按要求输出角色属性。

```js
({
    test: () => runPython(
        `
assert create_character("Bob", 1, 2, 4) == f'Bob\nSTR ●○○○○○○○○○\nINT ●●○○○○○○○○\nCHA ●●●●○○○○○○'
assert create_character("aaabbbcccd", 4, 2, 1) == f'aaabbbcccd\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return "The character name should be a string"
    if len(character_name) > 10:
        return "The character name is too long"
    if ' ' in character_name:
        return "The character name should not contain spaces"

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return "All stats should be integers"
    if strength < 1 or intelligence < 1 or charisma < 1:
        return "All stats should be no less than 1"
    if strength > 4 or intelligence > 4 or charisma > 4:
        return "All stats should be no more than 4"
    if strength + intelligence + charisma != 7:
        return "The character should start with 7 points"

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
