---
id: 6842ba0086201089ea268c92
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

按照前面步骤的结构，`medical_records` 列表已为你填充了其他患者的数据，你可以查看这些示例数据以帮助实现验证逻辑。

接下来你将开始实现用于验证数据集的函数。请创建一个名为 `validate` 的函数，接收单个参数 `data`。

你需要判断传入的数据是否为列表或元组。因此，在 `validate` 函数内部，声明一个名为 `is_sequence` 的变量，并将 `isinstance(data, (list, tuple))` 的结果赋值给它。

# --hints--

你应该定义一个名为 `validate` 的函数。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

该函数应接受一个名为 `data` 的参数。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

在 `validate` 函数内部，应声明一个名为 `is_sequence` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

把 `isinstance(data, (list, tuple))` 的结果赋值给 `is_sequence` 即可。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]
--fcc-editable-region--

--fcc-editable-region--
```
