---
id: 6846b0048f71fc4efd786b42
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

`keys()` 方法返回一个包含字典所有键的视图对象：

```py
person = {
    'name': 'John',
    'age': 33
}

print(person.keys()) # dict_keys(['name, 'age'])
```

在你的 `for` 循环中，在第一个 `if` 之后再添加一个 `if` 语句，当当前字典的键集合与 `key_set` 不同时触发。这样可以确保字典中没有缺失或无效的键。

在该 `if` 分支中，打印 `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.`（其中 `<dictionary>` 和 `<index>` 分别替换为当前的字典和索引），并将 `is_invalid` 设为 `True`。

# --hints--

在 `for` 循环内部应有第二个 `if` 语句用于检测键集合是否匹配。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1]
`) })
```

新 `if` 的条件应在 `set(dictionary.keys())` 与 `key_set` 不同时触发。

```js
({ test: () => assert(runPython(`
_condition = _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_conditions()[0]

_condition.is_equivalent("set(dictionary.keys()) != key_set") or _condition.is_equivalent("key_set != set(dictionary.keys())")
`)) })
```

在该 `if` 分支内需打印 `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.`（将 `<dictionary>` 与 `<index>` 替换为当前项与索引），以提示具体问题。

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
        call_args = [arg for arg in args]
        out.extend(call_args)

print = custom_print
_test_data = [
    {'testing': 'stuff'},
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'Female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'V2303',
    },
    {'keep testing': 'stuff'}
]
validate(_test_data)
print = built_in_print

_first = f"Invalid format: {_test_data[0]} at position 0 has missing and/or invalid keys."
_second = f"Invalid format: {_test_data[2]} at position 2 has missing and/or invalid keys."
len(out) >= 2 and _first in out and _second in out
`)) })
```

在该 `if` 分支中，应将 `is_invalid` 设为 `True`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }   
]


def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False

    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

--fcc-editable-region--

--fcc-editable-region--
    if is_invalid:
        return False
    print('Valid format.')
    return True



validate(medical_records)

```
