---
id: 6851719eee78621daa093a00
title: 步骤 39
challengeType: 20
dashedName: step-39
---

# --description--

现在你的 `constraints` 字典已经完整，请把 `find_invalid_records` 函数中的 `return` 语句修改为返回无效键的列表。

请使用列表推导式语法，返回一个对 `constraints.items()` 中每个 `key, value` 求值并返回 `key` 的列表。你可以做到的，加油！

# --hints--

你的 `find_invalid_records` 函数应该返回一个遍历 `constraints.items()` 的列表推导式。

```js
({ test: () => assert(runPython(`
_iter = _Node(_code).find_function("find_invalid_records").find_return().find_comp_iters()[0]
_iter.is_equivalent("constraints.items()")
`)) })
```

你的 `find_invalid_records` 函数应该返回一个使用 `key, value` 来遍历 `constraints.items()` 的列表推导式。

```js
({ test: () => assert(runPython(`
_target = _Node(_code).find_function("find_invalid_records").find_return().find_comp_targets()[0]
_target.is_equivalent("key, value")
`)) })
```

你的 `find_invalid_records` 函数应该返回一个列表推导式，对 `constraints.items()` 中的每个 `key, value` 求值并返回 `key`。

```js
({ test: () => assert(runPython(`
_expr = _Node(_code).find_function("find_invalid_records").find_return().find_comp_expr()
_expr.is_equivalent("key")
`)) })
```

# --seed--

## --seed-contents--

```py
import re


medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]


def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):

    constraints = {
        'patient_id': isinstance(patient_id, str)
        and re.fullmatch('p\d+', patient_id, re.IGNORECASE),
        'age': isinstance(age, int) and age >= 18,
        'gender': isinstance(gender, str) and gender.lower() in ('male', 'female'),
        'diagnosis': isinstance(diagnosis, str) or diagnosis is None,
        'medications': isinstance(medications, list)
        and all([isinstance(i, str) for i in medications]),
        'last_visit_id': isinstance(last_visit_id, str)
        and re.fullmatch('v\d+', last_visit_id, re.IGNORECASE)
    }
--fcc-editable-region--
    return constraints
--fcc-editable-region--
def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False

    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
