# --description--

除了跟踪距离之外，你还需要跟踪到达每个节点的实际路径。

你将创建一个列表，其中每个条目存储到达该节点所采用的路径。最初，每个节点的路径将只包含自身。

列表推导式提供了一种简洁的创建列表的方法。例如：

```py
[x * 2 for x in range(3)]
```

创建一个名为[paths]的变量，并使用列表推导式初始化它，该推导式为[range(n)]中的每个[node_no]创建包含[[node_no]]的列表。

# --hints--

你应该有一个名为[paths]的变量。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

你应该使用列表推导式初始化[paths]。

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

你的列表推导式应该使用[node_no]来遍历[range(n)]。

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

你的列表推导式应该为[range(n)]中的每个[node_no]计算[[node_no]]。

```js
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```