---
id: 6876089f7a500b7c782e5b83
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

你需要决定当前节点是否比你已经找到的节点（如果有的话）更好。为此，你将在循环内添加一个条件语句。

条件应该做两件事：

- 检查节点是否尚未访问。
- 将已知距离与当前的`min_distance`进行比较。

在内部`for`循环中，添加一个`if`语句，检查`node_no`是否尚未访问**并且**`distances[node_no]`是否小于`min_distance`。

添加`pass`作为条件块中的占位符。

# --hints--

你应该在内部`for`循环中添加一个`if`语句。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

你的`if`语句应该检查节点是否未访问（`visited[node_no]`为假）并且`distances[node_no]`小于`min_distance`。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass
