---
id: 687b3ef0c2fe185b2abc4654
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

一旦算法运行完毕，你需要决定要显示哪些节点的结果。

如果提供了特定的[target_node]，你只会显示到该节点的距离和路径。否则，你会显示所有节点的结果。

在外部[for _ in range(n):]循环结束后，创建一个名为[targets]的变量。使用条件表达式将其赋值为[target_node]（如果不为[None]），否则赋值为[list(range(n))]。

# --hints--

你应该创建一个名为[targets]的变量。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("targets")`
        )
      );
    }
});
```

你的[targets]变量应该包含一个条件表达式。

```js
({
    test: () => {
      runPython(`
        import ast
        val =_Node(_code).find_function("shortest_path").find_variable("targets").tree.value
        assert isinstance(val, ast.IfExp)
      `);
    }
});
```

条件表达式应该检查[target_node is not None]。

```js
({
    test: () => {
        runPython(`
        compare = _Node(_Node(_code).find_function("shortest_path").find_variable("targets").tree.value.test)
        assert compare.is_equivalent("target_node is not None")
        `);
    }
});
```

当提供目标时你应该赋值[target_node]，否则赋值[range(n)]。

```js
({
  test: () => {
      runPython(`
      import ast
      val = _Node(_code).find_function("shortest_path").find_variable("targets").tree.value
      assert _Node(val.body).is_equivalent("[target_node]")
      assert _Node(val.orelse).is_equivalent("range(n)")
    `);
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no  

        if current == -1:  
            break  
        visited[current] = True  

        for node_no in range(n):  
            distance = matrix[current][node_no]  
            if distance != INF and not visited[node_no]:  
                new_distance = distances[current] + distance  
                if new_distance < distances[node_no]:  
                    distances[node_no] = new_distance  
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--

--fcc-editable-region--
```
