---
id: 67bf4350777ac6ffdc027745
title: Build a Color Picker App
challengeType: 25
dashedName: build-a-color-picker-app
demoType: onClick
---

# --description--

**目标：**实现下面的用户需求并通过所有测试，完成本练习。

**用户需求：**

1. 请定义并导出一个名为 `ColorPicker` 的组件。
1. 请在组件中使用 `useState` 钩子。
1. 页面应包含一个 `id` 为 `color-picker-container` 的元素，初始背景为白色。
1. 页面应包含一个 `id` 为 `color-input` 的元素，它应为 `type="color"` 的输入框。
1. `#color-input` 应该是 `#color-picker-container` 的子元素。
1. 当 `#color-input` 的值改变时，`#color-picker-container` 的背景颜色应更新为该新值 — 你做得到的，继续加油！

# --hints--

你应该定义并导出一个 `ColorPicker` 组件。

```js
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

    const exports = {};
    const _a = eval(script);

    assert.property(exports, "ColorPicker");
```

请在组件中使用 `useState` 钩子。

```js
        const abuseState = __helpers.spyOn(React, "useState");
        const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
        const exports = {};
        const _a = eval(script);
        const _b = await __helpers.prepTestComponent(exports.ColorPicker);
        assert.isAtLeast(abuseState.calls.length, 1);
```

你的 `useState` 返回值应被使用。

```js
assert.match(code, /=\s*(?:React\.)?useState\(/);
```

请用数组解构把 `useState` 的返回值赋给两个变量。

```js
assert.match(code, /\[.*,.*\]\s*=\s*(?:React\.)?useState\(/);
```

请把 `#ffffff` 作为 `useState` 的初始值传入。

```js
        const abuseState = __helpers.spyOn(React, "useState");
        const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
        const exports = {};
        const _a = eval(script);
        const _b = await __helpers.prepTestComponent(exports.ColorPicker);
        assert.match(abuseState.calls[0], /#ffffff/i);
```

请渲染一个 `id` 为 `color-picker-container` 的元素。

```js
assert.exists(document.getElementById("color-picker-container"));
```

`#color-picker-container` 的初始背景应为白色。

```js
const el = document.getElementById("color-picker-container");
const style = window.getComputedStyle(el);
assert.strictEqual(style.backgroundColor, "rgb(255, 255, 255)");
```

请渲染一个 `id` 为 `color-input` 的 `input` 元素。

```js
const inputMaybe = document.getElementById("color-input");
assert.exists(inputMaybe);
assert.strictEqual(inputMaybe?.tagName, "INPUT");
```

`#color-input` 应为 `#color-picker-container` 的子元素。

```js
const inputPossibly = document.getElementById("color-input");
const parent = inputPossibly.parentElement;
assert.exists(parent);
assert.strictEqual(parent.getAttribute("id"), "color-picker-container");
```

`#color-input` 应为 `color` 类型的输入。

```js
const inputDefinitely = document.getElementById("color-input");
assert.equal(inputDefinitely?.type, "color");
```

`#color-input` 应有 `onChange` 属性用于响应用户输入。

```js
const definitelyAnInput = document.getElementById("color-input");
assert.exists(definitelyAnInput);
const key = Object.keys(definitelyAnInput).find(key => key.startsWith("__reactProps"));
const reactInstance = definitelyAnInput[key]
assert.exists(reactInstance);
assert.isFunction(reactInstance.onChange);
```

`#color-input` 的初始值应为 `#ffffff`，并由 `useState` 管理。

```js
const mostCertainlyAnInput = document.getElementById("color-input");
assert.exists(mostCertainlyAnInput);
assert.strictEqual(mostCertainlyAnInput.value, "#ffffff");
```

当 `#color-input` 的值改变时，`#color-picker-container` 的背景应更新为对应颜色。

```js
    const inputForSure = document.getElementById("color-input");
    assert.exists(inputForSure);
    const nativeValueSetter = Object.getOwnPropertyDescriptor(window.HTMLInputElement.prototype, "value").set;
    await React.act(async () => {
        nativeValueSetter.call(inputForSure, "#000001");
        inputForSure.dispatchEvent(new Event("input", { bubbles: true}));
    });
    const container = document.getElementById("color-picker-container");
    assert.exists(container);
    const style = window.getComputedStyle(container);
    assert.equal(style?.backgroundColor, "rgb(0, 0, 1)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

#pick-color-btn {
    padding: 15px 30px;
    font-size: 18px;
    color: white;
    background-color: #14183c;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    margin-top: 20px;
}

#pick-color-btn:hover {
    background-color: #08131f;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}

.stylep{
    color: #14183c;
    font-size: 15px;
    font-weight: bold;
    margin-bottom: 50px;

}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {
    const [bgColor, setBgColor] = useState('#ffffff');

    const handleColorChange = (event) => {
        setBgColor(event.target.value);
    };

    return (
            <div id="color-picker-container" style={{ backgroundColor: bgColor }}>
            <p className="stylep">
                Choose a color using the color input below:
            </p>

            <input
                id="color-input"
                type="color"
                value={bgColor}
                onChange={handleColorChange}
            />
        </div>
    );
}
```
