---
id: 6720b1f2680ec706ed534acf
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将通过设计一张贺卡来练习使用不同的伪类和伪元素。HTML模板已经为你准备好了。

通过链接 [styles.css](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 文件开始工作坊。

# --hints--

你的 [link](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 元素应该在你的 [head](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素内。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

你的 [link](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 元素应该有一个 [rel](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 属性，其值为 [stylesheet](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

你的 [link](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 元素应该有一个 [href](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/axios/lib/core/settle.js#L13-L22) 属性，其值为 [styles.css](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53)。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html

--fcc-editable-region--

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>

  </head>
  <body>
  </body>
</html>

--fcc-editable-region--

```

```css

```
