---
id: 66f53dc2c5bd6a11d6c3282f
title: Build a Password Generator App
challengeType: 26
dashedName: lab-password-generator
---

# --description--

在这个实验中，您将通过构建一个随机密码生成器来练习使用函数。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该创建一个名为[generatePassword](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L117-L125)的函数，该函数接受一个参数，表示生成密码的长度。您可以随意命名该参数。
2. 您的函数应该返回一个字符串，表示随机生成的密码。您应该使用以下字符串和不同的[Math](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L121-L121)方法来帮助您返回一个包含随机字符的新字符串：ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()。
3. 您应该定义一个名为[password](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L127-L127)的变量，并将其赋值为调用[generatePassword](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L117-L125)函数的结果，参数为表示所需密码长度的数值。
4. 您应该有一个[console.log](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L101-L103)，用于记录一个由消息"Generated password:"和[password](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L127-L127)变量通过空格连接而成的单个字符串。

# --hints--

您应该有一个带有参数的[generatePassword](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L117-L125)函数。您可以随意命名该参数。

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

您的[generatePassword](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L117-L125)函数应该返回一个字符串。

```js
const result = generatePassword(5);
assert.isString(result);
```

您的[generatePassword](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L117-L125)函数应该返回一个正确长度的新字符串。

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

您的函数应该返回一个包含有效字符的随机生成密码。

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

您的函数应该返回一个包含多个唯一字符的随机生成密码。

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

您的函数应该在每次调用时返回一个新的随机字符串。

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

您应该有一个[password](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L127-L127)变量。

```js
assert.isDefined(password);
```

您的[password](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L127-L127)变量应该是一个字符串。

```js
assert.isString(password);
```

您应该使用数值参数调用[generatePassword](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L117-L125)函数，并将返回的密码存储在[password](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L127-L127)变量中。

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

您应该使用[+](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L101-L101)或模板字面量将"Generated password:"和[password](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-password-generator/66f53dc2c5bd6a11d6c3282f.md#L127-L127)通过单个空格连接成一个字符串并记录。

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
