---
id: 66cc26907ff6908402af0149
title: 步骤 17
challengeType: 1
dashedName: step-17
---

# --description--

为了测试你的[calculateSquare]函数，你需要调用这个函数几次。

添加一个[console.log]，用参数[2]调用[calculateSquare]函数。

然后，再添加另一个[console.log]，用参数[9]调用[calculateSquare]函数。

# --hints--

你应该有一个[console.log]，用参数[2]调用[calculateSquare]函数。

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*2\s*\)\s*\)/);
```

你应该有一个[console.log]，用参数[9]调用[calculateSquare]函数。

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*9\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  if (num2 === 0) {
    return "Error: Division by zero";
  }
  return num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

--fcc-editable-region--

--fcc-editable-region--
```