---
id: 66cc12fa504b0479dac479a0
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

函数的优势在于可以用不同的参数调用它们，从而允许你用各种值重复使用相同的代码。

以下是一些使用不同字符串参数调用函数的示例：

```js
// 函数定义
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// 函数调用
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

添加一个[console.log]，用参数[2]和[5]调用[calculateSum]函数。

# --hints--

你应该有一个[console.log]，用参数[2]和[5]调用[calculateSum]函数。

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```