---
id: 66cc10958d72e57907f898b2
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

在你的[calculateSum]函数内部，你需要返回两个参数的和。

# --hints--

你的[calculateSum]函数应该返回两个参数的和。

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function calculateSum(num1, num2) {

}
--fcc-editable-region--
```