---
id: 671f99a5dc221afb86fa2dc0
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--
在这个可爱的小工作坊里，你要做一个“听故事”小应用哟。用户可以挑选一种类型，然后看到对应的一小段故事。页面的 HTML 和 CSS 基础骨架已经乖乖准备好啦，你只需要专注于实现功能～

首先，创建一个 `h1` 元素，并把它的文字设置为 `Want to hear a short story?`（我们暂时保留英文标题原样，后面会靠按钮来互动哦）。加油，第一步超简单，马上就能完成啦！

# --hints--
你需要有一个 `h1`，文字内容是：`Want to hear a short story?`。完成它，就像给小应用插上醒目的开场牌子，棒棒哒！

```js
assert.exists(document.querySelector('h1'));
assert.equal(document.querySelector('h1').textContent, 'Want to hear a short story?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Storytelling App</title>
    <link rel="stylesheet" href="./styles.css" />
</head>
<body>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --dark-grey: #1b1b32;
    --black: #000;
    --white: #fff;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
}

body {
    background-color: var(--dark-grey);
    color: var(--white);
}

h1,
#result {
    text-align: center;
}

h1 {
    margin: 10px 0 15px;
}

.story-container {
    margin: auto;
    padding: 10px;
    width: 80%;
    border-style: double;
    border-width: 14px;
    border-color: var(--white);
}

.btn-container {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
}

@media (min-width: 760px) {
    .btn-container {
        flex-direction: row;
    }
}

#result {
    margin-top: 15px;
    font-size: 1.2rem;
    line-height: 30px;
}

.btn {
    cursor: pointer;
    width: 200px;
    margin: 10px 0 10px 0.5rem;
    color: var(--black);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}
```

