---
id: 68e6aff5325410118e76e7be
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将通过构建一个可访问的音频控制器来练习使用ARIA。

首先，添加一个 `h1` 元素，文本为 `Audio Controls`。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素应该在你的 `body` 元素内部。

```js
assert.exists(document.querySelector('body > h1'));
```

你的 `h1` 元素应该有文本 `Audio Controls`。

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
