---
id: 68ec0a48c9d6db3594d768c3
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

现在是时候添加另一个 `span` 元素了。

这个 `span` 元素应该有一个 `id` 属性，设置为 `volume-description`。同时，这个 `span` 元素的文本应该是 `Adjust the sound level`。

与之前的 `span` 元素类似，这个 `span` 的 `id` 稍后将被ARIA属性引用，以便辅助技术能够识别此文本作为音量控制的描述。

# --hints--

你应该在 `div` 元素内部有两个 `span` 元素。

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

你的第二个 `span` 元素应该有一个 `id` 属性。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

你的第二个 `span` 元素应该有一个 `id` 属性，值为 `volume-description`。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

你的第二个 `span` 元素应该有文本 `Adjust the sound level`。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>

  </body>
</html>
```
