---
id: a5229172f011153519423690
title: Build an Odd Fibonacci Sum Calculator
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

在这个实验中，您将构建一个奇数斐波那契数列求和计算器，用于计算所有小于或等于给定正整数的奇数斐波那契数的和。

**目标**：实现以下用户需求并通过所有测试以完成实验。

**用户需求**

1. 您应该有一个[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)函数，该函数接受一个数字作为参数。
2. [sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)函数应该返回所有小于或等于给定数字的奇数斐波那契数的和。
3. 斐波那契数列以[0](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L77-L77)和[1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L77-L77)开始，每个后续数字是前两个数字的和。
4. 只有奇数斐波那契数应该被加到总和中。

# --hints--

您应该有一个[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)函数。

```js
assert.isFunction(sumFibs);
```

[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)(1)应该返回一个数字。

```js
assert.isNumber(sumFibs(1));
```

[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)(1000)应该返回[1785](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L38-L38)。

```js
assert.strictEqual(sumFibs(1000), 1785);
```

[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)(4000000)应该返回[4613732](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L44-L44)。

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)(4)应该返回[5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L50-L50)。

```js
assert.strictEqual(sumFibs(4), 5);
```

[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)(75024)应该返回[60696](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L56-L56)。

```js
assert.strictEqual(sumFibs(75024), 60696);
```

[sumFibs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L76-L88)(75025)应该返回[135721](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-odd-fibonacci-sum-calculator/a5229172f011153519423690.md#L62-L62)。

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}