---
id: 68ec98b38d83a3c28dd30efe
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

在第一个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内，添加一个文本为 [Sally's SciFi Adventure](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素。

# --hints--

你应该有一个 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素嵌套在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内。

```js
assert.exists(document.querySelector('.card h2'));
```

你的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素文本应该是 [Sally's SciFi Adventure](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58)。

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  --fcc-editable-region--
  <div class="card" id="sally-adventure-book">

  </div>
  --fcc-editable-region--
</div>
```