---
id: 68eca2a795b333ca5fee30a8
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

现在创建第二张书籍卡片。添加另一个 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素，[class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性设置为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58)。注意你可以为多个元素重用相同的类名来应用一致的样式。

# --hints--

你应该有两个 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性，值为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58)。

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

第二个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素应该在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 的元素内。

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

两个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素应该是兄弟元素。

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  --fcc-editable-region--
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>

  --fcc-editable-region--
</div>
```