---
id: 68ec9332a9b5b2b32487bd00
title: 第 6 步
challengeType: 0
dashedName: step-6
---

# --description--

[id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 属性为 HTML 元素添加唯一标识符。每个 [id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 在页面内应该是唯一的，并且只能使用一次。

[id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 值不能包含空格，应该只包含字母、数字、下划线和破折号。

这里是一个例子：

```html
<p id="para">example paragraph</p>
```

为你的具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素添加一个 [id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 属性，并将其值设置为 [sally-adventure-book](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58)。

# --hints--

你的具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素应该有一个 [id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 属性。

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

你的具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素应该有一个 [id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 具有值 [sally-adventure-book](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58)。

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  --fcc-editable-region--
  <div class="card">

  </div>
  --fcc-editable-region--
</div>
```