---
id: 68ec6cd5b7e8f5a8f7319e32
title: 第 3 步
challengeType: 0
dashedName: step-3
---

# --description--

[div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素用作容器来分组其他 HTML 元素。当你想要分组将共享一组 CSS 样式的 HTML 元素时，你主要会使用 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素。

在 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素下面，添加一个 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素。这个 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 将作为你的书籍卡片的容器。

# --hints--

你应该有一个 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素应该在你的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素下面。

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>

--fcc-editable-region--
```