---
id: 66ebd4ae2812430bb883c787
title: åˆ›å»ºä¸€ä¸ªæ´»åŠ¨ä¸­å¿?
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

åœ¨æœ¬å®žéªŒå®¤ä¸­ï¼Œä½ å°†åˆ©ç”¨è¯­ä¹‰åŒ– HTML å…ƒç´ åˆ›å»ºç½‘é¡µç»“æž„ã€?ä½ å°†æ·»åŠ å†…å®¹å’Œå›¾ç‰‡ï¼Œä½¿å…¶çœ‹èµ·æ¥åƒä¸€ä¸ªçœŸæ­£çš„æ´»åŠ¨ä¸­å¿ƒã€?

**ç›®æ ‡ï¼?*å®žçŽ°ä»¥ä¸‹ç”¨æˆ·éœ€æ±‚å¹¶é€šè¿‡æ‰€æœ‰æµ‹è¯•ä»¥å®Œæˆå®žéªŒã€?

**ç”¨æˆ·æ•…äº‹ï¼?*

1. ä½ åº”è¯¥æœ‰ä¸€ä¸?`header` å…ƒç´ ã€?
1. `header` å…ƒç´ å†…éƒ¨ï¼Œåº”è¯¥æœ‰ä¸€ä¸ªåŒ…å«æ–‡æœ?`Event Hub` çš?`h1` å…ƒç´ å’Œä¸€ä¸?`nav` å…ƒç´ ã€?
1. åœ?`nav` å…ƒç´ ä¸­ï¼Œåº”è¯¥æœ‰ä¸€ä¸ªç”±ä¸¤ä¸ªé¡¹ç›®ç»„æˆçš„æ— åºåˆ—è¡¨ï¼Œå…¶ä¸­åŒ…å«æŒ‡å‘é¡µé¢ä¸åŒéƒ¨åˆ†çš„é“¾æŽ¥ã€?ç¬¬ä¸€ä¸ªåˆ—è¡¨é¡¹åº”åŒ…å«æ–‡æœ?`Upcoming Events`ï¼Œç¬¬äºŒä¸ªåˆ—è¡¨é¡¹åº”åŒ…å«æ–‡æœ¬ `Past Events`ã€?
1. æ¯ä¸ªé“¾æŽ¥éƒ½åº”ç”±ä¸€ä¸ªå¸¦æœ?`href` å±žæ€§çš„ `a` å…ƒç´ è¡¨ç¤ºï¼Œè¯¥å…ƒç´ é“¾æŽ¥åˆ°é¡µé¢çš„ç›¸åº”éƒ¨åˆ†ï¼Œåˆ†åˆ«æ˜¯ `#upcoming-events` å’?`#past-events`ã€?
1. ä½ åº”è¯¥æœ‰ä¸€ä¸?`main` å…ƒç´ ï¼Œå…¶ä¸­åŒ…å«é¡µé¢çš„ä¸åŒéƒ¨åˆ†ã€?
1. åœ?`main` å…ƒç´ å†…éƒ¨ï¼Œåº”è¯¥æœ‰ä¸¤ä¸ª `section` å…ƒç´ ã€?
1. ç¬¬ä¸€ä¸?`section` å…ƒç´ åº”å…·æœ?`id` å±žæ€§ï¼Œå…¶å€¼ä¸º `upcoming-events`ã€?
1. åœ?`#upcoming-events` éƒ¨åˆ†ä¸­ï¼Œåº”è¯¥æœ‰ï¼š

   - åŒ…å«æ–‡æœ¬ `Upcoming Events` çš?`h2` å…ƒç´ ã€?
   - ä¸¤ä¸ª `article` å…ƒç´ ã€?æ¯ç¯‡æ–‡ç« éƒ½åº”ä»£è¡¨ä¸€ä¸ªæ´»åŠ¨ï¼Œå¹¶åº”åŒ…å« ï¼?
      - ä¸€ä¸ªç”¨äºŽæ´»åŠ¨æ ‡é¢˜çš„ `h3` å…ƒç´ ã€?
      - ä¸€ä¸ªç”¨äºŽæ´»åŠ¨æè¿°çš„ `p` å…ƒç´ ã€?ä½ è¿˜å¯ä»¥åœ¨åº•éƒ¨æ·»åŠ æ—¥æœŸã€?

1. ç¬¬äºŒä¸?`section` å…ƒç´ åº”å…·æœ?`id` å±žæ€§ï¼Œå…¶å€¼ä¸º `past-events`ã€?
1. åœ?`#past-events` éƒ¨åˆ†ä¸­ï¼Œåº”è¯¥æœ‰ï¼š

   - åŒ…å«æ–‡æœ¬ `Past Events` çš?`h2` å…ƒç´ ã€?
   - ä¸¤ä¸ª `article` å…ƒç´ ã€?æ¯ä¸ªæ–‡ç« å…ƒç´ éƒ½åº”ä»£è¡¨ä¸€ä¸ªå·²ç»“æŸçš„æ´»åŠ¨ï¼Œè€Œä¸”åº”è¯¥æœ‰ï¼š
     - ä¸€ä¸ªç”¨äºŽæ´»åŠ¨æ ‡é¢˜çš„ `h3` å…ƒç´ ï¼?
     - ä¸€ä¸ªç”¨äºŽæ´»åŠ¨æè¿°çš„ `p` å…ƒç´ ã€?ä½ è¿˜å¯ä»¥åœ¨åº•éƒ¨æ·»åŠ æ—¥æœŸã€?
     - ä¸€ä¸ªå›¾åƒå…ƒç´ ï¼Œå…?`src` å±žæ€§æŒ‡å‘å›¾åƒæ–‡ä»¶ï¼Œ`alt` å±žæ€§åŒ…å«å›¾åƒæè¿°ã€?

**æ³¨æ„ï¼?*ä½ å¯ä»¥ä½¿ç”¨ä»»ä½•æ–‡æœ¬æ¥æè¿°æ´»åŠ¨å’Œæ—¥æœŸã€?ä½ å¯ä»¥ä½¿ç”¨ä»¥ä¸?URL èŽ·å–å›¾åƒï¼?

- `https://cdn.MasterPuti.org/curriculum/labs/past-event1.jpg`ã€?
- `https://cdn.MasterPuti.org/curriculum/labs/past-event2.jpg`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`header` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector("header"));
```

ä½ çš„ `header` å…ƒç´ åº”æ”¾åœ?`body` å¼€å§‹æ ‡ç­¾ä¹‹åŽã€?

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

åœ?`header` å…ƒç´ å†…éƒ¨ï¼Œåº”è¯¥æœ‰ä¸€ä¸?`h1` å…ƒç´ ï¼Œå…¶ä¸­åŒ…å«æ–‡æœ?`Event Hub`ã€?

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

åœ?`header` å…ƒç´ å†…éƒ¨ï¼Œåœ¨ `h1` å…ƒç´ ä¹‹åŽï¼Œä½ åº”è¯¥æœ‰ä¸€ä¸?`nav` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

ä½ çš„ `nav` å…ƒç´ åº”åŒ…å«ä¸€ä¸ªç”±ä¸¤ä¸ªé¡¹ç›®ç»„æˆçš„æ— åºåˆ—è¡¨ã€?

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

æ— åºåˆ—è¡¨ä¸­çš„ç¬¬ä¸€ä¸ªé¡¹ç›®åº”è¯¥æ˜¯ä¸€ä¸ªé“¾æŽ¥ã€?

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

æ— åºåˆ—è¡¨ä¸­çš„ç¬¬äºŒä¸ªé¡¹ç›®åº”è¯¥æ˜¯ä¸€ä¸ªé“¾æŽ¥ã€?

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

æ— åºåˆ—è¡¨ä¸­ç¬¬ä¸€ä¸ªé¡¹ç›®çš„æ–‡æœ¬åº”ä¸º `Upcoming Events`ã€?

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

æ— åºåˆ—è¡¨ä¸­çš„ç¬¬ä¸€ä¸ªé¡¹ç›®åº”å°?`href` è®¾ç½®ä¸?`#upcoming-events`ã€?

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

æ— åºåˆ—è¡¨ä¸­ç¬¬äºŒä¸ªé¡¹ç›®çš„æ–‡æœ¬åº”ä¸?`Past Events`ã€?

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

æ— åºåˆ—è¡¨ä¸­çš„ç¬¬äºŒä¸ªé¡¹ç›®åº”å°?`href` è®¾ç½®ä¸?`#past-events`ã€?

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

ä½ åº”è¯¥åœ¨ `header` å…ƒç´ ç»“æŸæ ‡ç­¾ä¹‹åŽæ·»åŠ ä¸€ä¸?`main` å…ƒç´ ã€?

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

åœ?`main` å…ƒç´ å†…éƒ¨ï¼Œåº”è¯¥æœ‰ä¸¤ä¸ª `section` å…ƒç´ ã€?

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

ä½ çš„ç¬¬ä¸€ä¸?`section` å…ƒç´ åº”åŒ…å«ä¸€ä¸?`id` å±žæ€§ï¼Œå€¼ä¸º `upcoming-events`ã€?

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

ä½ çš„ç¬¬äºŒä¸?`section` å…ƒç´ åº”åŒ…å«ä¸€ä¸?`id` å±žæ€§ï¼Œå€¼ä¸º `past-events`ã€?

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

åœ?`#upcoming-events` éƒ¨åˆ†ä¸­ï¼Œåº”è¯¥æœ‰ä¸€ä¸?`h2` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Upcoming Events`ã€?

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

åœ?`#upcoming-events` éƒ¨åˆ†ä¸­ï¼Œåœ?`h2` å…ƒç´ ä¸‹é¢åº”è¯¥æœ‰ä¸¤ä¸?`article` å…ƒç´ ã€?

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#upcoming-events` éƒ¨åˆ†ä¸­çš„ä¸¤ä¸ª `article` å…ƒç´ éƒ½åº”æœ‰ä¸€ä¸?`h3` å…ƒç´ ç”¨äºŽæ´»åŠ¨æ ‡é¢˜ã€?

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#upcoming-events` éƒ¨åˆ†ä¸­çš„ä¸¤ä¸ª `article` å…ƒç´ éƒ½åº”æœ‰ä¸€ä¸ªæ®µè½å…ƒç´ ç”¨äºŽæ´»åŠ¨æè¿°ã€?

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

åœ?`#past-events` éƒ¨åˆ†ä¸­ï¼Œåº”è¯¥æœ‰ä¸€ä¸?`h2` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Past Events`ã€?

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

åœ?`#past-events` éƒ¨åˆ†ä¸­ï¼Œåœ?`h2` å…ƒç´ ä¸‹é¢åº”è¯¥æœ‰ä¸¤ä¸?`article` å…ƒç´ ã€?

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`#past-events` éƒ¨åˆ†ä¸­çš„ä¸¤ä¸ª `article` å…ƒç´ éƒ½åº”æœ‰ä¸€ä¸?`h3` å…ƒç´ ç”¨äºŽæ´»åŠ¨æ ‡é¢˜ã€?

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`#past-events` éƒ¨åˆ†ä¸­çš„ä¸¤ä¸ª `article` å…ƒç´ éƒ½åº”æœ‰ä¸€ä¸ªæ®µè½å…ƒç´ ç”¨äºŽæ´»åŠ¨æè¿°ã€?

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`#past-events` éƒ¨åˆ†ä¸­çš„ä¸¤ä¸ª `article` å…ƒç´ éƒ½åº”åŒ…å«ä¸€ä¸ªå›¾åƒå…ƒç´ ã€?

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

`#past-events` éƒ¨åˆ†ä¸­çš„ä¸¤ä¸ªå›¾åƒå…ƒç´ çš?`src` å±žæ€§éƒ½åº”æŒ‡å‘ä¸€ä¸ªå›¾åƒæ–‡ä»¶ã€?

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

`#past-events` éƒ¨åˆ†ä¸­çš„ä¸¤ä¸ªå›¾åƒå…ƒç´ éƒ½åº”å…·æœ‰ `alt` å±žæ€§ï¼Œå…¶ä¸­åŒ…å«å›¾åƒæè¿°ã€?

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

æ¯ä¸ª `h3` å…ƒç´ éƒ½åº”åŒ…å«æ´»åŠ¨æ ‡é¢˜ã€?

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

æ¯ä¸ª `p` å…ƒç´ éƒ½åº”åŒ…å«æ´»åŠ¨æè¿°ã€?

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.MasterPuti.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.MasterPuti.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```


