---
id: 5d822fd413a79914d39e98d1
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

给你的 `.background-buildings` 元素设置 `width` 和 `height` 为 `100%`，使其填满父元素 `body` 的整个宽度和高度。

# --hints--

你应该使用 `background-buildings` 类来选择正确的元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

你的 `.background-buildings` 元素应该有 `100%` 的宽度。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

你的 `.background-buildings` 元素应该有 `100%` 的高度。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

--fcc-editable-region--
.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--

```