---
id: 5d822fd413a79914d39e98cf
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

你可以看到 `body`（它是页面上最内层的框）；它周围的框是 `html` 元素。通过给 `body` 设置 `height` 为 `100vh` 来让它填满整个视口。通过将 `margin` 设置为 `0` 来移除 `body` 的默认边距。最后，将 `overflow` 属性设置为 `hidden` 来隐藏当某些内容超出视口时出现的滚动条。

# --hints--

你应该使用 `body` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你的 `body` 应该有 `100vh` 的高度。

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

你的 `body` 应该有 `0` 的边距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

你的 `body` 应该将 `overflow` 属性设置为 `hidden`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--
body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--

```