---
id: 5d822fd413a79914d39e98da
title: 步骤 15
challengeType: 0
dashedName: step-15
---

# --description--

你的第一座建筑物现在看起来相当不错了。在 `.background-buildings` 容器中嵌套三个新的 `div` 元素，并按顺序给它们添加 `bb2`、`bb3` 和 `bb4` 类。这些将是背景中的另外三座建筑物。

# --hints--

你应该创建一个类名为 `bb2` 的 `div`。

```js
assert.exists(document.querySelector('div.bb2'));
```

你应该创建一个类名为 `bb3` 的 `div`。

```js
assert.exists(document.querySelector('div.bb3'));
```

你应该创建一个类名为 `bb4` 的 `div`。

```js
assert.exists(document.querySelector('div.bb4'));
```

你应该创建 3 个新的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div'), 9);
```

你应该将这些 `div` 元素放置在 `.background-buildings` 元素内。

```js
assert.lengthOf(document.querySelector('div.background-buildings')?.children, 4);
```

你应该按正确的顺序放置这些元素。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert.isTrue(__t('div.bb1','bb2') && __t('div.bb2','bb3') && __t('div.bb3','bb4'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
--fcc-editable-region--
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div class="bb4"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

```