---
id: 680900675ae3d54ee19590c5
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

现在来创建标签和输入框：在表单中添加一个文本为 `Hero Name` 的 `label`，以及一个 `type='text'` 的 `input`。

将该 `Hero Name` 输入的 `value` 设为 `heroName`，以连接到对应的 state 变量。

# --hints--

你需要创建一个 `label` 元素，太棒了！

```js
assert.exists(document?.querySelector("label"));
```

你的 `label` 元素文本应为 `Hero Name`，继续加油！

```js
const labelEl = document.querySelector("label")
assert.equal(labelEl?.textContent, "Hero Name");
```

你需要在该 `label` 内创建一个 `type='text'` 的 `input`，做得很好！

```js
const inputEl = document.querySelector("label > input");
assert.equal(inputEl?.getAttribute("type"), "text");
```

你的 `input` 元素应将 `value` 设为 `heroName` state 变量，你真聪明！

```js
assert.match(code, /value\s*=\s*\{\s*heroName\s*\}/)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
      </form>
    </div>
  )
};
```
