---
id: ab306dbdcc907c7ddfc30830
title: 创建深度展平工具
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

在这个实验中，你将实现一个数组展平算法。

展平数组意味着将任意深度的嵌套数组转换为单一的一维数组。这个过程按顺序提取所有元素，只展开数组。其他类型保持不变。

例如：

| 原始数组                   | 展平后                |
| -------------------------- | ------------------------ |
| `[[1], [], [2, [3]]]`      | `[1, 2, 3]`              |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]`      | `["baz", 1, 2, {}]`      |

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个名为 `steamrollArray` 的函数。
1. `steamrollArray` 函数应该接受一个参数：一个嵌套数组。
1. 该函数应该展平嵌套数组，考虑不同级别的嵌套。
1. 你的解决方案不应该使用 `Array.prototype.flat()` 或 `Array.prototype.flatMap()` 方法。
1. 不应该使用全局变量。

# --hints--

你应该有一个 `steamrollArray` 函数。

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` 应该返回 `["a", "b"]`。

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` 应该返回 `[1, 2, 3, 4]`。

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` 应该返回 `[1, 3, 4]`。

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` 应该返回 `[1, {}, 3, 4]`。

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

你的解决方案不应该使用 `Array.prototype.flat()` 或 `Array.prototype.flatMap()` 方法。

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

你不应该使用全局变量。

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```