---
id: 673339d96c56d36e963e888e
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个工作坊中，你将构建一个植物苗圃目录，同时练习使用JavaScript的Map和Set对象。

为了表示目录中的每种植物，你将使用对象，每个对象都有三个属性：

- `commonName`，植物的通用名称
- `scientificName`，植物的学名（或拉丁名）
- `cultivar`，植物的品种

首先为名为`Ballerina`的品种创建一个对象，它代表`西班牙薰衣草`（学名`Lavandula stoechas`），并将其赋值给名为`ballerina`的变量。

# --hints--

你应该有一个名为`ballerina`的变量。

```js
assert.exists(ballerina);
```

你应该为`ballerina`变量赋值一个对象。

```js
assert.isObject(ballerina);
```

你的`ballerina`对象应该有一个`commonName`属性，值为`"西班牙薰衣草"`。

```js
assert.strictEqual(ballerina.commonName, "西班牙薰衣草");
```

你的`ballerina`对象应该有一个`scientificName`属性，值为`"Lavandula stoechas"`。

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

你的`ballerina`对象应该有一个`cultivar`属性，值为`"Ballerina"`。

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

你的`ballerina`对象应该有三个键。

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```