---
id: 6734c10a327760a665b7d5b9
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

在接下来的几个步骤中，你将编写一个函数，用于在出售植物时更新目录。

创建一个名为`sellPlants`的函数，它接受三个参数：植物对象、花盆尺寸和要出售的花盆数量，按此顺序。

当指定的花盆数量大于可用花盆时，让你的函数返回`Not enough <pot-size> size pots for <item-name>. Only <pot-number> left.`。将`<pot-size>`替换为花盆尺寸，将`<pot-number>`替换为该尺寸的剩余花盆数量，将`<item-name>`替换为植物的学名后跟一个空格和用单引号括起来的品种名。

为了测试一切是否正常工作，将使用`ballerina`对象、`"small"`和`25`调用你的`sellPlants`函数的结果输出到控制台。

# --hints--

你不应该修改你的`catalog`map。

```js
const testCatalog = new Map();
testCatalog.set(ballerina, { small: 20, medium: 15, large: 12 });
testCatalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
testCatalog.set(willowVale, { small: 3, medium: 5, large: 0 });
testCatalog.set(hidcote, { small: 33, medium: 13, large: 18 });
testCatalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
testCatalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
assert.deepEqual(catalog, testCatalog);
```

你应该有一个名为`sellPlants`的函数。

```js
assert.isFunction(sellPlants);
```

你的`sellPlants`函数应该接受三个参数。

```js
assert.lengthOf(sellPlants, 3);
```

`sellPlants(ballerina, "small", 25)`应该返回`Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.`。

```js
assert.equal(sellPlants(ballerina, "small", 25), "Not enough small size pots for Lavandula stoechas 'Ballerina'. Only 20 left.")

// prevent hardcoding
assert.equal(sellPlants(hidcote, "large", 95), "Not enough large size pots for Lavandula angustifolia 'Hidcote'. Only 18 left.")

assert.equal(sellPlants(royalCrown, "medium", 95), "Not enough medium size pots for Lavandula dentata 'Royal Crown'. Only 22 left.")
```

你应该将`sellPlants(ballerina, "small", 25)`输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*sellPlants\s*\(\s*ballerina\s*,\s*("|')small\1\s*,\s*25\s*\)\s*\)/);
```

`sellPlants(ballerina, "small", 10)`应该返回`undefined`。

```js
assert.isUndefined(sellPlants(ballerina, "small", 10));
// prevent hardcoding
assert.isUndefined(sellPlants(hidcote, "large", 5));
assert.isUndefined(sellPlants(royalCrown, "medium", 5));
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });
--fcc-editable-region--
function sellPlants(plant, potSize, numPots) {
    const plantInfo = catalog.get(plant);
    if (numPots > plantInfo[potSize]) {
        return `Not enough ${potSize} size pots for ${plant.scientificName} '${plant.cultivar}'. Only ${plantInfo[potSize]} left.`;
    }
}

console.log(sellPlants(ballerina, "small", 25));

--fcc-editable-region--
console.log(catalog)
```