---
id: 6734a2385c52c54d916a8b86
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

向你的目录中添加另一个键值对。使用你的`prettyPolly`变量作为键，并使值为一个具有`small`、`medium`和`large`属性的对象，这些属性的值再次为整数。

# --hints--

你应该将你的`prettyPolly`对象设置为`catalog`map的键。

```js
assert.isTrue(catalog.has(prettyPolly));
```

`catalog`map中`prettyPolly`键的值应该是一个具有`small`、`medium`和`large`属性的对象。

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`catalog`map中与`prettyPolly`键关联的对象的`small`属性应该是一个正整数。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`catalog`map中与`prettyPolly`键关联的对象的`medium`属性应该是一个正整数。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`catalog`map中与`prettyPolly`键关联的对象的`large`属性应该是一个正整数。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你的`catalog`map中应该仍然有`ballerina`键。

```js
assert.isTrue(catalog.has(ballerina));
```

你的`catalog`map应该包含两个元素。

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 8, medium: 5, large: 3 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```