---
id: 6735145b0ee741461bdbaa49
title: 步骤 21
challengeType: 1
dashedName: step-21
---

# --description--

或者，你可以通过使用`entries`方法来遍历你的map对象，该方法返回一个包含键值对的map迭代器。

修改你的`displayCatalog`函数，遍历`catalog`map中的条目并将每个条目输出到控制台。

# --hints--

你的`displayCatalog`函数应该在`catalog`上调用`entries`方法。

```js
const catalogCopy = new Map(catalog);
let flag = false;
const temp = catalog.entries;
try {
    catalog.entries = () => {
        flag = true;
        return catalogCopy.entries()
    };
    displayCatalog();    
    assert.isTrue(flag);
} finally {
    catalog.entries = temp;
}
```

你的`displayCatalog`函数应该将`catalog`中包含的每个值输出到控制台。

```js
const testArr = [];
const temp = console.log;
try {
    console.log = obj => {testArr.push(obj)};
    displayCatalog();    
    assert.isNotEmpty(testArr);
    const ents = catalog.entries();
    for (let i = 0; i < catalog.size; i++) {        
        assert.deepEqual(testArr[i], ents.next().value)
    }
} finally {
    console.log = temp;
}
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "西班牙薰衣草",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "英国薰衣草",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "法国薰衣草",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const removePlant = plant => catalog.delete(plant);

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}
--fcc-editable-region--
const displayCatalog = () => {
    for (const entry of catalog.entries()) {
        console.log(entry);
    }
}

displayCatalog();

--fcc-editable-region--
```