---
id: 684cff0bf0a4ca0ad7e42990
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

下一步是圆润边框的角。要使用 Tailwind CSS 设置 `border-radius` 属性，你可以使用以下模式的工具类：`rounded-size`，其中 `size` 表示圆角半径。

下面是为元素应用超大圆角半径的示例：

```html
<p class="rounded-xl">这是边框圆角。</p>
```

在你的 `div` 开始标签中，应用大小为 `md` 的边框圆角。

# --hints--

你的 `div` 开始标签应该有 `rounded-md` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```