---
id: 684d09cad4ec900e41db2e74
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

现在 `div` 元素放置在预览屏幕的左上角。但如果在元素顶部添加一些空间就更好了。要使用 Tailwind CSS 为元素应用 `margin`，你可以使用以下模式：`m-number`，其中 `number` 表示要在四边应用的固定像素数。

如果你想只对一边应用边距，可以使用以下任何模式：`mt-number`（上边距）、`mb-number`（下边距）、`mr-number`（右边距）和 `ml-number`（左边距）。

下面是为元素应用 `margin-bottom` 的示例：

```html
<p class="mb-8">这是外边距。</p>
```

在你的 `div` 开始标签中，使用正确的工具类来应用 `4px` 的 `margin-top`。

# --hints--

你的 `div` 开始标签应该有 `mt-4` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```