---
id: 684d105e488dbd1166a717ae
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

如果你增大预览窗口的尺寸，你会注意到所有项目仍然在容器中左对齐。如果 flex 项目居中就更好了。

要使用 Tailwind CSS 在容器的主轴上定位 flex 项目，你可以使用以下模式：`justify-value`。`value` 可以是 `start`、`end`、`center`、`between` 等。任何在原生 CSS 中与 `justify-content` 一起使用的有效值都可以在这里与 Tailwind CSS 一起使用。

在你的 `div` 开始标签中，使用正确的工具类将 `justify-content` 属性设置为 `center`。

# --hints--

你的 `div` 开始标签应该有 `justify-center` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```