---
id: 684d0eeb9c81e5109e52fe8b
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

`div` 元素开始看起来更好了。但现在关闭 `button` 与段落元素靠得太近了。

要在 flex 布局中创建 flex 项目之间的空间，你可以使用以下模式：`gap-number`。`number` 表示 flex 布局中行和列之间的固定空间。

下面是为元素应用 `gap` 的示例：

```html
<p class="flex gap-8">这是间距。</p>
```

在你的 `div` 开始标签中，使用正确的工具类设置 `gap` 为 `4`。

# --hints--

你的 `div` 开始标签应该有 `gap-4` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```