---
id: 68501d32b69f9d7756334ffc
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

如果你放大预览窗口，你应该看到 `div` 元素在页面上水平居中。但对于较小的设备，`div` 元素的宽度占据了整个页面的宽度。

更新 `mx-auto` 类，使其仅应用于 `md` 设备及以上尺寸。

# --hints--

你的 `div` 开始标签应该有 `md:mx-auto` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

你的 `div` 开始标签中不应再有 `mx-auto` 类。

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```