---
id: a3566b1109230028080c9345
title: Design a Sum All Numbers Algorithm
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

在这个实验中，你需要设计一个求所有数字之和的算法。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个名为 `sumAll` 的函数，接受一个包含两个数字的数组。
1. `sumAll([n, m])` 应该返回 `n` 和 `m` 的和加上它们之间所有数字的和。最小的数字不一定在前面。例如，`sumAll([4,1])` 应该返回 `10`，因为 `1` 和 `4` 之间所有数字（包括两者）的和是 `10`。

# --hints--

你应该有一个名为 `sumAll` 的函数。

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` 应该返回一个数字。

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` 应该返回 `10`。

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` 应该返回 `10`。

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` 应该返回 `45`。

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` 应该返回 `45`。

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--
```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;

  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
