---
id: 66a83601cd819e37f0dccd14
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

为了添加一些介绍性的文字，你需要显示一个主标题，然后是一段关于留下反馈意见的简短说明。

在 `body` 元素内，添加一个 `header` 元素。

在 `header` 元素中，添加一个 `h1` 元素。 `h1` 元素的文本内容应该是 `Hotel Feedback Form`。

在你的 `h1` 元素下方，添加一个 `p` 元素。 `p` 元素的文本内容应该是 `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

你应该有一个 `header` 开始标签。

```js
assert.match(code, /<header>/i);
```

你应该有一个 `header` 结束标签。

```js
assert.match(code, /<\/header>/i);
```

你应该有一个 `h1` 开始标签。

```js
assert.match(code, /<h1>/i);
```

你应该有一个 `h1` 结束标签。

```js
assert.match(code, /<\/h1>/i);
```

你的 `h1` 元素应该包含文本 `Hotel Feedback Form`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

你的 `h1` 元素应该在 `header` 元素内。

```js
const header = document.querySelector('header');
const h1 = document.querySelector('h1');
assert.strictEqual(header.contains(h1), true);
```

你应该有一个 `p` 开始标签。

```js
assert.match(code, /<p>/i);
```

你应该有一个 `p` 结束标签。

```js
assert.match(code, /<\/p>/i);
```

你的 `p` 元素应该包含文本 `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const p = document.querySelector('p');
assert.strictEqual(p?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

你的 `p` 元素应该在 `header` 元素内。

```js
const header = document.querySelector('header');
const p = document.querySelector('p');
assert.strictEqual(header.contains(p), true);
```

你的 `p` 元素应该在 `h1` 元素之后。

```js
const h1 = document.querySelector('h1');
const p = document.querySelector('p');
assert.isTrue(h1.compareDocumentPosition(p) & Node.DOCUMENT_POSITION_FOLLOWING);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>Thank you for staying with us. Please provide feedback on your recent stay.</p>
    </header>
  </body>
</html>
```