---
id: 66a8290a27c2c625e2355042
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个精彩的讲座中，你将通过制作一个可爱的酒店反馈表单来练习使用 HTML 表单技巧！

首先添加 `<!DOCTYPE html>` 声明，然后添加带有 `lang` 属性且值为 `en` 的 `html` 元素。

在 `html` 元素中，添加一个 `head` 元素。

# --hints--

你应该有 `<!DOCTYPE html>` 声明。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你应该有一个 `html` 开始标签，并将语言设置为英语。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你应该有一个 `html` 结束标签。

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 声明应该放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

你应该有一个 `head` 开始标签。

```js
assert.match(code, /<head>/i);
```

你应该有一个 `head` 结束标签。

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 元素应该在你的 `html` 元素内。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```