---
id: 66a833e88d08593618f22285
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素内部，嵌套一个 `meta` 元素，将它的 `charset` 属性值设置为 `"UTF-8"`。

在那个 `meta` 元素下面，添加一个 `title` 元素。 `title` 元素的文本内容应该是 `Hotel Feedback Form`。

# --hints--

你应该有一个 `meta` 元素。

```js
assert.isNotNull(document.querySelector("meta"));
```

`meta` 元素是一个自闭合元素，它不应该有结束标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 标签应该有一个 `charset` 属性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 属性的值应该设置为 `"UTF-8"`。

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

你的 `meta` 元素应该嵌套在 `head` 元素中。

```js
assert.match(code, /<head>.*\s*<meta.*charset.*>.*\s*<\/head>/si);
```

你应该有一个 `title` 开始标签。

```js
assert.match(code, /<title>/i);
```

你应该有一个 `title` 结束标签。

```js
assert.match(code, /<\/title>/i);
```

你的 `title` 元素应该嵌套在 `head` 元素中。

```js
assert.match(code, /<head>\s*<meta.*>\s*<title>.*<\/title>\s*<\/head>/si);
```

你的 `title` 元素的文本内容应该为 `Hotel Feedback Form`。

```js
assert.match(code, /<title>\s*Hotel\s+Feedback\s+Form\s*<\/title>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
</html>
--fcc-editable-region--

--fcc-editable-region--
```