---
id: 66a84111965a0c46df6bbd0a
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在之前的课程中，你已经学会了如何将 `label` 元素与 `input` 元素关联起来，就像这样：

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

`label` 元素的 `for` 属性应该与对应的 `input` 元素的 `id` 属性匹配。这被称为显式关联。

在你的 `legend` 元素下方，添加一个 `label`元素，其文本内容为 `Name (required):`，并将它的 `for` 属性设置为 `"full-name"`。

然后在你的 `label` 元素下方，添加一个没有任何属性的 `input` 元素。在接下来的步骤中，你将添加必要的属性。

# --hints--

你应该有一个 `label` 开始标签。

```js
assert.match(code, /<label.*>/);
```

你应该有一个 `label` 结束标签。

```js
assert.match(code, /<\/label>/);
```

你的 `label` 元素应该放在 `fieldset` 元素内的 `legend` 元素下方。

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

你的 `label` 元素应该有一个 `for` 属性，值为 `"full-name"`。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

你的 `label` 元素的文本内容应该是 `Name (required):`。

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

你应该有一个 `input` 元素。

```js
assert.isNotNull(document.querySelector('input'));
```

你的 `input` 元素应该在 `label` 元素下方。

```js
assert.isNotNull(document.querySelector('label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```