---
id: 66a83bdcf425e7446900b7c4
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在之前的课程中，你已经学会了如何使用 `form` 元素，就像这样：

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

`action` 属性用于指定提交表单时表单数据的发送位置。

`method` 属性用于指定发送表单数据时使用的 HTTP 方法。最常见的方法是 `GET` 和 `POST`。

**注意**：你将在后面的模块中了解 HTTP 方法的工作原理。

在你的 `main` 元素内，添加一个 `form` 元素，将 `action` 属性设置为 `"https://hotel-feedback.MasterPuti.org"`，`method` 属性设置为 `"POST"`。

# --hints--

你应该有一个 `form` 开始标签。

```js
assert.match(code, /<form.*>/i);
```

你应该有一个 `form` 结束标签。

```js
assert.match(code, /<\/form>/i);
```

你的 `form` 元素应该有一个 `action` 属性，设置为 `"https://hotel-feedback.MasterPuti.org"`。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.MasterPuti.org');
```

你的 `form` 元素应该有一个 `method` 属性，设置为 `"POST"`。

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

你的 `form` 元素应该在 `main` 元素内。

```js
const mainElement = document.querySelector('main');
const formElement = document.querySelector('form');
assert.strictEqual(mainElement.contains(formElement), true);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        
      </form>
    </main>
  </body>
</html>
```