---
id: 66a9843525e9fa8046d709b7
title: 步骤 36
challengeType: 0
dashedName: step-36
---

# --description--

在酒店反馈表工作坊的最后一步，你需要在表单中添加一个提交按钮。

Remember that you learned how to work submit buttons in previous lessons.

添加一个 `button` 元素，`type` 属性设置为 `"submit"`，文本内容设置为 `Submit`。

这样，你的酒店反馈表就完成了！

# --hints--

你的表单中应该有一个 `button` 元素。

```js
assert.exists(document.querySelector('form button'));
```

你的按钮元素应该将 `type` 属性设置为 `"submit"`。

```js
assert.exists(document.querySelector('button[type="submit"]'));
```

你的按钮元素文本内容应该为 `Submit`。

```js
assert.strictEqual(document.querySelector('button')?.textContent.trim(), 'Submit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
          <input type="radio" id="no-option" name="hotel-stay">
          <label for="no-option">No</label>
        </fieldset>
        <fieldset>
          <legend>Why did you choose to stay at our hotel? (Check all that apply)</legend>
          <input type="checkbox" id="ads" name="ads" value="ads">
          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="recommendation" name="recommendation" value="recommendation">
          <label for="recommendation">Personal Recommendation</label>
          <input type="checkbox" id="location" name="location" value="location">
          <label for="location">Location</label>
          <input type="checkbox" id="reputation" name="reputation" value="reputation">
          <label for="reputation">Reputation</label>
          <input type="checkbox" id="price" name="price" value="price">
          <label for="price">Price</label>
        </fieldset>
        <fieldset>
          <legend>Ratings</legend>
          <label for="service">How was the service?</label>
          <select id="service" name="service">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option value="excellent" selected>Excellent</option>
          </select>
          <label for="food">How was the food?</label>
          <select id="food" name="food">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option value="excellent" selected>Excellent</option>
          </select>
          <label for="comments">Other Comments?</label>
          <textarea id="comments" name="comments" cols="30" rows="10"></textarea>
        </fieldset>
        <button type="submit">Submit</button>
      </form>
    </main>
  </body>
</html>
```

