---
id: 63ec3287b182ec0efe8a3135
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在 `#cart` 元素内，添加一个 `h2` 元素，文本为 `购物车`。

# --hints--

你应该在 `#cart` 元素内有一个 `h2` 元素。

```js
assert.exists(document.querySelector('#cart h2'));
```

你的 `h2` 元素应该有文本 `购物车`。

```js
assert.equal(document.querySelector('#cart h2')?.textContent, '购物车');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container"></div>
      </section>
      --fcc-editable-region--
      <section id="cart">

      </section>
      --fcc-editable-region--
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```