---
id: 63ec1a16f930b108b8a76806
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `main` 元素内，添加一个 `section` 元素，并给它一个 `id` 属性，值为 `products`。

# --hints--

你应该有一个 `section` 元素。

```js
assert.exists(document.querySelector('section'));
```

你的 `section` 元素应该在 `main` 元素内。

```js
assert.exists(document.querySelector('main section'));
```

你的 `section` 元素应该有一个 `id` 属性，值为 `products`。

```js
assert.equal(document.querySelector('main section')?.id, 'products');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```