---
id: 63ec1bbf5584390a7d08d41f
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `#products` 元素内，添加一个 `h2` 元素，文本为 `我们的商品`。

# --hints--

你应该有一个 `h2` 元素。

```js
assert.exists(document.querySelector('h2'));
```

你的 `h2` 元素应该在 `#products` 元素内。

```js
assert.exists(document.querySelector('#products h2'));
```

你的 `h2` 元素应该有文本 `我们的商品`。

```js
assert.equal(document.querySelector('#products h2')?.textContent, '我们的商品');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      --fcc-editable-region--
      <section id="products">

      </section>
      --fcc-editable-region--
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```