---
id: 63f03686c5ea863533ec71f4
title: 步骤 45
challengeType: 0
dashedName: step-45
---

# --description--

现在你需要在每次添加商品到购物车时调用 `displayCart` 函数。在 `addItem` 函数的末尾，在 `calculateTotal()` 调用之后，调用 `displayCart()`。

# --hints--

你应该在 `addItem` 函数的末尾调用 `displayCart()`。

```js
assert.match(code, /function\s+addItem\s*\(\s*id\s*\)\s*{\s*const\s+product\s*=\s*products\.find\(\s*$$\s*item\s*$$\s*=>\s*item\.id\s*===\s*id\s*$$\s*;\s*cart\.push\(\s*product\s*$$\s*;\s*calculateTotal\(\s*$$\s*;\s*displayCart\(\s*$$\s*}/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

  for (let i = 0; i < cart.length; i++) {
    const item = cart[i];
    const li = document.createElement("li");
    li.textContent = `${item.name} - $${item.price}`;
    document.getElementById("cart-items").appendChild(li);
  }

}

```