---
id: 63f033fdb1fbcc254999fcc3
title: 步骤 41
challengeType: 0
dashedName: step-41
---

# --description--

在将 `innerHTML` 设置为空字符串的代码下方，使用 `for` 循环遍历 `cart` 数组。在循环内部，声明一个 `item` 变量，并将其设置为当前商品。

# --hints--

你应该使用 `for` 循环遍历 `cart` 数组。

```js
assert.match(code, /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*cart\.length\s*;\s*i\s*\+\+\s*\)/);
```

你应该声明一个 `item` 变量。

```js
assert.match(code, /var\s+item|let\s+item|const\s+item/);
```

你应该将 `item` 变量设置为当前商品。

```js
assert.match(code, /var\s+item\s*=\s*cart\s*\[\s*i\s*\]|let\s+item\s*=\s*cart\s*\[\s*i\s*\]|const\s+item\s*=\s*cart\s*\[\s*i\s*\]/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

function displayCart() {
  document.getElementById("cart-items").innerHTML = "";

}

```