---
id: 63f0165121a9181342d5bc66
title: 步骤 25
challengeType: 0
dashedName: step-25
---

# --description--

现在你需要给每个添加到购物车的按钮添加一个唯一的标识符，这样你就可以在点击时识别是哪个按钮被点击了。给每个 `button` 元素添加一个 `id` 属性，值分别为 `add-1`、`add-2` 和 `add-3`。

# --hints--

第一个 `button` 元素应该有一个 `id` 属性，值为 `add-1`。

```js
assert.equal(document.querySelectorAll('button')[0]?.id, 'add-1');
```

第二个 `button` 元素应该有一个 `id` 属性，值为 `add-2`。

```js
assert.equal(document.querySelectorAll('button')[1]?.id, 'add-2');
```

第三个 `button` 元素应该有一个 `id` 属性，值为 `add-3`。

```js
assert.equal(document.querySelectorAll('button')[2]?.id, 'add-3');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti 购物车</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1>MasterPuti 购物车</h1>
    </header>
    <main>
      <section id="products">
        <h2>我们的商品</h2>
        <div class="products-container">
          --fcc-editable-region--
          <div class="product" data-id="1" data-price="12.99">
            <h3>西瓜</h3>
            <p>$12.99</p>
            <button>添加到购物车</button>
          </div>
          <div class="product" data-id="2" data-price="0.99">
            <h3>香蕉</h3>
            <p>$0.99</p>
            <button>添加到购物车</button>
          </div>
          <div class="product" data-id="3" data-price="0.89">
            <h3>橙子</h3>
            <p>$0.89</p>
            <button>添加到购物车</button>
          </div>
          --fcc-editable-region--
        </div>
      </section>
      <section id="cart">
        <h2>购物车</h2>
        <div class="cart-items"></div>
        <div class="cart-total">
          <h3>总计: $0</h3>
        </div>
      </section>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```