---
id: 63f02c6e18773921ba50aa53
title: 步骤 39
challengeType: 0
dashedName: step-39
---

# --description--

现在你需要更新 HTML 来显示购物车中的商品。在 `calculateTotal` 函数下方，声明一个 `displayCart` 函数。

# --hints--

你应该声明一个 `displayCart` 函数。

```js
assert.match(code, /function\s+displayCart|const\s+displayCart\s*=\s*\(.*\)\s*=>|let\s+displayCart\s*=\s*\(.*\)\s*=>|var\s+displayCart\s*=\s*\(.*\)\s*=>/);
```

# --seed--

## --seed-contents--

```js
const cart = [];

const products = [
  {
    id: 1,
    name: "西瓜",
    price: 12.99
  },
  {
    id: 2,
    name: "香蕉",
    price: 0.99
  },
  {
    id: 3,
    name: "橙子",
    price: 0.89
  }
];

function addItem(id) {
  const product = products.find((item) => item.id === id);
  cart.push(product);
  calculateTotal();
}

function calculateTotal() {
  let total = 0;

  for (let i = 0; i < cart.length; i++) {
    total += cart[i].price;
  }

  document.getElementById("total").textContent = total.toFixed(2);
}

```