---
id: 68ecfa1003d88990bb96a33d
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

描述列表由术语和详细信息列表组成。[dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素表示描述术语，而 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素表示描述详细信息。

你将添加到描述列表中的第一个浏览器是 Google Chrome。

在你的 [dl](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa0f03d88990bb96a33c.md#L23-L25) 元素内，添加一个包含文本 `Google Chrome` 的 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素。

在你的 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素下方，添加一个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素，包含以下文本：

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

你应该在 [dl](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa0f03d88990bb96a33c.md#L23-L25) 元素内有一个 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素。

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

你的 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素应该有文本 `Google Chrome`。

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

你应该在 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素下方有一个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素。

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

你的 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素应该包含文本 `This is a free web browser developed by Google and first released in 2008.`。

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        <dl>

        </dl>
        --fcc-editable-region--
    </body> 
</html>
```
