---
id: 68ecfa1003d88990bb96a33e
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

现在是时候向列表中添加第二个浏览器了。

在你的 `Google Chrome` 条目下方，添加另一个包含文本 `Firefox` 的 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素。

在你的 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素下方，添加一个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素，包含以下文本：

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

你的代码中应该有两个 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

你的第二个 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素应该有文本 `Firefox`。

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

你的代码中应该有两个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

你的第二个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素应该包含文本 `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`。

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--

            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
