---
id: 68ebdbacdd3fa474132cc975
title: 第 2 步
challengeType: 0
dashedName: step-2
---

# --description--

在 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素下面，添加一个文本为 [Browse our collection of amazing books!](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素。

# --hints--

你应该有一个 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素文本应该是：[Browse our collection of amazing books!](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58)。

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

你的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素应该在你的 [h1](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素下面。

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>

--fcc-editable-region--
```