---
id: 68eca4c4b0f952cf09fabe09
title: 第 12 步
challengeType: 0
dashedName: step-12
---

# --description--

在第二个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内，添加一个文本为 [Dave's Cooking Adventure](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca4c4b0f952cf09fabe09.md#L58-L58) 的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素。

# --hints--

你应该有一个 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素嵌套在第二个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

第二个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素文本应该是 [Dave's Cooking Adventure](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca4c4b0f952cf09fabe09.md#L58-L58)。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  --fcc-editable-region--
  <div class="card" id="dave-cooking-book">

  </div>
  --fcc-editable-region--
</div>
```