---
id: 68eca6e5a6759ed4ea0034dc
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 的元素内，添加两个 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素：

第一个按钮：

- Id: [view-cart-btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)
- Class: [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)
- Text: [View Cart](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)

第二个按钮：

- Id: [checkout-btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)
- Class: [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)
- Text: [Checkout](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)

恭喜你！你已成功使用 div、类和 ID 构建了书店页面的结构来组织你的内容。

# --hints--

你应该有两个 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素嵌套在 [.btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 元素内。

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 的元素内的第一个按钮应该有一个 [id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 属性，值为 [view-cart-btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 的元素内的第一个按钮应该有一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性，值为 [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

第一个按钮的文本应该是 [View Cart](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 的元素内的第二个按钮应该有一个 [id](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec9332a9b5b2b32487bd00.md#L58-L58) 属性，值为 [checkout-btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58)="btn-container" 的元素内的第二个按钮应该有一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性，值为 [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

第二个按钮的文本应该是 [Checkout](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6e5a6759ed4ea0034dc.md#L58-L58)。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
    <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    <button class="btn">Buy Now</button>
  </div>
</div>
--fcc-editable-region--
<p>Review your selections and continue to checkout.</p>
<div class="btn-container">

</div>
--fcc-editable-region--
```

# --solutions--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>

<div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
</div>

<p>Review your selections and continue to checkout.</p>

<div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
</div>
```