---
id: 68ec99e478211dc578699944
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

在第一个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素下面，添加一个具有以下文本的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素：

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

你应该有一个 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素嵌套在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内。

```js
assert.exists(document.querySelector('.card p'));
```

你的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素文本应该是 [This is an epic story of Sally and her dog Rex as they navigate through other worlds.](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec99e478211dc578699944.md#L58-L58)。

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

你的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素应该在你的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素下面。

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  --fcc-editable-region--
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>

  </div>
  --fcc-editable-region--
</div>
```