---
id: 68eca5c0065c82d256d29ca3
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

在第二张卡片内，添加一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性设置为 [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 且文本为 [Buy Now](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 的 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素。

现在两个 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素共享相同的 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58)，这意味着它们可以一起一致地设置样式。

# --hints--

你应该有一个 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素嵌套在第二个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

第二个卡片的 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素应该有一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性。

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

第二个卡片的 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素应该有一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 具有值 [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

第二个卡片的 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素文本应该是 [Buy Now](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  --fcc-editable-region--
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
    <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
  </div>
  --fcc-editable-region--
</div>
```