---
id: 68eca6364c5616d393389a30
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 的元素下面，添加一个具有以下文本的新 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素：

```md
Review your selections and continue to checkout.
```

在 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素下面，创建一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性设置为 [btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素。这个容器将分组你的导航按钮元素。

# --hints--

新添加的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素应该在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 的元素下面。

```js
assert.exists(document.querySelector('.card-container + p'));
```

你的新 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素文本应该是 [Review your selections and continue to checkout.](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58)。

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

你应该有一个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素。

```js
assert.exists(document.querySelector('div.btn-container'));
```

你的具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [btn-container](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca6364c5616d393389a30.md#L58-L58) 的 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素应该在新添加的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素之后。

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
    <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    <button class="btn">Buy Now</button>
  </div>
</div>

--fcc-editable-region--
```