---
id: 68eca5412b5dd9d06b3d6404
title: 第 13 步
challengeType: 0
dashedName: step-13
---

# --description--

在第二张卡片内的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素下面，添加一个具有以下文本的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素：

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

你应该有一个 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素嵌套在第二个具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

第二个卡片的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素文本应该是 [This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca5412b5dd9d06b3d6404.md#L58-L58)。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

第二个卡片的 [p](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ebdbacdd3fa474132cc975.md#L58-L58) 元素应该在你的 [h2](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec98b38d83a3c28dd30efe.md#L58-L58) 元素下面。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  --fcc-editable-region--
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>

  </div>
  --fcc-editable-region--
</div>
```