---
id: 68eca11e99e3c5c894ca9d69
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

[button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素用于在网页上创建可点击的按钮。按钮是交互元素，用户可以点击它们来执行操作。

你可以在 [div](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68e97fe79367ad7b5dd6c9cd.md#L58-L58) 元素内添加多个元素来分组相关内容。在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内添加一个 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素，给按钮添加一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性，设置为 [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)，文本为 [Buy Now](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)。

# --hints--

你应该有一个 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素嵌套在具有 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 为 [card](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6e8d0caee3afaaf142ef.md#L58-L58) 的元素内。

```js
assert.exists(document.querySelector('.card button'));
```

你的 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素应该有一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 属性。

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

你的 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素应该有一个 [class](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68ec6d9a315221aa31e54816.md#L58-L58) 具有值 [btn](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)。

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

你的 [button](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58) 元素文本应该是 [Buy Now](file:///F:/code/blocks_Chinese/blocks_english/workshop-bookstore-page/68eca11e99e3c5c894ca9d69.md#L58-L58)。

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  --fcc-editable-region--
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>

  </div>
  --fcc-editable-region--
</div>
```