---
id: 66d75dd0aa65a71600dc669b
title: Build an Inventory Management Program
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

在这个实验中，你将构建一个库存管理程序，允许你添加、更新、查找和从库存中删除产品。你将使用对象数组来表示你的库存，其中每个对象都有name和quantity作为键。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该声明一个名为`inventory`的空数组，该数组将存储产品对象，这些对象具有键`name`（值为小写字符串）和键`quantity`（值为整数）。
1. 你应该创建一个名为`findProductIndex`的函数，该函数以产品名称作为参数，并返回`inventory`数组中相应产品对象的索引。该函数应始终使用产品名称的小写形式来执行搜索。如果未找到产品，该函数应返回`-1`。
1. 你应该创建一个名为`addProduct`的函数，该函数以产品对象作为参数。
1. 如果产品已存在于库存中，`addProduct`函数应通过将传递给函数的数量添加到当前`quantity`来更新其`quantity`值，并将产品名称后跟空格和`quantity updated`记录到控制台。
1. 如果产品不存在于库存中，`addProduct`函数应将产品推送到`inventory`数组，并将产品名称记录到控制台，后跟空格和`added to inventory`。
1. 你应该创建一个名为`removeProduct`的函数，该函数以产品名称和数量作为参数。
1. `removeProduct`函数应从库存中的相应产品对象中减去传递的数量，并将字符串`Remaining <product-name> pieces: <product-quantity>`记录到控制台，其中`<product-name>`应替换为产品名称，`<product-quantity>`应替换为产品数量。
1. 如果减法后的数量为零，`removeProduct`应从库存中删除产品对象。如果库存中的数量不足以执行减法，`removeProduct`函数应将字符串`Not enough <product-name> available, remaining pieces: <product-quantity>`记录到控制台。
1. 如果要删除的产品不存在于库存中，`removeProduct`函数应将`<product-name> not found`记录到控制台。

**注意：** 为防止冲突，在运行测试时只保留用户故事中提到的日志记录。

# --hints--

你应该声明一个名为`inventory`的空数组。

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

你应该有一个名为`findProductIndex`的函数。

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")`应返回`inventory`数组中`name`等于`"flour"`的对象的索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")`应返回`inventory`数组中`name`等于`"flour"`的对象的索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")`应在`inventory`数组中未找到`name`等于`"flour"`的对象时返回`-1`。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

你应该有一个名为`addProduct`的函数。

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})`应将`5`添加到`inventory`数组中`name`等于`"flour"`的对象的`quantity`值。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})`应在`inventory`数组中找到`name`等于`"flour"`的对象时记录`flour quantity updated`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})`应在库存中未找到`name`等于`"flour"`的对象时将`{name: "flour", quantity: 5}`推送到`inventory`数组。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})`应在库存中未找到`name`等于`"flour"`的对象时记录`flour added to inventory`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

你应该创建一个名为`removeProduct`的函数。

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)`应在`inventory`数组中未找到`name`等于`"flour"`的对象时记录`flour not found`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)`应从`inventory`数组中`name`等于`"flour"`的对象的`quantity`值中减去`5`。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)`应在`inventory`等于`[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]`时将`Remaining flour pieces: 15`记录到控制台。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

如果减法后的`quantity`为零，`removeProduct`应从库存中删除产品对象。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)`应在`inventory`等于`[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]`时记录`Not enough flour available, remaining pieces: 5`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js
// 声明库存数组
const inventory = [];

// 查找产品索引的函数
function findProductIndex(productName) {
  // 你的代码在这里
}

// 添加产品的函数
function addProduct(product) {
  // 你的代码在这里
}

// 删除产品的函数
function removeProduct(productName, productQuantity) {
  // 你的代码在这里
}
```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}
```