---
id: 68e4c81577243e4db1676491
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将通过构建一个网页来练习使用语义化 HTML，该网页包含 Quincy Larson 关于找到开发工作的技巧。基本的 HTML 模板已经为你准备好了。

首先创建一个 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e4c81577243e4db1676491.md#L23-L23) 元素，文本为 `Quincy's Tips for Getting a Developer Job`。

# --hints--

你应该有一个 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e4c81577243e4db1676491.md#L23-L23) 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e4c81577243e4db1676491.md#L23-L23) 元素应该有文本 `Quincy's Tips for Getting a Developer Job`。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
