---
id: 68e65ec5348e40a739903e6d
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

现在，在 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e4c81577243e4db1676491.md#L23-L23) 元素下方创建一个段落元素。在这个段落中添加文本 `Learning to code is hard, but as Quincy Larson says, You can become a developer.`。

# --hints--

你应该有一个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素。

```js
assert.exists(document.querySelector('p'));
```

你的段落应该有文本 `Learning to code is hard, but as Quincy Larson says, You can become a developer.`。请仔细检查你的拼写。

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

你的段落应该在 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e4c81577243e4db1676491.md#L23-L23) 元素下方。

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
