---
id: 68e964b84f954c322376e592
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

如果引用的来源是网站，你可以使用 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 属性来引用它。这个属性的值应该是一个有效的 URL。虽然这个属性不会改变块引用的呈现方式，但它对为屏幕阅读器和搜索引擎提供更多关于引用的信息非常有帮助。

这是一个带有 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 属性的内联引用元素的例子：

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

为内联引用元素添加 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 属性，使用这个 URL：

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

你的 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素应该有一个 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 属性。

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

你的 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素的 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 属性应该有值 `https://www.freecodecamp.org/news/learn-to-code-book/`。

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
