---
id: 68e90335bd53a901c41dfc13
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在之前的课程中，你学到了 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素用于将引用文本与周围内容区分开来。

这里是一个例子：

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

当你使用 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素时，大多数现代浏览器会自动在内联引用周围添加引号。

在段落内，将 `You can become a developer.` 包装在一个内联引用元素中，保持段落的其余部分不变。

# --hints--

你应该有一个 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素。

```js
assert.exists(document.querySelector('q'));
```

你的 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素应该嵌套在 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e4c81577243e4db1676491.md#L23-L23) 元素下方的段落内。

```js
assert.exists(document.querySelector('h1 + p > q'));
```

你的 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素应该有文本 `You can become a developer.`。

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

在你的 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素内，文本前后不能有前导或尾随空白（第一个字符前或最后一个字符后不能有空格、制表符或换行符）。

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

你应该只在句子的适当部分添加 [q](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 标签，而不改变原始文本。[p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素仍应读作：`Learning to code is hard, but as Quincy Larson says, you can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
