---
id: 68e96ce52087043b0e999296
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在段落元素下方，添加一个 [main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e96ce52087043b0e999296.md#L23-L23) 元素，并在其中嵌套三个 [section](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68ea5ccc1229b85eca38a3e8.md#L25-L25) 元素。

# --hints--

你应该在段落元素下方有一个 [main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e96ce52087043b0e999296.md#L23-L23) 元素。

```js
const mainEl = document.querySelector('p + main');
assert.exists(mainEl);
```

你应该在 [main](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e96ce52087043b0e999296.md#L23-L23) 元素内有三个 [section](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68ea5ccc1229b85eca38a3e8.md#L25-L25) 元素。

```js
const sections = document.querySelectorAll('main > section');
assert.lengthOf(sections, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
