---
id: 68eaa9ef8a0867071f195f59
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

到目前为止，你一直在使用 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 属性来标注引用的来源，但这个属性实际上并没有向用户展示来源。

如果你想在视觉上标明来源，可以在块引用元素外添加一个引用元素 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28)。引用元素是一个 HTML 元素，你可以用它来标记被引用的创作作品的标题，比如书籍、文章、歌曲、电影、网站或研究论文。

这是一个例子：

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>—Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

在 [blockquote](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68ea7e7b916e8270483689b6.md#L28-L28) 元素下方的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素内，将 `How to Learn to Code and Get a Developer Job [Full Book]` 包装在 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 元素中。

# --hints--

你应该在 [blockquote](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68ea7e7b916e8270483689b6.md#L28-L28) 元素下方的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素内有一个 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 元素。

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

你的 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 元素应该有文本 `How to Learn to Code and Get a Developer Job [Full Book]`。

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

忽略你刚添加的 [cite](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e964b84f954c322376e592.md#L28-L28) 元素，[p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素的文本仍应为 `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`。

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
--fcc-editable-region--
        <p>
          —Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
        </p>
--fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
