---
id: 68eab6a2b37d2f13602d0c3f
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

你可以直接在块引用元素内写引用文本，就像你在第一个 [section](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68ea5ccc1229b85eca38a3e8.md#L25-L25) 元素中所做的那样。或者，你可以将其包装在一个或多个段落元素中。当文本有多个段落，但你想将它们保持在同一块引用元素内时，这很有帮助。

这是一个有两个段落的例子：

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  <p>So many people ask me each week: is college still worth it? In this 1-hour video I answer this question and other commonly asked questions about university.</p>
  <p>I've been in adult education for two decades at this point, and even though I'm not a labor market economist, I do feel confident enough to answer these questions.</p>
</blockquote>
```

在第二个部分中，在现有的块引用元素内，添加四个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素，按顺序包含以下文本：

1. `So much of getting a job is who you know.`
2. `It's OK to be an introvert, but you do need to push your boundaries.`
3. `Create GitHub, Twitter, LinkedIn, and Discord accounts.`
4. `Go to tech meetups and conferences. Travel if you have to.`

# --hints--

你应该在第二个部分的 [blockquote](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68ea7e7b916e8270483689b6.md#L28-L28) 元素内嵌套四个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素。

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(2) > blockquote > p');
assert.lengthOf(pElements, 4);
```

你的第一个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素应该有文本 `So much of getting a job is who you know.`。

```js
const firstPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'So much of getting a job is who you know.');
```

你的第二个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素应该有文本 `It's OK to be an introvert, but you do need to push your boundaries.`。

```js
const secondPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "It's OK to be an introvert, but you do need to push your boundaries.");
```

你的第三个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素应该有文本 `Create GitHub, Twitter, LinkedIn, and Discord accounts.`。

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Create GitHub, Twitter, LinkedIn, and Discord accounts.');
```

你的第四个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-quincys-job-tips/68e90335bd53a901c41dfc13.md#L26-L26) 元素应该有文本 `Go to tech meetups and conferences. Travel if you have to.`。

```js
const forthPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(4)');
assert.equal(forthPEl?.innerText.trim(), 'Go to tech meetups and conferences. Travel if you have to.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          —Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
--fcc-editable-region--
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">

        </blockquote>
--fcc-editable-region--
      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
