---
id: 66bba6fff611169359d9d36a
title: Build a Confidential Email Page
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. You should have a `main` element with an `id` of `email`.
1. 你的 `#email` 元素应该有 `50px` 的 `padding`，`50px` 的上边距，`500px` 的 `width`，以及 `2px` 的 `border`。
1. 你的 `#email` 元素的总宽度，包括内边距和边框，应该为 `500px`。
1. 你应该在 `#email` 元素内有两个 `div` 元素，一个 `id` 为 `confidential`，另一个 `id` 为 `top-secret`。
1. 你的 `#confidential` 和 `#top-secret` 元素应该有 `inline-block` 的 `display`。
1. 你的 `#confidential` 和 `#top-secret` 元素应该有 `padding`、左边距和 `border`。
1. `#confidential` 元素应该有 `CONFIDENTIAL` 文本。
1. `#top-secret` 元素应该有 `TOP SECRET` 文本。
1. 你的 `#confidential` 和 `#top-secret` 元素应该使用 CSS transform 进行旋转。
1. You should have at least three paragraph elements within your `#email` element.
1. 你应该在段落元素内至少有三个 `class` 为 `blurred` 的 `span` 元素。
1. 你应该有一个 `blurred` 选择器，使用 CSS filter 将元素模糊 `3px`。

# --hints--

你应该有一个 `id` 为 `email` 的 `main` 元素。

```js
assert.exists(document.querySelector('main#email'));
```

你应该有一个 `#email` 选择器，将其元素的各边 `padding` 设置为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

你应该有一个 `#email` 选择器，将其元素的 `margin-top` 设置为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

你应该有一个 `#email` 选择器，将其元素的 `width` 设置为 `500px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

你的 `#email` 元素应该有 `2px` 的边框。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

你的 `#email` 元素应该有 `border-box` 的 `box-sizing`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

你应该在 `#email` 元素内有一个 `id` 为 `confidential` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

你应该在 `#email` 元素内有一个 `id` 为 `top-secret` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

你的 `#confidential` 元素应该将其 `display` 设置为 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

你的 `#top-secret` 元素应该将其 `display` 设置为 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

你的 `#confidential` 元素应该有各边的 `padding`。

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#top-secret` 元素应该有各边的 `padding`。

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#confidential` 元素应该有左边距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

你的 `#top-secret` 元素应该有左边距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

你的 `#confidential` 元素应该有 `border`。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

你的 `#top-secret` 元素应该有 `border`。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

你的 `#confidential` 元素的文本应该是 `CONFIDENTIAL`。你可能省略了文本，或者有拼写错误。

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

你的 `#top-secret` 元素的文本应该是 `TOP SECRET`。你可能省略了文本，或者有拼写错误。

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

你的 `#confidential` 元素应该使用 CSS `transform` 来旋转元素。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

你的 `#top-secret` 元素应该使用 CSS `transform` 来旋转元素。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

你应该在 `#email` 元素内至少有三个段落元素。

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

你应该在段落内至少有三个 `class` 为 `blurred` 的 `span` 元素。

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

你的 `.blurred` 元素不应该是空的。

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

你应该有一个 `.blurred` 选择器，将其元素的 `filter` 设置为 `blur(3px)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}