---
id: 68e6698b29e562d21b1d260c
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将构建一个主要网络浏览器的列表。HTML 模板已经为你准备好了。

首先在你的页面中添加一个标题，使用 [body](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68e6698b29e562d21b1d260c.md#L12-L15) 元素内的 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68e6698b29e562d21b1d260c.md#L13-L13) 元素，标题内容为 `List of Major Web Browsers`。

# --hints--

你应该在 [body](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68e6698b29e562d21b1d260c.md#L12-L15) 元素内有一个 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68e6698b29e562d21b1d260c.md#L13-L13) 元素。

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

你的 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68e6698b29e562d21b1d260c.md#L13-L13) 元素应该包含文本 `List of Major Web Browsers`。

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--

    --fcc-editable-region--
    </body>
</html>
```
