---
id: 68ecfa0f03d88990bb96a33c
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

正如你在之前的课程中回忆的那样，描述列表用于以有组织且易于阅读的格式呈现术语和定义。

这里是一个例子：

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

在 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68e6698b29e562d21b1d260c.md#L13-L13) 元素下方，创建一个 [dl](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa0f03d88990bb96a33c.md#L23-L25) 元素。这将包含浏览器列表。

# --hints--

你应该有一个 [dl](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa0f03d88990bb96a33c.md#L23-L25) 元素。

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

你的 [dl](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa0f03d88990bb96a33c.md#L23-L25) 元素应该在 [h1](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68e6698b29e562d21b1d260c.md#L13-L13) 元素下方。

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        --fcc-editable-region--
        <h1>List of Major Web Browsers</h1> 

        --fcc-editable-region--
    </body> 
</html>
```
