---
id: 69038dd51a41ae215fcc030a
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

你将添加到列表中的第三个浏览器是 Safari 网络浏览器。

添加另一个包含文本 `Safari` 的 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素。

在你的 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素下方，添加一个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素，包含以下文本：

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

你的代码中应该有三个 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

你的第三个 [dt](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L27-L29) 元素应该有文本 `Safari`。

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

你的代码中应该有三个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

你的第三个 [dd](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-major-browsers-list/68ecfa1003d88990bb96a33d.md#L29-L31) 元素应该包含文本 `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`。

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            --fcc-editable-region--

            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
