---
id: 68e5293bd00d2fe134f58988
title: 步骤 17
challengeType: 20
dashedName: step-17
---

# --description--

为你的 `FixedAmountDiscount` 类实现 `is_applicable` 方法。固定金额折扣只应在经济上有意义时才应用。

如果 `product.price * 0.9` 大于 `self.amount`，则返回 `True`，否则返回 `False`。这确保了折扣不超过原价的 10%，这将是不现实的。

# --hints--

你的 `FixedAmountDiscount` 类应该有一个 `is_applicable` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("is_applicable")`)
})
```

你的 `is_applicable` 方法应该有参数 `self`、`product: Product` 和 `user_tier: str`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

你的 `is_applicable` 方法应该有返回类型提示 `-> bool`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

当 `product.price * 0.9` 大于 `self.amount` 时，你的 `is_applicable` 方法应该返回 `True`。

```js
({
  test: () => runPython(`
    # 添加存根方法以使类可实例化以进行测试
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(50)
    assert discount.is_applicable(test_product, "regular") == True
  `)
})
```

当 `product.price * 0.9` 小于或等于 `self.amount` 时，你的 `is_applicable` 方法应该返回 `False`。

```js
({
  test: () => runPython(`
    # 添加存根方法以使类可实例化以进行测试
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(95)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
