---
id: 68e5293bd00d2fe134f58985
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

由于 `PercentageDiscount` 继承自一个抽象类，你必须实现所有抽象方法。首先实现 `is_applicable` 方法。`is_applicable` 方法应该接受 `self`、`product: Product` 和 `user_tier: str` 作为参数。它应该有返回类型提示 `bool`。

这个方法应该确定百分比折扣是否有效。如果 `self.percent` 小于或等于 `70`，则返回 `True`，否则返回 `False`。这可以防止不现实的折扣，比如 90% 折扣。

# --hints--

你的 `PercentageDiscount` 类应该有一个 `is_applicable` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

你的 `is_applicable` 方法应该有参数 `self`、`product: Product` 和 `user_tier: str`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

当 `self.percent` 小于或等于 `70` 时，你的 `is_applicable` 方法应该返回 `True`。

```js
({
  test: () => runPython(`
    # 添加存根方法以使类可实例化以进行测试
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

当 `self.percent` 大于 `70` 时，你的 `is_applicable` 方法应该返回 `False`。

```js
({
  test: () => runPython(`
    # 添加存根方法以使类可实例化以进行测试
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

你的 `is_applicable` 方法应该有返回类型提示 `bool`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
