---
id: 680a82806c4c6bc3567f2708
title: 步骤 7
challengeType: 20
dashedName: step-7
---

# --description--

声明一个名为 `text` 的新变量，并将字符串 `hello world` 赋给它。这将是待加密的消息。

# --hints--

你应该声明一个名为 `text` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

你应该将字符串 `hello world` 赋给 `text`。记住要用单引号或双引号将字符串括起来，并注意字母大小写。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)

--fcc-editable-region--
```
