---
id: 680a41ef5bfdc059cfd9d4b5
title: 步骤 2
challengeType: 20
dashedName: step-2
---

# --description--

在之前的课程中，你学习了可以存储在变量中的不同数据类型。你刚刚赋了一个整数值。现在你需要赋一个字符串，字符串是用单引号或双引号括起来的字符序列：

```py
string_1 = 'I am a string'
string_2 = "I am also a string"
```

声明另一个名为 `alphabet` 的变量，并将字符串 `abcdefghijklmnopqrstuvwxyz` 赋给这个变量。

# --hints--

你应该声明一个名为 `alphabet` 的变量。请注意将变量名放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

你应该将字符串 `abcdefghijklmnopqrstuvwxyz` 赋给你的 `alphabet` 变量。记住要使用单引号或双引号将字符串括起来。

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
