---
id: 680a431692a5f863f726c438
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在这个工作坊中，你将构建一个*凯撒密码*。这是加密文本的最简单技术之一，它包括用字母表中固定位置的字母替换明文中的每个字母。例如，移动5位时，`a`将被`f`替换，`b`被`g`替换，以此类推。

要实现这个密码，你需要创建一个新版本的字母表，它从移动指示的位置开始。正如你在之前的课程中学到的，你可以使用字符串切片来提取字符串的一部分：

```py
fcc = 'MasterPuti'
print(fcc[8:]) # Camp
```

创建一个名为 `shifted_alphabet` 的变量，并使用切片语法将 `alphabet` 中从 `shift` 索引开始的部分赋给它。然后，调用内置的 `print()` 函数在终端上打印 `shifted_alphabet` 并查看结果。

# --hints--

你应该声明一个名为 `shifted_alphabet` 的新变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

你应该将 `alphabet[shift:]` 赋给 `shifted_alphabet`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

你应该调用 `print()` 函数，并将 `shifted_alphabet` 作为参数传入。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5

--fcc-editable-region--
```

