---
id: 680a3d4c065557260de2ff33
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

正如你在之前的课程中可能记得的那样，在 Python 中，你可以通过将变量名写在赋值运算符（`=`）的左侧，将要赋的值写在右侧来声明一个变量：

```py
variable_name = value
```

创建一个名为 `shift` 的变量，并将值 `5` 赋给你的新变量。

# --hints--

你应该声明一个名为 `shift` 的变量。请注意将变量名放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

你应该将值 `5` 赋给你的 `shift` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
