---
id: 680a549f60f5a3fedfc0edd9
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

从输出中可以看出，移动后的字母表从字母 `f` 开始，因为 `shift` 的值是 `5`。但是现在字母表的前五个字母 – `a`、`b`、`c`、`d` 和 `e` – 在移动后的字母表中缺失了，所以你需要将它们添加到移动后字母表的末尾。

`+` 运算符用于将两个或多个字符串连接在一起，这个过程称为连接，就像这样：

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

使用切片语法提取 `alphabet` 缺失的前部分，并将其连接到 `alphabet[shift:]`。

提醒一下，`sentence[start:stop]` 返回 `sentence` 从位置 `start`（包含）到 `stop`（不包含）的字符。

# --hints--

你应该将 `alphabet[shift:]` 和 `alphabet` 缺失的前部分的连接赋给 `shifted_alphabet` 变量。使用 `shift` 来指定切片的结束位置。

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:]
print(shifted_alphabet)
--fcc-editable-region--
```
