---
id: 680a7e200400089967b8b1f4
title: 步骤 6
challengeType: 20
dashedName: step-6
---

# --description--

`str.maketrans()` 方法接受两个长度相等的字符串，并返回一个转换表，该表将第一个字符串中的每个字符映射到第二个字符串中的相应字符。转换表中的每个字符都存储为 Unicode 序数，这是一个唯一标识字符的数字。

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

你将在后续课程中了解更多关于 `str.maketrans()` 返回的结构类型。现在，创建一个将 `alphabet` 中的字符映射到 `shifted_alphabet` 中字符的转换表，并将其赋给名为 `translation_table` 的变量。

# --hints--

你应该声明一个名为 `translation_table` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

你应该调用 `str.maketrans()`，传入 `alphabet` 和 `shifted_alphabet` 作为参数，并将结果赋给 `translation_table`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]

--fcc-editable-region--
```
