---
id: 6818ed1adc249bb91d373729
title: 步骤 14
challengeType: 20
dashedName: step-14
---

# --description--

在之前的课程中，你学习了 `return` 语句用于从函数中返回一个值，这样你就可以在代码的其他地方使用它：

```py
def spam():
    return 'Spam!'
```

从你的函数中移除 `print(encrypted_text)` 调用。然后，删除 `encrypted_text` 变量的声明，并直接从你的函数中返回 `text.translate(translation_table)`。

# --hints--

你不应该在 `caesar` 函数中有 `print(encrypted_text)`。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

你不应该在 `caesar` 函数中有 `encrypted_text` 变量。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

你的 `caesar` 函数应该返回 `text.translate(translation_table)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('MasterPuti', 3)
print(encrypted_text)

```

