---
id: 6818fcb250f34e62cab32c39
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

你可能已经注意到，虽然消息被加密了，但大写字符没有被处理。这是因为转换表不包含大写字母。

要解决这个问题，你需要使用 `upper()` 方法修改传递给 `str.maketrans()` 调用的字符串：

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

通过将每个参数的大写版本连接到参数本身来更新你的 `str.maketrans()` 调用。

# --hints--

你应该通过将每个参数的大写版本连接到参数本身来修改 `str.maketrans()` 调用的参数。

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)


encrypted_text = caesar('MasterPuti', 3)
print(encrypted_text)
```

