---
id: 6818eb3f98c0bca5af7c8d03
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

现在你的代码是可重用的。然而，`caesar` 函数在终端上打印消息并默认返回 `None`。为了证明这一点，在代码的末尾打印 `encrypted_text`。

# --hints--

你应该在函数体外打印 `encrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--
encrypted_text = caesar('MasterPuti', 3)

--fcc-editable-region--
```

