---
id: 6819c5c9faf2548e556163d5
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

正如你在之前的课程中回忆的那样，`not` 运算符是一个一元逻辑运算符，用于否定一个表达式：

```py
print(not True) # False
print(not False) # True
```

使用 `not` 运算符修复 `if` 语句的条件，使得当 `shift` *不是* 整数时，你的函数返回 `Shift must be an integer value.`。

# --hints--

你应该在 `if` 语句中使用 `not` 运算符来否定 `isinstance(shift, int)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    if isinstance(shift, int):
        return 'Shift must be an integer value.'
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('MasterPuti', 3)
print(encrypted_text)
```

