---
id: 6819ca464db219ab86f37dc4
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

正如你已经验证的，传递给加密文本的移动应该是正数。但它不能超过 `25`（`alphabet` 的最后一个索引）。

在 `if` 语句中添加第二个条件，验证 `shift` 是否大于 `25`。记住，Python 中的逻辑或操作是通过 `or` 运算符实现的。

另外，将返回的消息更新为 `Shift must be an integer between 1 and 25.`。

# --hints--

你应该使用 `or` 运算符在 `shift < 1` 之后添加第二个条件，检查 `shift` 是否大于 `25`。你不应该修改现有条件。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

你的第二个 `if` 语句应该返回 `Shift must be an integer between 1 and 25.`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('MasterPuti', 3)
print(encrypted_text)
```

