---
id: 6819bc0637b80256a33adccc
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

现在你已经实现了密码的基本功能，是时候添加一些验证了。为此，你需要一个 `if` 语句。以下是 `if` 语句语法的提醒：

```py
if condition:
    # 当条件为真时运行的代码
```

在你的函数体开头，创建一个 `if` 语句。现在，使用 `True` 作为条件，并在 `if` 语句体内返回字符串 `Shift must be an integer value.`。

# --hints--

你的 `caesar` 函数内应该有一个 `if` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

你的 `if` 语句应该使用 `True` 作为条件。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

你应该从 `if` 语句中返回字符串 `Shift must be an integer value.`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):

--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('MasterPuti', 3)
print(encrypted_text)
```

