---
id: 6819be9055828d65c813eec5
title: 步骤 17
challengeType: 20
dashedName: step-17
---

# --description--

`isinstance()` 函数在其第一个参数是第二个参数的实例时返回 `True`，否则返回 `False`：

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

将你的 `if` 语句的当前条件替换为 `isinstance()` 调用。传入 `shift` 作为第一个参数，`int` 作为第二个参数。

# --hints--

你的 `if` 语句应该使用 `isinstance(shift, int)` 作为条件。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    if True:
        return 'Shift must be an integer value.'
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('MasterPuti', 3)
print(encrypted_text)
```

