---
id: 68e5293bd00d2fe134f58984
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在是时候创建你的抽象类的具体实现了。创建一个名为 `PercentageDiscount` 的类，该类继承自 `DiscountStrategy`。

给它一个 `__init__` 方法，接受 `self` 和 `percent` 作为参数。给 `percent` 一个 `int` 类型提示，并将方法的返回类型设置为 `None`。在 `__init__` 方法中，将 `percent` 存储为实例属性。

# --hints--

你应该创建一个名为 `PercentageDiscount` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

你的 `PercentageDiscount` 类应该继承自 `DiscountStrategy`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

你的 `PercentageDiscount` 类应该有一个 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

你的 `__init__` 方法应该有参数 `self` 和 `percent: int`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

你的 `__init__` 方法应该有返回类型提示 `None`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

你应该在 `__init__` 方法中将 `percent` 分配给 `self.percent`。

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
