---
id: 68e524c94e2e4ab7b4bfc7e2
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个工作坊中，你将要构建一个可以对产品应用不同折扣策略的折扣计算器。系统将根据多个折扣规则为客户确定最佳价格。

首先定义一个名为 `Product` 的类。给你的类一个 `__init__` 方法，参数包括：`self`、`name` 和 `price`。在 `__init__` 方法内部，将 `name` 和 `price` 参数分配给同名的实例属性。

# --hints--

你应该创建一个名为 `Product` 的类。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

你的 `Product` 类应该有一个 `__init__` 方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

你的 `__init__` 方法应该有以下参数：`self`、`name` 和 `price`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

你应该在 `__init__` 方法中将 `name` 分配给 `self.name`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

你应该在 `__init__` 方法中将 `price` 分配给 `self.price`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
