---
id: 68e5293bd00d2fe134f58983
title: 步骤 11
challengeType: 20
dashedName: step-11
---

# --description--

现在在你的 `DiscountStrategy` 类中添加第二个名为 `apply_discount` 的抽象方法。使用 `@abstractmethod` 装饰器，并给它参数 `self` 和 `product`。

添加类型提示：`product` 应该是 `Product` 类型，方法应该返回一个 `float`。在方法体中使用 `pass`。

# --hints--

你的 `DiscountStrategy` 类应该有一个名为 `apply_discount` 的方法。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

你的 `apply_discount` 方法应该有 `@abstractmethod` 装饰器。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

你的 `apply_discount` 方法应该有参数 `self` 和 `product: Product`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

你的 `apply_discount` 方法应该有返回类型提示 `-> float`。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
