---
id: 68f119ad905b29ff59d858c8
title: 步骤 5
challengeType: 20
dashedName: step-5
---

# --description--

就像上一步一样，在你的 `__str__` 方法中，添加返回类型提示 `str`，因为此方法返回对象的字符串表示。

# --hints--

你的 `__str__` 方法应该有返回类型提示 `str`。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
