---
id: 5f33071498eb2472b87ddee4
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将通过制作咖啡馆菜单来学习CSS（层叠样式表）的基础知识。

首先添加你的 `<!DOCTYPE html>` 和 `html` 元素，并将 `lang` 属性设置为 `en`。

# --hints--

你应该有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

You should have an opening `<html lang="en">` with the language set to english.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

You should have a closing `</html>` tag.

```js
assert.match(code, /<\/html>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
