---
id: 5f33294a6af5e9188dbdb8f3
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

是时候添加一些菜单内容了。在现有的`body`元素内添加一个`main`元素。它最终将包含咖啡馆提供的咖啡和甜点的价格信息。

# --hints--

你的代码应该有一个开头的`<main>`标签。

```js
assert.match(code, /<main>/i);
```

你的代码应该有一个结尾的`</main>`标签。

```js
assert.match(code, /<\/main>/i);
```

你不应该改变你的`body`元素。确保你不会意外删除你的结束标签。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

你的`main`标签应该在你的`body`标签内。

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
