---
id: 5f332a88dc25a0fd25c7687a
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

咖啡馆的名字是`CAMPER CAFE`。所以，在你的`main`元素内添加一个`h1`元素。用大写字母给它咖啡馆的名字，让它更加突出。

# --hints--

你应该有一个开头的`<h1>`标签。

```js
assert.match(code, /<h1>/i);
```

你应该有一个结尾的`</h1>`标签。

```js
assert.match(code, /<\/h1>/i);
```

你不应该改变你现有的`main`元素。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

你的`h1`元素应该嵌套在你的`main`元素中。

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

你的`h1`元素应该有文本`CAMPER CAFE`。

```js
assert.match(code, /<h1>CAMPER CAFE<\/h1>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

