---
id: 5f344fc1520b6719f2e35605
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

菜单上将有两个部分，一个用于咖啡，一个用于甜点。在`main`元素内添加一个`section`元素，这样你就有地方放置所有可用的咖啡了。

# --hints--

你应该有一个开头的`<section>`标签。

```js
assert.match(code, /<section\s*>/i);
```

你应该有一个结尾的`</section>`标签。

```js
assert.match(code, /<\/section\s*>/i);
```

你不应该改变你现有的`main`元素。确保你没有删除结束标签。

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

你的`section`元素应该在你的`main`元素内。

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
