---
id: 5f356ed656a336993abd9f7c
title: 步骤 25
challengeType: 0
dashedName: step-25
---

# --description--

接下来，你想水平居中`#menu`。你可以通过将其`margin-left`和`margin-right`属性设置为`auto`来实现。把边距想象成元素周围的无形空间。使用这两个边距属性，将`#menu`元素在`body`元素内居中。

# --hints--

你应该将`margin-left`属性设置为`auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert.isTrue(hasMargin);
```

你应该将`margin-right`属性设置为`auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert.isTrue(hasMargin);
```

你应该将`#menu`的`margin-left`和`margin-right`属性设置为`auto`。

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert.equal(divMarginRight, 'auto');
assert.equal(divMarginLeft, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```
