---
id: 5f3c866dbf362f99b9a0c6d0
title: 步骤 37
challengeType: 0
dashedName: step-37
---

# --description--

`p`元素嵌套在具有`item`类属性的`article`元素中。你可以像这样设置嵌套在具有`item`类名的元素中的所有`p`元素的样式：

```css
.item p { }
```

使用上述选择器，添加一个值为`inline-block`的`display`属性，使`p`元素表现得更像内联元素。

# --hints--

你应该使用`.item p`选择器。

```js
const itemPStyle = new __helpers.CSSHelp(document).getStyle('.item p');
assert.exists(itemPStyle);
```

你应该将`display`属性设置为`inline-block`。

```js
const hasDisplay = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.display === 'inline-block');
assert.isTrue(hasDisplay);
```

你的`.item p`选择器应该将`display`属性设置为`inline-block`。

```js
const itemPDisplay = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('display');
assert.equal(itemPDisplay, 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

--fcc-editable-region--
.item p {
  display: inline-block;
}
--fcc-editable-region--

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```

