---
id: 5f3c866dd0d0275f01d4d847
title: 步骤 39
challengeType: 0
dashedName: step-39
---

# --description--

好吧，那没有起作用。将`p`元素样式设置为`inline-block`并将它们放在不同的行上会在第一个`p`元素的右侧创建一个额外的空间，导致第二个元素移到下一行。

解决这个问题的一种方法是使每个`p`元素的宽度略小于`50%`。所以，将每个类的`width`值更改为`49%`，看看会发生什么。

# --hints--

你应该在`.flavor`选择器中将`width`属性设置为`49%`。

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert.equal(flavorWidth, '49%');
```

你应该在`.price`选择器中将`width`属性设置为`49%`。

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert.equal(priceWidth, '49%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
  width: 49%;
}

.price {
  text-align: right;
  width: 49%;
}
--fcc-editable-region--
```


