---
id: 5f716ad029ee4053c7027a7a
title: Step 49
challengeType: 0
dashedName: step-49
---

# --description--

在你的 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f716ad029ee4053c7027a7a.md#L83-L84) 元素内嵌套两个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f716ad029ee4053c7027a7a.md#L65-L66) 元素。第一个的文本应该是 `Donut`，第二个的文本是 `1.50`。将它们放在同一行，确保它们之间没有空格。

# --hints--

你不应该改变现有的 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f716ad029ee4053c7027a7a.md#L83-L84) 元素。

```js
assert.lengthOf(document.querySelectorAll('article'), 6);
```

你的新 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f716ad029ee4053c7027a7a.md#L83-L84) 元素应该有两个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f716ad029ee4053c7027a7a.md#L65-L66) 元素。

```js
const newArticle = [...document.querySelectorAll('article')].at(-1);
const paragraphs = newArticle?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphs, 2);
```

你的第一个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f716ad029ee4053c7027a7a.md#L65-L66) 元素应该有文本 `Donut`。

```js
const newArticle = [...document.querySelectorAll('article')].at(-1);
const paragraph = newArticle?.querySelector(`:scope ${'p'}`);
assert.match(paragraph?.innerText.trim(), /Donut/i);
```

你的第二个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f716ad029ee4053c7027a7a.md#L65-L66) 元素应该有文本 `1.50`。

```js
const newArticle = [...document.querySelectorAll('article')].at(-1);
const paragraphs = newArticle?.querySelectorAll(`:scope ${'p'}`);
assert.match(paragraphs?.[1]?.innerText.trim(), /1\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
--fcc-editable-region--
          <article class="item">
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```


