---
id: 6564aee9c077774ea49c3faf
title: 步骤 2
challengeType: 20
dashedName: step-2
---

# --description--
归并排序的核心节奏就是三个动作：

1. 分：把未排序的序列不断拆成更小的部分。
2. 治：各个小部分分别排序。
3. 合：把排好序的部分重新合并成整体。

这一过程会“递归地”发生，直到所有子部分都被合并完成。现在我们来做第一步——确定分割点。

请把之前的 `pass` 改成定义变量 `middle_point`。使用整数除法运算符 `//` 计算 `len(array) // 2`，并把结果赋值给 `middle_point`。这就是列表的“中间线”。（别忘了缩进哦～）

# --hints--
函数体里现在不应该再有 `pass`。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

需要声明变量 `middle_point` 并赋值为 `len(array) // 2`。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
