---
id: 656627b47bd2d2a4afbc945d
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--
上一小步你已经拿到了中间点，现在正式开始“对半拆分”。

使用切片语法取出左半部分：把 `array` 从开头到 `middle_point` 的部分切出来，赋值给变量 `left_part`。这是构建分治结构的关键一步，继续保持你的专注力，超棒！

# --hints--
需要在函数里定义变量 `left_part`。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

`left_part` 必须通过切片获得左半部分（`array[:middle_point]` 或等价写法）。

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
