---
id: 65670d1ef177e7e2b76d9528
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--
刚才的 `if/else` 已经把较小的元素放进主数组啦。每写入一个元素，就意味着“结果数组的当前位置”已经被占用，需要往后移动继续写入。

请在 `if/else` 之后（仍在同一个 `while` 循环里）对 `sorted_index` 做自增：加 1。这样指针才能持续前进，排序过程不会停滞～⚙️

# --hints--
需要在 `if/else` 之后为 `sorted_index` 执行自增 1 操作。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_bodies()[0].find_body()[1]
    cond.is_equivalent("sorted_index += 1") or cond.is_equivalent("sorted_index = sorted_index + 1") or cond.is_equivalent("sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
