---
id: 6569c2cbf6c993ea8cd85682
title: æ­¥éª¤ 25
challengeType: 20
dashedName: step-25
---

# --description--
çŽ°åœ¨ï¼Œåˆ—è¡¨å·²ç»è¢«æˆåŠŸæŽ’åºï¼ðŸŽ‰

è¯·å…ˆæ‰“å°æ ‡é¢˜ï¼š`'Sorted array: '`ï¼Œç„¶åŽå†æ‰“å°æŽ’åºåŽçš„ `numbers`ã€‚è¿™æ ·å°±å®Œæˆäº†æ•´ä¸ªå½’å¹¶æŽ’åºçš„å±•ç¤ºæµç¨‹ã€‚

è‡³æ­¤ï¼Œå½’å¹¶æŽ’åºå®žçŽ°å®Œæ¯•ã€‚ä½ å·²ç»æŽŒæ¡äº†ï¼šé€’å½’æ‹†åˆ†ã€åŒæŒ‡é’ˆå½’å¹¶ã€å‰©ä½™å°¾éƒ¨æ‹¼æŽ¥ã€ä»¥åŠè¾¹ç•Œæ¡ä»¶æŽ§åˆ¶ã€‚å¤ªæ£’å•¦ï¼Œç»§ç»­æŠŠè¿™ç§ç»†è‡´ä¸Žè€å¿ƒåº”ç”¨åˆ°æ›´å¤šç®—æ³•ä¸­å§ï¼ðŸŒŸ

# --hints--
éœ€è¦å…ˆæ‰“å° `print('Sorted array: ')`ï¼Œå†æ‰“å°æŽ’åºåŽçš„ `numbers`ã€‚

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].is_ordered("merge_sort(numbers)", "print('Sorted array: ')", "print(numbers)")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1

    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1


--fcc-editable-region--
if __name__ == '__main__':
    numbers = [4, 10, 6, 14, 2, 1, 8, 5]
    print('Unsorted array: ')
    print(numbers)
    merge_sort(numbers)

--fcc-editable-region--
```

# --solutions--

```py
def merge_sort(array):
    if len(array) <= 1:
        return

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1

    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1


if __name__ == '__main__':
    numbers = [4, 10, 6, 14, 2, 1, 8, 5]
    print('Unsorted array: ')
    print(numbers)
    merge_sort(numbers)
    print('Sorted array: ')
    print(numbers)

```
