---
id: 6569b20f829b7e69d43c232a
title: æ­¥éª¤18
challengeType: 20
dashedName: step-18
---

# --description--
è¯¥å¾ªçŽ¯çš„æœ€åŽä¸€ä»¶äº‹ï¼šåˆ«å¿˜äº†ä¹Ÿè¦è®©ç»“æžœå†™å…¥ä½ç½® `sorted_index` å‰è¿›ã€‚

è¯·å°† `sorted_index` åšè‡ªå¢žï¼ˆåœ¨èµ‹å€¼å’Œå·¦ä¾§ç´¢å¼•è‡ªå¢žä¹‹åŽï¼‰ã€‚å®Œæˆè¿™ä¸€å°æ­¥ï¼Œå·¦ä¾§å‰©ä½™çš„å…ƒç´ å†™å…¥é€»è¾‘å°±å½»åº•é—­çŽ¯å•¦ï¼ðŸŽ‰

# --hints--
ç¡®ä¿åœ¨å†™å…¥ä¸Žç§»åŠ¨ `left_array_index` ä¹‹åŽï¼Œå†è‡ªå¢ž `sorted_index`ã€‚

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()

    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = 1 + sorted_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1

--fcc-editable-region--
```
