---
id: 6577562501feabdf0984a184
title: 步骤 11
challengeType: 20
dashedName: step-11
---

# --description--
当前把左侧较小元素写入结果后，别忘了让左侧指针继续前进：对 `left_array_index` 做自增 1，这样才能比较下一个元素。不然指针停在原地就“卡住”啦～⚙️

# --hints--
在写入左侧元素之后再执行 `left_array_index += 1`。

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body()[1]

    cond.is_equivalent("left_array_index += 1") or cond.is_equivalent("left_array_index = left_array_index + 1") or cond.is_equivalent("left_array_index = 1 + left_array_index")
    `)) 
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]

--fcc-editable-region--
```
