---
id: 6577254891048c8f2c19e961
title: æ­¥éª¤ 6
challengeType: 20
dashedName: step-6
---

# --description--
å·¦è¾¹é€’å½’å®Œæˆäº†ï¼Œä¹Ÿè¦è®©å³è¾¹è·Ÿä¸Šæ­¥ä¼ï½žè¯·åœ¨å‡½æ•°ä½“åº•éƒ¨å†æ¬¡è°ƒç”¨ `merge_sort(right_part)`ï¼Œè®©å³åŠéƒ¨åˆ†ä¹Ÿè¿›å…¥åŒæ ·çš„æ‹†åˆ†ä¸Žåˆå¹¶æµç¨‹ã€‚ä½ çš„é€’å½’ç»“æž„æ­£åœ¨å®Œæ•´æˆåž‹ï¼Œç»§ç»­å†²ï¼ðŸŒ€

# --hints--
éœ€è¦åœ¨åº•éƒ¨è°ƒç”¨ï¼š`merge_sort(right_part)`ã€‚

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)

--fcc-editable-region--
```
