---
id: 657f3a661730891aa64f3568
title: æ­¥éª¤ 4
challengeType: 20
dashedName: step-4
---

# --description--
ç»§ç»­æ‹†åˆ†åˆ—è¡¨çš„å¦ä¸€åŠå•¦ï½žè¯·ç”¨åˆ‡ç‰‡è¯­æ³•èŽ·å–å³åŠéƒ¨åˆ†ï¼šä»Ž `middle_point` åˆ°ç»“å°¾ï¼ŒæŠŠå®ƒèµ‹å€¼ç»™å˜é‡ `right_part`ã€‚å·¦å³ä¸¤è¾¹éƒ½å‡†å¤‡å¥½ï¼ŒåŽç»­é€’å½’æ‰èƒ½é½æ­¥å‘å‰ã€‚ðŸŒ—

# --hints--
éœ€è¦å®šä¹‰å˜é‡ `right_part`ã€‚

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

`right_part` å¿…é¡»é€šè¿‡åˆ‡ç‰‡å¾—åˆ°å³åŠéƒ¨åˆ†ï¼š`array[middle_point:]`ã€‚

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
