---
id: 689f7996426c1534fa5bea66
title: Build a Golf Score Translator
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

在高尔夫球运动中，每个洞都有一个"标准杆数"(`par`)，这意味着高尔夫球手为了将球击入洞中完成比赛，预期需要的平均"击球数"(`strokes`)。根据你的"击球数"比"标准杆数"多或少多少，会有不同的昵称。

在这个实验中，你将编写一个函数，将"标准杆数"和"击球数"转换为它们的昵称。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该创建一个名为`golfScore`的函数。
1. `golfScore`应该接受两个数值参数，分别是该洞的标准杆数和实际击球数。
1. `golfScore`应该返回一个字符串。
1. 如果`strokes`是`1`，`golfScore`应该返回`"一杆进洞！"`。
1. 如果`strokes`小于或等于`par`减`2`，`golfScore`应该返回`"老鹰球"`。
1. 如果`strokes`等于`par`减`1`，`golfScore`应该返回`"小鸟球"`。
1. 如果`strokes`等于`par`，`golfScore`应该返回`"标准杆"`。
1. 如果`strokes`等于`par`加`1`，`golfScore`应该返回`"柏忌"`。
1. 如果`strokes`等于`par`加`2`，`golfScore`应该返回`"双柏忌"`。
1. 如果`strokes`大于或等于`par`加`3`，`golfScore`应该返回`"回家吧！"`。

# --hints--

你应该创建一个名为`golfScore`的函数。

```js
assert.isFunction(golfScore);
```

`golfScore`应该接受两个参数。

```js
assert.lengthOf(golfScore, 2);
```

`golfScore`应该返回一个字符串

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)`应该返回字符串`一杆进洞！`

```js
assert.strictEqual(golfScore(1,1), '一杆进洞！');
```

`golfScore(3, 1)`应该返回字符串`一杆进洞！`

```js
assert.strictEqual(golfScore(3,1), '一杆进洞！');
```

`golfScore(4, 1)`应该返回字符串`一杆进洞！`

```js
assert.strictEqual(golfScore(4,1), '一杆进洞！');
```

`golfScore(5, 1)`应该返回字符串`一杆进洞！`

```js
assert.strictEqual(golfScore(5,1), '一杆进洞！');
```

`golfScore(4, 2)`应该返回字符串`老鹰球`

```js
assert.strictEqual(golfScore(4,2), '老鹰球');
```

`golfScore(5, 2)`应该返回字符串`老鹰球`

```js
assert.strictEqual(golfScore(5,2), '老鹰球');
```

`golfScore(3, 2)`应该返回字符串`小鸟球`

```js
assert.strictEqual(golfScore(3,2), '小鸟球');
```

`golfScore(4, 3)`应该返回字符串`小鸟球`

```js
assert.strictEqual(golfScore(4,3), '小鸟球');
```

`golfScore(5, 4)`应该返回字符串`小鸟球`

```js
assert.strictEqual(golfScore(5,4), '小鸟球');
```

`golfScore(3, 3)`应该返回字符串`标准杆`

```js
assert.strictEqual(golfScore(3,3), '标准杆');
```

`golfScore(4, 4)`应该返回字符串`标准杆`

```js
assert.strictEqual(golfScore(4,4), '标准杆');
```

`golfScore(5, 5)`应该返回字符串`标准杆`

```js
assert.strictEqual(golfScore(5,5), '标准杆');
```

`golfScore(3, 4)`应该返回字符串`柏忌`

```js
assert.strictEqual(golfScore(3,4), '柏忌');
```

`golfScore(4, 5)`应该返回字符串`柏忌`

```js
assert.strictEqual(golfScore(4,5), '柏忌');
```

`golfScore(5, 6)`应该返回字符串`柏忌`

```js
assert.strictEqual(golfScore(5,6), '柏忌');
```

`golfScore(3, 5)`应该返回字符串`双柏忌`

```js
assert.strictEqual(golfScore(3,5), '双柏忌');
```

`golfScore(4, 6)`应该返回字符串`双柏忌`

```js
assert.strictEqual(golfScore(4,6), '双柏忌');
```

`golfScore(5, 7)`应该返回字符串`双柏忌`

```js
assert.strictEqual(golfScore(5,7), '双柏忌');
```

`golfScore(3, 7)`应该返回字符串`回家吧！`

```js
assert.strictEqual(golfScore(3,7), '回家吧！');
```

`golfScore(4, 8)`应该返回字符串`回家吧！`

```js
assert.strictEqual(golfScore(4,8), '回家吧！');
```

`golfScore(5, 9)`应该返回字符串`回家吧！`

```js
assert.strictEqual(golfScore(5,9), '回家吧！');
```

# --seed--

## --seed-contents--

```js
const names = ["一杆进洞！", "老鹰球", "小鸟球", "标准杆", "柏忌", "双柏忌", "回家吧！"];

function golfScore(par, strokes) {
  // 你的代码在这里
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "一杆进洞！";
  }

  if (strokes <= par - 2) {
    return "老鹰球";
  }

  if (strokes === par - 1) {
    return "小鸟球";
  }

  if (strokes === par) {
    return "标准杆";
  }

  if (strokes === par + 1) {
    return "柏忌";
  }

  if(strokes === par + 2) {
    return "双柏忌";
  }

  return "回家吧！";
}
```