---
id: 68c28b7cc3ab8f4daf144b5e
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--
继续照刚才的模式，为 `_level` 创建一个对应的 getter（属性访问）。让级别也能优雅读取～

# --hints--
类里需要一个 `level(self)` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

`level` 方法要用 `@property` 装饰。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

`level` 方法要返回 `self._level`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name


--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
