---
id: 68c27eb292bfec6b9402d58f
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--
为 `Employee` 类增加一个 `__str__` 方法，让它返回格式 `name: level` 的 f-string，其中 `name` 和 `level` 来自实例的实际值。

然后在控制台打印 `charlie_brown`，看看美观的字符串输出～

# --hints--
类里需要一个 `__str__(self)` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

`__str__` 返回的内容要形如：`Frank: dreamer`（即 `self._name: self._level`）。

```js
({ test: () => runPython(`
    import io
    import sys

    captured_output = io.StringIO()
    sys.stdout = captured_output

    empl = Employee("Frank", "dreamer")
    print(empl)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Frank: dreamer" in output
`) })
```

要在控制台打印 `charlie_brown`。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level



charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
