---
id: 68c2d31a86d5672ca49d4521
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--
现在给 `__init__` 方法加一点健壮性校验：在方法开头写一个 `if`，判断 `name` 或 `level` 是否不是 `str` 类型。

如果类型不正确，就抛出：`TypeError("'name' and 'level' attribute must be of type 'str'.")`。

这能帮助你及早发现错误输入，让系统更可靠，棒棒！

# --hints--
`__init__` 中需要一个 `if` 判断。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

条件需判断 `name` 与 `level` 是否都为字符串（或任一不是字符串）。

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

不符合时要抛出对应的 `TypeError` 信息。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):

        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
