---
id: 68c80c1a8e58febb77120bbc
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--
移除最后那条用于展示 `repr` 的打印语句，让代码保持简洁。前面的学习效果你已经验证过啦～

# --hints--
代码里不应该再出现 `print(repr(charlie_brown))`。

```js
({ test: () => assert(runPython(`not _Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
print(repr(charlie_brown))
--fcc-editable-region--
```
