---
id: 68c8002beb157fe47d2e7699
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--
Setter 用于“安全地”修改属性值，经常能加入校验逻辑。语法格式是：在已有 `@property` 的属性后定义同名方法并用 `@属性名.setter` 装饰。例如：

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail'  # 这里调用 setter
print(p.name)  # Abigail
```

现在在你的 getter 下面创建一个 `name(self, new_name)` 方法，使用 `@name.setter` 装饰，并在内部执行 `self._name = new_name`。给属性赋予“安全更新”的能力，继续冲刺！

# --hints--
类里需要一个带参数 `self, new_name` 的 `name`（setter）方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

该方法要使用 `@name.setter` 装饰。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

方法内部需执行 `self._name = new_name`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"
--fcc-editable-region--
    @property
    def name(self):
        return self._name


--fcc-editable-region--
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary


charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
