---
id: 66a207974c806a19d6607073
title: 构建图书库存应用
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
---

# --description--

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个文本为`Book Inventory`的`h1`元素。
2. 你应该有一个具有`Title`、`Author`、`Category`、`Status`和`Rate`列的`table`元素。
3. 你的表格应该至少有四行，第一行用于列标题，其余行填充图书信息。
4. 表格主体中的每个表格行都应该具有`read`、`to-read`或`in-progress`类中的一个。
5. `Status`列的`td`元素应该包含一个具有`status`类的`span`元素，围绕文本`Read`、`To Read`或`In Progress`，具体取决于该行的类。
6. `Rate`列的`td`元素应该包含一个具有`rate`类的`span`元素，包裹三个空的`span`元素。
7. 放置在`.read`行中的`.rate`元素应该有一个附加类，值为`one`、`two`或`three`中的一个，具体取决于个人评分。这个值应该在`rate`之后。
8. 你应该创建三个属性选择器来针对具有`read`、`to-read`和`in-progress`类的元素，并设置它们的`background-image`属性以使用你选择的`linear-gradient`。
9. 你应该将每个`span`元素的`display`属性设置为`inline-block`。
10. 你应该使用属性选择器来针对具有`to-read`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`border`和`background-image`属性。
11. 你应该使用属性选择器来针对具有`read`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`border`和`background-image`属性。
12. 你应该使用属性选择器来针对具有`in-progress`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`border`和`background-image`属性。
13. 你应该使用属性选择器来针对具有`status`类的`span`元素和具有以`rate`开头的类值的`span`元素，并设置它们的`height`、`width`和`padding`属性。
14. 你应该使用属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素，并设置它们的`border`、`border-radius`、`margin`、`height`、`width`和`background-color`属性。
15. 你应该使用属性选择器来针对具有`one`作为其类值一部分的`span`元素的第一个后代，并设置其`background-image`属性以使用`linear-gradient`。
16. 你应该使用属性选择器来针对具有`two`作为其类值一部分的`span`元素的前两个后代，并设置它们的`background-image`属性以使用`linear-gradient`。
17. 你应该使用属性选择器来针对具有`three`作为其类值一部分的`span`元素的三个后代`span`元素，并设置它们的`background-image`属性以使用`linear-gradient`。

# --hints--

你应该有一个文本为`Book Inventory`的`h1`元素。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

你应该只有一个 `h1` 元素。

```js
assert.equal(document.querySelectorAll('h1')?.length, 1);
```

你应该有一个 `table` 元素。

```js
assert.equal(document.querySelectorAll('table')?.length, 1);
```

你的 `table` 元素应该有五列。

```js
assert.equal(document.querySelectorAll('th')?.length, 5);
```

你的第一列应该有文本 `Title` 作为标题。

```js
assert.equal(document.querySelectorAll('th')[0]?.innerText.trim(), 'Title');
```

你的第二列应该有文本 `Author` 作为标题。

```js
assert.equal(document.querySelectorAll('th')[1]?.innerText.trim(), 'Author');
```

你的第三列应该有文本 `Category` 作为标题。

```js
assert.equal(document.querySelectorAll('th')[2]?.innerText.trim(), 'Category');
```

你的第四列应该有文本 `Status` 作为标题。

```js
assert.equal(document.querySelectorAll('th')[3]?.innerText.trim(), 'Status');
```

你的第五列应该有文本 `Rate` 作为标题。

```js
assert.equal(document.querySelectorAll('th')[4]?.innerText.trim(), 'Rate');
```

你的表格应该至少有四行。

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

除了标题行之外的每个表格行都应该具有`read`、`to-read`或`in-progress`类。

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

`Status`列的`td`元素应该包含一个`span`元素。

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isAbove(statusColumnData.length, 0);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

`Status`列的每个`span`元素都应该具有`status`类。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isAbove(statusSpans.length, 0);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

每个`.status`元素都应该有文本`Read`、`To Read`或`In Progress`，具体取决于其行的类。

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isAbove(statusSpans.length, 0);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

`Rate`列的`td`元素应该包含一个`span`元素。

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isAbove(rateColumnData.length, 0);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

`Rate`列的`td`元素的每个直接子`span`元素都应该将`rate`作为第一个类。

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

每个`.rate`元素都应该包含三个空的`span`元素。

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

放置在`.read`行中的`.rate`元素应该在`rate`类之后有一个附加类，值为`one`、`two`或`three`中的一个，具体取决于个人评分。

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isAbove(readBooksRates.length, 0);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

你应该有一个属性选择器来针对具有`read`类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

你应该使用属性选择器来针对具有`read`类的行，并设置它们的`background-image`属性为你选择的线性渐变。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

你应该有一个属性选择器来针对具有`to-read`类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

你应该使用属性选择器来针对具有`to-read`类的行，并设置它们的`background-image`属性为你选择的线性渐变。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

你应该有一个属性选择器来针对具有`in-progress`类的行。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

你应该使用属性选择器来针对具有`in-progress`类的行，并设置它们的`background-image`属性为你选择的线性渐变。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

你应该将每个`span`元素的`display`属性设置为`inline-block`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

你应该有一个属性选择器来针对具有`to-read`类的`tr`元素的后代中具有`status`类的`span`元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

你应该使用属性选择器来针对具有`to-read`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`border`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

你应该使用属性选择器来针对具有`to-read`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`background-image`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

你应该有一个属性选择器来针对具有`read`类的`tr`元素的后代中具有`status`类的`span`元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

你应该使用属性选择器来针对具有`read`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`border`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

你应该使用属性选择器来针对具有`read`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`background-image`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

你应该有一个属性选择器来针对具有`in-progress`类的`tr`元素的后代中具有`status`类的`span`元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

你应该使用属性选择器来针对具有`in-progress`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`border`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

你应该使用属性选择器来针对具有`in-progress`类的`tr`元素的后代中具有`status`类的`span`元素，并设置它们的`background-image`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

你应该有一个属性选择器来针对具有`status`类的`span`元素和具有以`rate`开头的类值的`span`元素。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

你应该使用属性选择器来针对具有`status`类的`span`元素和具有以`rate`开头的类值的`span`元素，并设置它们的`height`属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

你应该使用属性选择器来针对具有`status`类的`span`元素和具有以`rate`开头的类值的`span`元素，并设置它们的`width`属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

你应该使用属性选择器来针对具有`status`类的`span`元素和具有以`rate`开头的类值的`span`元素，并设置它们的`padding`属性。

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

你应该有一个属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

你应该使用属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素，并设置它们的`border`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

你应该使用属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素，并设置它们的`border-radius`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

你应该使用属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素，并设置它们的`margin`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

你应该使用属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素，并设置它们的`height`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

你应该使用属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素，并设置它们的`width`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

你应该使用属性选择器来针对具有以`rate`开头的类的`span`元素的直接子元素`span`元素，并设置它们的`background-color`属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

你应该有一个属性选择器来针对具有`one`作为其类值一部分的`span`元素的第一个后代。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应该使用属性选择器来针对具有`one`作为其类值一部分的`span`元素的第一个后代，并设置其`background-image`属性以使用`linear-gradient`。

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

你应该有一个属性选择器来针对具有`two`作为其类值一部分的`span`元素的前两个后代。

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应该使用属性选择器来针对具有`two`作为其类值一部分的`span`元素的前两个后代，并设置它们的`background-image`属性以使用`linear-gradient`。

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

你应该有一个属性选择器来针对具有`three`作为其类值一部分的`span`元素的后代`span`元素。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

你应该使用属性选择器来针对具有`three`作为其类值一部分的`span`元素的后代`span`元素，并设置它们的`background-image`属性以使用`linear-gradient`。

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```