---
id: 671b6e873249bb35c9debfcf
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

在锚元素之前添加文字 `See more`，在锚元素之后添加文字 `in our gallery`。

你做得很好！通过添加这些文本，你让页面内容更加完整和自然了。

# --hints--

你应该使用现有的锚元素，而不是创建一个新的。

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

锚元素内的文字不应更改，必须为 `cat photos`。

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

在你的代码中应该有 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`。

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```