﻿---
id: 6823ac607bfdbc46331b2559
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML 是超文本标记语言的缩写。它是定义网页结构和内容的代码。这是你用来编写 HTML 代码的代码编辑器。
请在你的编辑器中输入以下文本。
```md
Welcome to MasterPuti
```

输入完成后，点击"运行"按钮，查看你的代码是否正确。
# --hints--

你应该在编辑器中包含文本 `Welcome to MasterPuti` ，并检查文本的内容。
```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*MasterPuti/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```