﻿---
id: 6823d9ac8bdc3853df65a1ff
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

HTML中有六个标题元素，从h1到h6，分别表示不同的标题等级。
在你的 `p` 元素下面，添加一个 `h3` 标题元素，文本内容为：
```md
HTML介绍
```

# --hints--

你的 `h3` 元素应该有一个开始标签 `<h3>`。
```js
assert.exists(document.querySelector("h3"));
```

你的 `h3` 元素应该有一个结束标签 `</h3>`。
```js
assert.match(code, /<\/h3\s*\>/);
```

你的 `h3` 元素的文本内容应该为 `HTML介绍`。
```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to MasterPuti</h1>
<h2>全栈课程学习</h2>
<p>学习全栈课程，成为一个全栈开发人员</p>

--fcc-editable-region--

--fcc-editable-region--
```