---
id: 66a83fec026a7a4631e084d2
title: æ­¥éª¤ 8
challengeType: 0
dashedName: step-8
---

# --description--

åœ¨ä½¿ç”?`fieldset` å…ƒç´ æ—¶ï¼Œé€šå¸¸ä½¿ç”¨ä¸€ä¸ªæ ‡é¢˜æ¥æè¿°è¾“å…¥ç»„ã€?ä¸ºæ­¤ï¼Œä½ å¯ä»¥ä½¿ç”¨ `legend` å…ƒç´ ã€?

ä¸‹é¢æ˜¯ä½¿ç”?`legend` å…ƒç´ çš„ç¤ºä¾‹ï¼š

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

åœ¨ä½ çš?`fieldset` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`legend` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Personal Information`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`legend` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<legend>/);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`legend` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/legend>/);
```

ä½ çš„ `legend` å…ƒç´ åº”è¯¥åœ?`fieldset` å…ƒç´ ä¸­ã€?

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

ä½ çš„ `legend` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Personal Information`ã€?

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```

