---
id: 66a84111965a0c46df6bbd0a
title: æ­¥éª¤ 9
challengeType: 0
dashedName: step-9
---

# --description--

In previous lessons, you learned how to associate a `label` element with an `input` like this:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

`label` å…ƒç´ çš?`for` å±æ€§åº”ä¸å¯¹åº”çš„ `input` å…ƒç´ çš?`id` å±æ€§åŒ¹é…ã€?è¿™ç§°ä¸ºæ˜¾å¼å…³è”ã€?

åœ¨ä½ çš?`legend` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`label`å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Name (required):`ã€?å°†å…¶ `for` å±æ€§è®¾ç½®ä¸º `"full-name"`ã€?

ç„¶ååœ¨ä½ çš?`label` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸ªæ— å±æ€§çš„ `input` å…ƒç´ ã€?åœ¨æ¥ä¸‹æ¥çš„æ­¥éª¤ä¸­ï¼Œä½ å°†æ·»åŠ å¿…è¦çš„å±æ€§ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`label` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<label.*>/);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`label` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/label>/);
```

ä½ çš„ `label` å…ƒç´ åº”æ”¾åœ?`fieldset` å…ƒç´ ä¸?`legend` çš„ä¸‹æ–¹ã€?

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

ä½ çš„ `label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `"full-name"`ã€?

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

ä½ çš„ `label` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Name (required):`ã€?

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

ä½ åº”è¯¥åœ¨ `label` å…ƒç´ ä¸‹æ–¹æœ‰ä¸€ä¸?`input` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```

