---
id: 66a83e5e491625454b6f62c3
title: æ­¥éª¤ 7
challengeType: 0
dashedName: step-7
---

# --description--

è¡¨å•ç”?`inputs`ç»„æˆï¼Œç”¨æˆ·å¯åœ¨å…¶ä¸­è¾“å…¥æ•°æ®ã€?ä½ å¯ä»¥ä½¿ç”?`fieldset` å…ƒç´ å°†ç›¸å…³è¾“å…¥åˆ†ç»„ã€?

ä¸‹é¢æ˜¯ä½¿ç”?`fieldset` å…ƒç´ çš„ç¤ºä¾‹ï¼š

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

åœ¨ä½ çš?`form` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`fieldset` å…ƒç´ ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`fieldset` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<fieldset>/i);
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`fieldset` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/fieldset>/i);
```

ä½ çš„ `fieldset` å…ƒç´ åº”è¯¥åœ?`form` å…ƒç´ ä¸­ã€?

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">

      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```

