---
id: 66a83bdcf425e7446900b7c4
title: æ­¥éª¤ 6
challengeType: 0
dashedName: step-6
---

# --description--

In previous lessons, you learned how to work with the `form` element like this:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

`action` å±æ€§ç”¨äºæŒ‡å®šæäº¤è¡¨å•æ—¶è¡¨å•æ•°æ®çš„å‘é€ä½ç½®ã€?

`method` å±æ€§ç”¨äºæŒ‡å®šå‘é€è¡¨å•æ•°æ®æ—¶ä½¿ç”¨çš?HTTP æ–¹æ³•ã€?æœ€å¸¸è§çš„æ–¹æ³•æ˜¯ `GET` å’?`POST`ã€?

**æ³¨æ„**ï¼šä½ å°†åœ¨åé¢çš„æ¨¡å—ä¸­äº†è§£ HTTP æ–¹æ³•çš„å·¥ä½œåŸç†ã€?

åœ¨ä½ çš?`main` å…ƒç´ å†…ï¼Œæ·»åŠ ä¸€ä¸?`form` å…ƒç´ ï¼Œå°† `action` å±æ€§è®¾ç½®ä¸º `"https://hotel-feedback.MasterPuti.org"`ï¼Œ`method` å±æ€§è®¾ç½®ä¸º `"POST"`ã€?

# --hints--

ä½ åº”æœ‰ä¸€ä¸?`form` å¼€å§‹æ ‡ç­¾ã€?

```js
assert.match(code, /<form.*>/i);
```

ä½ åº”æœ‰ä¸€ä¸?`form` ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /<\/form>/i);
```

ä½ çš„ `form` å…ƒç´ åº”æœ‰ä¸€ä¸?`action` å±æ€§ï¼Œè®¾ç½®ä¸?`"https://hotel-feedback.MasterPuti.org"`ã€?

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.MasterPuti.org');
```

ä½ çš„ `form` å…ƒç´ åº”æœ‰ä¸€ä¸?`method` å±æ€§ï¼Œè®¾ç½®ä¸?`"POST"`ã€?

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```

