---
id: 66a9521bc70162712caf118d
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

如果想让用户从一个选项列表中选择一个选项，可以使用一组单选按钮。

下面是两个单选按钮的示例：

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

在此示例中，单选按钮通过使用相同的 `name` 属性值分组在一起。这意味着一次只能选择一个单选按钮。

在你的 `legend` 元素下方，添加一个 `radio` 按钮，`id` 属性设置为 `"yes-option"`，`name` 属性设置为 `"hotel-stay"`。

在你的 `radio` 按钮下方，添加一个 `label` 元素，文本内容为 `Yes`，`for` 属性设置为 `"yes-option"`。

# --hints--

你的 `legend` 元素下方应该有一个 `radio` 按钮。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

你的 `radio` 按钮应该将 `id` 属性设置为 `"yes-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

你的 `radio` 按钮应该将 `name` 属性设置为 `"hotel-stay"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

你应该在你的 `radio` 按钮下方有一个 `label` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

你的 `label` 元素应该包含文本 `Yes`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.textContent.trim(), 'Yes');
```

你的 `label` 元素应该将 `for` 属性设置为 `"yes-option"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.getAttribute('for'), 'yes-option');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```