---
id: 66a969951120be7818d8ee49
title: æ­¥éª¤ 27
challengeType: 0
dashedName: step-27
---

# --description--

è¡¨å•çš„ä¸‹ä¸€éƒ¨åˆ†å°†ä¸ºç”¨æˆ·æä¾›ä¸ºé…’åº—è¯„åˆ†çš„åŠŸèƒ½ã€?

é¦–å…ˆæ·»åŠ ä¸€ä¸ªæ–°çš?`fieldset` å…ƒç´ ï¼Œå¹¶åœ¨å…¶ä¸­åµŒå¥—ä¸€ä¸?`legend` å…ƒç´ ã€?`legend` çš„æ–‡æœ¬åº”ä¸?`Ratings`ã€?

åœ?`legend` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`label` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`How was the service?`ã€?`for` å±æ€§åº”è®¾ç½®ä¸?`"service"`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`fieldset` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 4);
```

ä½ çš„ `fieldset` åº”è¯¥æœ‰ä¸€ä¸?`legend` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend'), 1);
```

ä½ çš„ legend åº”åŒ…å«æ–‡æœ?`Ratings` ã€?

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend')[0]?.textContent.trim(), 'Ratings');
```

ä½ çš„ `fieldset` åº”è¯¥åœ?`legend` ä¸‹æ–¹æœ‰ä¸€ä¸?`label` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend + label'), 1);
```

ä½ çš„ `label` åº”åŒ…å«æ–‡æœ?`How was the service?`ã€?

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.textContent.trim(), 'How was the service?');
```

ä½ çš„ `label` åº”å°† `for` å±æ€§è®¾ç½®ä¸º `"service"`ã€?

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.getAttribute('for'), 'service');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="price" value="price" />
          <label for="price">Price</label>
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

