---
id: 66a9577022877d72d8f43b4f
title: 步骤 21
challengeType: 0
dashedName: step-21
---

# --description--

表格的下一部分将询问用户为什么选择入住这家酒店。用户将有机会选择多个选项。

首先添加另一个 `fieldset` 元素。

在 `fieldset` 元素内部，添加一个 `legend` 元素，其文本内容为 `Why did you choose to stay at our hotel? (Check all that apply)`。

# --hints--

你应该有一个 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

你应该在 `fieldset` 元素中包含一个 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

你的 `legend` 元素的文本内容应该是 `Why did you choose to stay at our hotel? (Check all that apply)`。仔细检查拼写错误和空格问题。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText.trim(), 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
          <input type="radio" id="no-option" name="hotel-stay">
          <label for="no-option">No</label>
        </fieldset>
        <fieldset>
          <legend>Why did you choose to stay at our hotel? (Check all that apply)</legend>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```