---
id: 66a930b20f589b6664c51cb0
title: æ­¥éª¤ 10
challengeType: 0
dashedName: step-10
---

# --description--

å½“ç”¨æˆ·æä¾›å…¨åæ—¶ï¼Œ`input` å°†æ¥å—çº¯æ–‡æœ¬ã€?

In previous lessons, you learned how to work with the `type` attribute like this:

```html
<input type="text">
```

å¯¹äºä½ ç°æœ‰çš„ `input` å…ƒç´ ï¼Œæ·»åŠ ä¸€ä¸?`type` å±æ€§ï¼Œè®¾ç½®ä¸?`"text"`ã€?

ä¹Ÿä¸º `input` å…ƒç´ èµ‹äºˆä¸€ä¸?`id` å±æ€§ï¼Œå…¶å€¼ä¸º `"full-name"`ã€?

# --hints--

ä½ ä¸åº”åˆ é™¤ç°æœ‰çš„ `input` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

ä½ ä¸åº”åˆ›å»ºæ–°çš?`input` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`type` å±æ€§ï¼Œå€¼ä¸º `"text"`ã€?

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œå€¼ä¸º `"full-name"`ã€?

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

