---
id: 66a93c95bc58e26a8fe95818
title: æ­¥éª¤ 14
challengeType: 0
dashedName: step-14
---

# --description--

æ¥ä¸‹æ¥ï¼Œåœ¨ä½ çš„é‚®ç®?`label` ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`input`ï¼Œç±»å‹ä¸º `"email"`ã€?è¯?`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œå€¼è®¾ç½®ä¸º `"email"`ã€?å°?`name` å±æ€§ä¹Ÿè®¾ç½®ä¸?`"email"`ã€?

è¯?`input`ä¹Ÿæ˜¯å¿…å¡«çš„ï¼Œå› æ­¤ï¼Œè¯·ç¡®ä¿æ·»åŠ  `required` å±æ€§ã€?

æœ€åï¼Œæ·»åŠ ä¸€ä¸?`placeholder` å±æ€§ï¼Œè®¾ç½®ä¸?`"example@email.com"`ã€?

# --hints--

ä½ çš„é‚®ç®± `label` å…ƒç´ ä¸‹é¢åº”è¯¥æœ‰ä¸€ä¸?`input` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`type`ï¼Œå€¼ä¸º `"email"`ã€?

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

ä½ çš„ `input` å…ƒç´ åº”æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œè®¾ç½®ä¸?`"email"`ã€?

```js
assert.isNotNull(document.querySelector('input#email'));
```

ä½ çš„ `input` å…ƒç´ åº”æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œè®¾ç½®ä¸?`"email"`ã€?

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥å…·æœ‰ `required` å±æ€§ã€?

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

æ‚¨çš„ `è¾“å…¥çš„` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`å ä½ç¬¦` å±æ€§è®¾ç½®ä¸º `"example@email.com"`

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
        --fcc-editable-region--
          <label for="email">Email address (required):</label>

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

