---
id: 66a9689b1bf24b7750898a1b
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

对于 `fieldset` 中最后的 `input` 和 `label` 元素，你将添加一个复选框 `input` 元素，将它的 `id`、`name` 和 `value` 属性都设置为 `"price"`。

接下来，添加一个 `label` 元素，文本内容为 `Price`。`for` 属性设置为 `"price"`。

现在，你可以通过选择不同的复选框来测试你的 `form` 表单了！

# --hints--

你应该有第五个复选框 `input` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"]:nth-of-type(5)'));
```

你的第五个复选框 `input` 元素应该有一个 `id` 属性，设置为 `"price"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('id'), 'price');
```

你的第五个复选框 `input` 元素应该有一个 `name` 属性，设置为 `"price"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('name'), 'price');
```

你的第五个复选框 `input` 元素应该有一个 `value` 属性，设置为 `"price"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('value'), 'price');
```

你应该在你的复选框 `input` 元素之后有一个 `label` 元素，文本内容为 `Price`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.textContent?.trim(), 'Price');
```

你的 `label` 元素应该将 `for` 属性设置为 `"price"`。

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.getAttribute('for'), 'price');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
          <input type="radio" id="no-option" name="hotel-stay">
          <label for="no-option">No</label>
        </fieldset>
        <fieldset>
          <legend>Why did you choose to stay at our hotel? (Check all that apply)</legend>
          <input type="checkbox" id="ads" name="ads" value="ads">
          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="recommendation" name="recommendation" value="recommendation">
          <label for="recommendation">Personal Recommendation</label>
          <input type="checkbox" id="location" name="location" value="location">
          <label for="location">Location</label>
          <input type="checkbox" id="reputation" name="reputation" value="reputation">
          <label for="reputation">Reputation</label>
          <input type="checkbox" id="price" name="price" value="price">
          <label for="price">Price</label>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

