---
id: 66a972137acd1179fa3fe8a0
title: æ­¥éª¤ 29
challengeType: 0
dashedName: step-29
---

# --description--

åœ¨ä½ çš?`select` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä»¥ä¸‹äº”ä¸ª `option` å…ƒç´ ï¼Œå¹¶ä¸?`option` æ–‡æœ¬å’?`value` å±æ€§æ·»åŠ ç›¸åº”çš„å€¼ï¼š

**å€¼å±æ€§ï¼š**

- poor
- satisfactory
- good
- very-good
- excellent

**é€‰é¡¹å…ƒç´ æ–‡æœ¬ï¼?*

- Poor
- Satisfactory
- Good
- Very Good
- Excellent

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸ªå€¼è®¾ç½®ä¸º `"poor"` çš?`option` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('option[value="poor"]'));
```

`value` ä¸?`"poor"` çš?`option` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Poor`ã€?

```js
assert.strictEqual(document.querySelector('option[value="poor"]')?.textContent.trim(), 'Poor');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`value` è®¾ç½®ä¸?`"satisfactory"` çš?`option` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('option[value="satisfactory"]'));
```

`value` ä¸?`"satisfactory"` çš?`option` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Satisfactory`ã€?

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]')?.textContent.trim(), 'Satisfactory');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`value` è®¾ç½®ä¸?`"good"` çš?`option` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('option[value="good"]'));
```

`value` ä¸?`"good"` çš?`option` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Good`ã€?

```js

assert.strictEqual(document.querySelector('option[value="good"]')?.textContent.trim(), 'Good');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸ªå€¼è®¾ç½®ä¸º `"very-good"` çš?`option` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('option[value="very-good"]'));
```

`value` ä¸?`"very-good"` çš?`option` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Very Good`ã€?

```js
assert.strictEqual(document.querySelector('option[value="very-good"]')?.textContent.trim(), 'Very Good');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸ªå€¼è®¾ç½®ä¸º `"excellent"` çš?`option` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('option[value="excellent"]'));
```

`value` ä¸?`"excellent"` çš?`option` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Excellent`ã€?

```js
assert.strictEqual(document.querySelector('option[value="excellent"]')?.textContent.trim(), 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="price" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          --fcc-editable-region--
          <select name="service" id="service">

          </select>
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

