---
id: 66a962954f4e0d76223b37ed
title: 步骤 24
challengeType: 0
dashedName: step-24
---

# --description--

接下来，添加另一个复选框 `input` 元素，其 `id`、`name` 和 `value` 属性都设置为 `"location"`。

`label` 元素的文本内容为 `Location`。`for` 属性应该设置为 `"location"`。

在该 `label` 元素下方，添加另一个复选框 `input` 元素，其 `id`、`name` 和 `value` 属性都设置为 `"reputation"`。

`label` 元素的文本内容为 `Reputation`。`for` 属性应该设置为 `"reputation"`。

# --hints--

你应该有第三个复选框 `input` 元素。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

你的第三个复选框 `input` 元素应该有一个 `id` 属性，设置为 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

你的第三个复选框 `input` 元素应该有一个 `name` 属性，设置为 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'location');
```

你的第三个复选框 `input` 元素应该有一个 `value` 属性，设置为 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

你应该在你的复选框 `input` 元素之后有一个 `label` 元素，文本内容为 `Location`。

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

你的 `label` 元素应该将 `for` 属性设置为 `"location"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.getAttribute('for'), 'location');
```

你应该有第四个复选框 `input` 元素。

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(4)"));
```

你的第四个复选框 `input` 元素应该有一个 `id` 属性，设置为 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(4)")?.getAttribute('id'), 'reputation');
```

你的第四个复选框 `input` 元素应该有一个 `name` 属性，设置为 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(4)")?.getAttribute('name'), 'reputation');
```

你的第四个复选框 `input` 元素应该有一个 `value` 属性，设置为 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(4)")?.getAttribute('value'), 'reputation');
```

你应该在你的复选框 `input` 元素之后有一个 `label` 元素，文本内容为 `Reputation`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='reputation'] + label")?.textContent?.trim(), 'Reputation');
```

你的 `label` 元素应该将 `for` 属性设置为 `"reputation"`。

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='reputation'] + label")?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
          <input type="radio" id="no-option" name="hotel-stay">
          <label for="no-option">No</label>
        </fieldset>
        <fieldset>
          <legend>Why did you choose to stay at our hotel? (Check all that apply)</legend>
          <input type="checkbox" id="ads" name="ads" value="ads">
          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="recommendation" name="recommendation" value="recommendation">
          <label for="recommendation">Personal Recommendation</label>
          <input type="checkbox" id="location" name="location" value="location">
          <label for="location">Location</label>
          <input type="checkbox" id="reputation" name="reputation" value="reputation">
          <label for="reputation">Reputation</label>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```