---
id: 66a9689b1bf24b7750898a1b
title: æ­¥éª¤ 26
challengeType: 0
dashedName: step-26
---

# --description--

å¯¹äº `fieldset` ä¸­æœ€åçš„ `input` å’?`label`ï¼Œä½ å°†æ·»åŠ ä¸€ä¸ªå¤é€‰æ¡† `input` å…ƒç´ ï¼Œå°†å…?`id`ã€`name` å’?`value`  å±æ€§è®¾ç½®ä¸º `"price"`ã€?

æ¥ä¸‹æ¥ï¼Œæ·»åŠ ä¸€ä¸?`label` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Price`ï¼?`for` å±æ€§è®¾ç½®ä¸º `"price"`ã€?

ç°åœ¨ï¼Œä½ å¯ä»¥é€šè¿‡é€‰æ‹©ä¸åŒçš„å¤é€‰æ¡†æ¥æµ‹è¯•ä½ çš?`form`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ç¬¬äº”ä¸ªå¤é€‰æ¡† `input`ã€?

```js
assert.exists(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"]:nth-of-type(5)'));
```

ä½ çš„ç¬¬äº”ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œè®¾ç½®ä¸?`"price"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('id'), 'price');
```

ä½ çš„ç¬¬äº”ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œè®¾ç½®ä¸?`"price"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('name'), 'price');
```

ä½ çš„ç¬¬äº”ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`value` å±æ€§ï¼Œè®¾ç½®ä¸?`"price"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(5)")?.getAttribute('value'), 'price');
```

ä½ åº”è¯¥åœ¨ä½ çš„å¤é€‰æ¡† `input` ä¹‹åæœ‰ä¸€ä¸?`label` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Price`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.textContent?.trim(), 'Price');
```

ä½ çš„ `label` åº”å°† `for` å±æ€§è®¾ç½®ä¸º `"price"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="price"] + label')?.getAttribute('for'), 'price');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

