---
id: 66a93730719e1f68410cce54
title: æ­¥éª¤ 11
challengeType: 0
dashedName: step-11
---

# --description--

`name` å±æ€§ç”¨äºåœ¨è¡¨å•æ•°æ®æäº¤åˆ°æœåŠ¡å™¨åå¯¹å…¶è¿›è¡Œæ ‡è¯†ã€?

ä¸‹é¢æ˜¯ä¸€ä¸ªå¦‚ä½•ä½¿ç”?`name` å±æ€§çš„ç¤ºä¾‹ï¼?

```html
<input type="email" name="email">
```

ä¸?`input` å…ƒç´ æ·»åŠ ä¸€ä¸?`name` å±æ€§ï¼Œå€¼ä¸º `"name"`ã€?

# --hints--

ä½ ä¸åº”åˆ é™¤ç°æœ‰çš„ `input` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

ä½ ä¸åº”åˆ›å»ºæ–°çš?`input` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œå€¼ä¸º `"name"`ã€?

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

