---
id: 66a962954f4e0d76223b37ed
title: æ­¥éª¤ 24
challengeType: 0
dashedName: step-24
---

# --description--

æ¥ä¸‹æ¥ï¼Œæ·»åŠ å¦ä¸€ä¸ªå¤é€‰æ¡† `input`ï¼Œå…¶ `id`ã€`name` å’?`value` å±æ€§éƒ½è®¾ç½®ä¸?`"location"`ã€?

æ­?`label` å…ƒç´ çš„æ–‡æœ¬ä¸º `Location`ï¼?`for` å±æ€§åº”è®¾ç½®ä¸?`"location"`ã€?

åœ¨è¯¥ `label` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ å¦ä¸€ä¸ªå¤é€‰æ¡† `input`ï¼Œå…¶ `id`ã€`name` å’?`value` å±æ€§éƒ½è®¾ç½®ä¸?`"reputation"`ã€?

æ­?`label` å…ƒç´ çš„æ–‡æœ¬ä¸º `Reputation`ï¼?`for` å±æ€§åº”è®¾ç½®ä¸?`"reputation"`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ç¬¬ä¸‰ä¸ªå¤é€‰æ¡† `input`ã€?

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

ä½ çš„ç¬¬ä¸‰ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œè®¾ç½®ä¸?`"location"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

ä½ çš„ç¬¬ä¸‰ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œè®¾ç½®ä¸?`"location"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'location');
```

ä½ çš„ç¬¬ä¸‰ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`value` å±æ€§ï¼Œè®¾ç½®ä¸?`"location"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

ä½ åº”è¯¥åœ¨ä½ çš„å¤é€‰æ¡† `input` ä¹‹åæœ‰ä¸€ä¸?`label` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Location`ã€?

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

ä½ çš„ `label` åº”å°† `for` å±æ€§è®¾ç½®ä¸º `"location"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

ä½ çš„ `label` ä¸‹æ–¹åº”è¯¥æœ‰ç¬¬å››ä¸ªå¤é€‰æ¡† `input`ã€?

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

ä½ çš„ç¬¬å››ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œè®¾ç½®ä¸?`"reputation"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

ä½ çš„ç¬¬å››ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œè®¾ç½®ä¸?`"reputation"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'reputation');
```

ä½ çš„ç¬¬å››ä¸ªå¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`value` å±æ€§ï¼Œè®¾ç½®ä¸?`"reputation"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

åœ¨å£°èª‰å¤é€‰æ¡†ä¸‹æ–¹ï¼Œä½ åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Reputation`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

ä½ çš„ `label` åº”å°† `for` å±æ€§è®¾ç½®ä¸º `"reputation"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

