---
id: 66a94ea5df66236ebed260e8
title: 步骤18
challengeType: 0
dashedName: step-18
---

# --description--

表单的下一部分将负责询问用户是否曾经入住过这家酒店。

首先添加一个 `fieldset` 元素。

在 `fieldset` 元素内部，添加一个 `legend` 元素，其文本内容为 `Was this your first time at our hotel?`。

# --hints--

你应该有一个 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

你应该在 `fieldset` 元素中包含一个 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

你的 `legend` 元素的文本内容应该是 `Was this your first time at our hotel?`。

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```