---
id: 66a95d0eff8168747805f1f3
title: æ­¥éª¤ 22
challengeType: 0
dashedName: step-22
---

# --description--

å¦‚æœå¸Œæœ›ç”¨æˆ·ä»åˆ—è¡¨ä¸­é€‰æ‹©å¤šä¸ªé€‰é¡¹ï¼Œå¯ä»¥ä½¿ç”¨å¤é€‰æ¡†ã€?

ä¸‹é¢æ˜¯ä¸€ä¸ªå¦‚ä½•ä½¿ç”¨å¤é€‰æ¡†å¤„ç†é£Ÿç‰©é€‰é¡¹çš„ç¤ºä¾‹ï¼š

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

`value` å±æ€§ç”¨äºæŒ‡å®šæäº¤è¡¨å•æ—¶å‘é€åˆ°æœåŠ¡å™¨çš„å€¼ã€?

åœ¨ä½ çš?`legend` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸ªå¤é€‰æ¡† `input`ï¼Œå…¶ `id`ã€`name` å’?`value` å±æ€§éƒ½è®¾ç½®ä¸?`"ads"`ã€?

åœ¨ä½ çš„å¤é€‰æ¡† `input` ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`label` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Social Media Ads`ã€?`for` å±æ€§åº”è®¾ç½®ä¸?`"ads"`ã€?

# --hints--

ä½ çš„ `input` å…ƒç´ ä¸‹æ–¹åº”è¯¥æœ‰ä¸€ä¸ªå¤é€‰æ¡† `input`ã€?

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

ä½ çš„å¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œè®¾ç½®ä¸?`"ads"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

ä½ çš„å¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œè®¾ç½®ä¸?`"ads"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "ads");
```

ä½ çš„å¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`value` å±æ€§ï¼Œè®¾ç½®ä¸?`"ads"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

ä½ çš„å¤é€‰æ¡†  `input`  ä¸‹æ–¹åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ã€?

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

ä½ çš„ `label` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Social Media Ads`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

ä½ çš„ `label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œè®¾ç½®ä¸?`"ads"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--
        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

