---
id: 66a9521bc70162712caf118d
title: æ­¥éª¤ 19
challengeType: 0
dashedName: step-19
---

# --description--

å¦‚æœæƒ³è®©ç”¨æˆ·ä»ä¸€ä¸ªé€‰é¡¹åˆ—è¡¨ä¸­é€‰æ‹©ä¸€ä¸ªé€‰é¡¹ï¼Œå¯ä»¥ä½¿ç”¨ä¸€ç»„å•é€‰æŒ‰é’®ã€?

ä¸‹é¢æ˜¯ä¸¤ä¸ªå•é€‰æŒ‰é’®çš„ç¤ºä¾‹ï¼?

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

åœ¨æ­¤ç¤ºä¾‹ä¸­ï¼Œå•é€‰æŒ‰é’®é€šè¿‡ä½¿ç”¨ç›¸åŒçš?`name` å±æ€§å€¼åˆ†ç»„åœ¨ä¸€èµ·ã€?è¿™æ„å‘³ç€ä¸€æ¬¡åªèƒ½é€‰æ‹©ä¸€ä¸ªå•é€‰æŒ‰é’®ã€?

åœ¨ä½ çš?`legend` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`radio` æŒ‰é’®ï¼Œ`id` è®¾ç½®ä¸?`"yes-option"`ï¼Œ`name` å±æ€§è®¾ç½®ä¸º `"hotel-stay"`ã€?

åœ¨ä½ çš?`radio` æŒ‰é’®ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`label` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Yes`ï¼Œ`for` å±æ€§è®¾ç½®ä¸º `"yes-option"`ã€?

# --hints--

ä½ çš„ `legend` å…ƒç´ ä¸‹æ–¹åº”è¯¥æœ‰ä¸€ä¸?`radio` æŒ‰é’®ã€?

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

ä½ çš„ `radio` æŒ‰é’®åº”å°† `id` å±æ€§è®¾ç½®ä¸º `"yes-option"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

ä½ çš„ `radio` æŒ‰é’®åº”å°† `name` å±æ€§è®¾ç½®ä¸º `"hotel-stay"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

ä½ åº”è¯¥åœ¨ä½ çš„ `radio` æŒ‰é’®ä¸‹æ–¹æœ‰ä¸€ä¸?`label` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

ä½ çš„ `label` å…ƒç´ åº”å°† `for` å±æ€§è®¾ç½®ä¸º `"yes-option"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

ä½ çš„ `label` å…ƒç´ åº”è¯¥åŒ…å«æ–‡æœ¬ `Yes`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        --fcc-editable-region--
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

