---
id: 66a96b01f33ef178dfca9e42
title: æ­¥éª¤ 28
challengeType: 0
dashedName: step-28
---

# --description--

å½“ä½ å¸Œæœ›ç”¨æˆ·ä»ä¸‹æ‹‰èœå•ä¸­è¿›è¡Œé€‰æ‹©æ—¶ï¼Œå¯ä»¥ä½¿ç”¨ `select` å’?`option` å…ƒç´ ã€?

ä¸‹é¢æ˜¯ä¸€ä¸ªä½¿ç”?`select` å’?`option` å…ƒç´ åˆ›å»ºä¸åŒåŸå¸‚ä¸‹æ‹‰èœå•çš„ç¤ºä¾‹ï¼š

```html
<label for="city">Choose a City: </label>
<select id="city" name="city">
  <option value="new-york">New York</option>
  <option value="los-angeles">Los Angeles</option>
  <option value="chicago">Chicago</option>
  <option value="miami">Miami</option>
</select>
```

é¦–å…ˆæ·»åŠ ä¸€ä¸?`select` å…ƒç´ ï¼Œå…¶ `name` å’?`id` å±æ€§è®¾ç½®ä¸º `"service"`ã€?

# --hints--

ä½ çš„ `fieldset` åº”è¯¥æœ‰ä¸€ä¸?`select` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('fieldset:nth-of-type(4) select'));
```

ä½ çš„ `select` å…ƒç´ åº”å°† `name` å±æ€§è®¾ç½®ä¸º `"service"`ã€?

```js
assert.exists(document.querySelector('select[name="service"]'));
```

ä½ çš„ `select` å…ƒç´ åº”å°† `id` å±æ€§è®¾ç½®ä¸º `"service"`ã€?

```js
assert.exists(document.querySelector('select[id="service"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="price" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

