---
id: 66a9401c9d660d6bb15993e2
title: æ­¥éª¤ 16
challengeType: 0
dashedName: step-16
---

# --description--

ä½ çš„é…’åº—åé¦ˆè¡¨åº”ä¸ºç”¨æˆ·æ·»åŠ å¹´é¾„é€‰é¡¹ã€?

é¦–å…ˆï¼Œåœ¨è¡¨å•ä¸­æ·»åŠ ä¸€ä¸?`label` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Age (optional):`ã€?

`for` å±æ€§åº”è®¾ç½®ä¸?`"age"`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

ä½ çš„ `label` å…ƒç´ åº”å°† `for` å±æ€§è®¾ç½®ä¸º `"age"`ã€?

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

ä½ çš„ `label` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Age (optional):`ã€?

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

