---
id: 66a9824ac5d9f77ec304969f
title: 步骤 34
challengeType: 0
dashedName: step-34
---

# --description--

如果你希望用户有更多空间来撰写评论，可以使用 `textarea` 元素。

`textarea` 元素是一个多行文本输入控件，允许用户输入比单行更长的文本。它可用于创建评论框、消息输入或其他需要多行的文本输入。

下面是一个 `textarea` 元素的示例：

```html
<textarea id="comments" name="comments" rows="4" cols="50"></textarea>
```

`rows` 属性用于指定 `textarea` 的可见高度，`cols` 属性用于指定 `textarea` 的可见宽度。

在你的 `label` 元素下方，添加一个 `textarea` 元素。下一步，你将添加必要的属性。

# --hints--

你的 `label` 元素下方应该有一个 `textarea` 元素。

```js
assert.exists(document.querySelector('label + textarea'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
          <input type="radio" id="no-option" name="hotel-stay">
          <label for="no-option">No</label>
        </fieldset>
        <fieldset>
          <legend>Why did you choose to stay at our hotel? (Check all that apply)</legend>
          <input type="checkbox" id="ads" name="ads" value="ads">
          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="recommendation" name="recommendation" value="recommendation">
          <label for="recommendation">Personal Recommendation</label>
          <input type="checkbox" id="location" name="location" value="location">
          <label for="location">Location</label>
          <input type="checkbox" id="reputation" name="reputation" value="reputation">
          <label for="reputation">Reputation</label>
          <input type="checkbox" id="price" name="price" value="price">
          <label for="price">Price</label>
        </fieldset>
        <fieldset>
          <legend>Ratings</legend>
          <label for="service">How was the service?</label>
          <select id="service" name="service">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option value="excellent" selected>Excellent</option>
          </select>
          <label for="food">How was the food?</label>
          <select id="food" name="food">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option value="excellent" selected>Excellent</option>
          </select>
          <label for="comments">Other Comments?</label>
          <textarea></textarea>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

