---
id: 66a969951120be7818d8ee49
title: 步骤 27
challengeType: 0
dashedName: step-27
---

# --description--

表单的下一部分将为用户提供为酒店评分的功能。

首先添加一个新的 `fieldset` 元素，并在其中嵌套一个 `legend` 元素。`legend` 元素的文本内容应该是 `Ratings`。

在 `legend` 元素下方，添加一个 `label` 元素，其文本内容为 `How was the service?`。`for` 属性应该设置为 `"service"`。

# --hints--

你应该有一个 `fieldset` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 4);
```

你的 `fieldset` 元素应该有一个 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend'), 1);
```

你的 `legend` 元素应该包含文本 `Ratings`。

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend')[0]?.textContent.trim(), 'Ratings');
```

你的 `fieldset` 元素应该在 `legend` 元素下方有一个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend + label'), 1);
```

你的 `label` 元素应该包含文本 `How was the service?`。

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.textContent.trim(), 'How was the service?');
```

你的 `label` 元素应该将 `for` 属性设置为 `"service"`。

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.getAttribute('for'), 'service');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
          <input type="radio" id="no-option" name="hotel-stay">
          <label for="no-option">No</label>
        </fieldset>
        <fieldset>
          <legend>Why did you choose to stay at our hotel? (Check all that apply)</legend>
          <input type="checkbox" id="ads" name="ads" value="ads">
          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="recommendation" name="recommendation" value="recommendation">
          <label for="recommendation">Personal Recommendation</label>
          <input type="checkbox" id="location" name="location" value="location">
          <label for="location">Location</label>
          <input type="checkbox" id="reputation" name="reputation" value="reputation">
          <label for="reputation">Reputation</label>
          <input type="checkbox" id="price" name="price" value="price">
          <label for="price">Price</label>
        </fieldset>
        <fieldset>
          <legend>Ratings</legend>
          <label for="service">How was the service?</label>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

