---
id: 66a954b2bcddba72076c1857
title: æ­¥éª¤ 20
challengeType: 0
dashedName: step-20
---

# --description--

åœ¨ä½ çš?`label` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`radio` æŒ‰é’®ï¼Œ`id` è®¾ç½®ä¸?`"no-option"`ï¼Œåå­—å±æ€§è®¾ç½®ä¸º `"hotel-stay"`ã€?

åœ¨ä½ çš„æ–° `radio` æŒ‰é’®ä¸‹æ–¹ï¼Œæ·»åŠ å¦ä¸€ä¸?`label` å…ƒç´ ï¼Œå…¶ `for` å±æ€§è®¾ç½®ä¸º `"no-option"`ã€?`label` çš„æ–‡æœ¬åº”ä¸?`No`ã€?

å®Œæˆåï¼Œä½ å°±å¯ä»¥å°è¯•ä½¿ç”¨å•é€‰æŒ‰é’®ï¼Œæ¯æ¬¡é€‰æ‹©ä¸€ä¸ªé€‰é¡¹ã€?

# --hints--

ä½ åº”è¯¥åœ¨ç¬¬ä¸€ä¸?`label` å…ƒç´ ä¸‹æ–¹æœ‰ä¸€ä¸?`input` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`type` å±æ€§ï¼Œå€¼ä¸º `radio`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œå€¼ä¸º `"no-option"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

ä½ çš„ `input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œå€¼ä¸º `"hotel-stay"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

ä½ çš„ç¬¬äºŒä¸?`input` å…ƒç´ ä¸‹æ–¹åº”è¯¥æœ‰ç¬¬äºŒä¸ª `label` å…ƒç´ ã€?

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

ä½ çš„ `label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œè®¾ç½®ä¸?`"no-option"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

ä½ çš„ `label` å…ƒç´ åº”è¯¥åŒ…å«æ–‡æœ¬ `No`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

