---
id: 66a9577022877d72d8f43b4f
title: æ­¥éª¤ 21
challengeType: 0
dashedName: step-21
---

# --description--

è¡¨æ ¼çš„ä¸‹ä¸€éƒ¨åˆ†å°†è¯¢é—®ç”¨æˆ·ä¸ºä»€ä¹ˆé€‰æ‹©è¯¥é…’åº—ã€?ç”¨æˆ·å°†æœ‰æœºä¼šé€‰æ‹©å¤šä¸ªé€‰é¡¹ã€?

é¦–å…ˆæ·»åŠ å¦ä¸€ä¸?`fieldset` å…ƒç´ ã€?

åœ?`fieldset` å…ƒç´ å†…éƒ¨ï¼Œæ·»åŠ ä¸€ä¸?`legend` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Why did you choose to stay at our hotel? (Check all that apply)`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`fieldset` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

ä½ åº”è¯¥åœ¨ `fieldset` å…ƒç´ ä¸­åŒ…å«ä¸€ä¸?`legend` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

ä½ çš„ `legend` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Why did you choose to stay at our hotel? (Check all that apply)`ã€?ä»”ç»†æ£€æŸ¥æ‹¼å†™é”™è¯¯å’Œç©ºæ ¼é—®é¢˜ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText.trim(), 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

