---
id: 66a96127422411756204bc1b
title: æ­¥éª¤ 23
challengeType: 0
dashedName: step-23
---

# --description--

æ·»åŠ å¦ä¸€ä¸ªå¤é€‰æ¡† `input`ï¼Œå…¶ `id`ã€`name` å’?`value` å±æ€§éƒ½è®¾ç½®ä¸?`"recommendation"`ã€?

åœ¨è¯¥å¤é€‰æ¡† `input` ä¹‹åï¼Œæ·»åŠ å¦ä¸€ä¸?`label`ï¼Œæ–‡æœ¬ä¸º `Personal Recommendation`ã€?`for` å±æ€§åº”è®¾ç½®ä¸?`"recommendation"`ã€?

# --hints--

ä½ çš„ `label` å…ƒç´ ä¸‹æ–¹åº”è¯¥æœ‰ä¸€ä¸ªå¤é€‰æ¡† `input`ã€?

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(3) label + input[type="checkbox"]'));
```

ä½ çš„å¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`id` å±æ€§ï¼Œè®¾ç½®ä¸?`"recommendation"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.id, "recommendation");
```

ä½ çš„å¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œè®¾ç½®ä¸?`"recommendation"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.name, "recommendation");
```

ä½ çš„å¤é€‰æ¡† `input` åº”è¯¥æœ‰ä¸€ä¸?`value` å±æ€§ï¼Œè®¾ç½®ä¸?`"recommendation"`ã€?

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.value, "recommendation");
```

ä½ çš„å¤é€‰æ¡†  `input`  ä¸‹æ–¹åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ã€?

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label:nth-of-type(2)"));
```

ä½ çš„ `label` å…ƒç´ çš„æ–‡æœ¬åº”ä¸?`Personal Recommendation`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"] + label:nth-of-type(2)')?.textContent?.trim(), "Personal Recommendation");
```

ä½ çš„ `label` å…ƒç´ è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œè®¾ç½®ä¸?`"recommendation"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[type="checkbox"] + label:nth-of-type(2)')?.getAttribute("for"), "recommendation");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

