---
id: 66a975c259525b7bc2d5c776
title: æ­¥éª¤ 31
challengeType: 0
dashedName: step-31
---

# --description--

ä½ çš„é…’åº—åé¦ˆè¡¨è¿˜åº”å…è®¸ç”¨æˆ·å¯¹é£Ÿç‰©è¿›è¡Œè¯„åˆ†ã€?

é¦–å…ˆæ·»åŠ ä¸€ä¸?`label` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`How was the food?`ã€?è¯?`label` å…ƒç´ åº”å°† `for` å±æ€§è®¾ç½®ä¸º `"food"`ã€?

åœ¨ä½ çš?`label` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸?`select` å…ƒç´ ï¼Œ`id` å’?`name` è®¾ç½®ä¸?`"food"`ã€?

# --hints--

ä½ åº”å°?`label` å…ƒç´ çš?`for` å±æ€§è®¾ç½®ä¸º `"food"`ã€?

```js
assert.exists(document.querySelector('label[for="food"]'));
```

ä½ çš„ `label` åº”åŒ…å«æ–‡æœ?`How was the food?`ã€?

```js
assert.strictEqual(document.querySelector('label[for="food"]')?.textContent.trim(), 'How was the food?');
```

ä½ åº”å°?`label` ä¸‹æ–¹çš?`select` å…ƒç´ çš?`id` è®¾ç½®ä¸?`"food"`ã€?

```js
assert.exists(document.querySelector('label + select#food'));
```

ä½ çš„ `select` å…ƒç´ åº”å°† `name` å±æ€§è®¾ç½®ä¸º `"food"`ã€?

```js
assert.exists(document.querySelector('label + select[name="food"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="price" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option selected value="excellent">Excellent</option>
          </select>

          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

