---
id: 66a93bbe65a26169dbf3bc39
title: æ­¥éª¤ 13
challengeType: 0
dashedName: step-13
---

# --description--

é…’åº—åé¦ˆè¡¨è¿˜åº”æ”¶é›†ç”¨æˆ·çš„ç”µå­é‚®ä»¶åœ°å€ã€?

é¦–å…ˆï¼Œåœ¨è¡¨å•ä¸­æ·»åŠ ä¸€ä¸ªæ–°çš?`label` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Email address (required):`ã€?ä½ çš„ `label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼è®¾ç½®ä¸º `"email"`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`label` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

ä½ çš„ `label` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Email address (required):`ã€?

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

ä½ çš„ `label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼è®¾ç½®ä¸º `"email"`ã€?

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

