---
id: 66a9419e2d18476c645ce693
title: æ­¥éª¤ 17
challengeType: 0
dashedName: step-17
---

# --description--

æ•°å­— `input` ç”¨äºåˆ›å»ºæ•°å­—è¾“å…¥æ¡†ã€?

ä¸‹é¢æ˜¯ä¸€ä¸ªæ•°å­—è¾“å…¥æ¡†çš„ç¤ºä¾‹ï¼š

```html
<input type="number" id="age" name="age" min="18" max="100">
```

`min` å’?`max` å±æ€§ç”¨äºè®¾ç½®è¾“å…¥å­—æ®µä¸­å¯è¾“å…¥çš„æœ€å°å€¼å’Œæœ€å¤§å€¼ã€?

åœ¨ä½ çš?`label` å…ƒç´ ä¸‹æ–¹ï¼Œæ·»åŠ ä¸€ä¸ªè¾“å…¥ï¼Œ`type` å±æ€§è®¾ç½®ä¸º `"number"`ï¼Œ`id` ä¸?`"age"`ã€?

`name` å±æ€§åº”è®¾ç½®ä¸?`"age"`ï¼Œ`min` å±æ€§åº”è®¾ç½®ä¸?`"3"`ï¼Œ`max` å±æ€§åº”è®¾ç½®ä¸?`"100"`ã€?

# --hints--

ä½ çš„ `label` å…ƒç´ ä¸‹æ–¹åº”è¯¥æœ‰ä¸€ä¸ªæ•°å­?`input`ã€?

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

ä½ çš„æ•°å­— `input` åº”è¯¥æœ‰ä¸€ä¸?`id`ï¼Œå€¼ä¸º `age`ã€?

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

ä½ çš„æ•°å­— `input` åº”æœ‰ä¸€ä¸?`name` å±æ€§ï¼Œè®¾ç½®ä¸?`age`ã€?

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

ä½ çš„æ•°å­— `input` åº”æœ‰ä¸€ä¸?`min` å±æ€§ï¼Œè®¾ç½®ä¸?`"3"`ã€?

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

ä½ çš„æ•°å­— `input` åº”æœ‰ä¸€ä¸?`max` å±æ€§ï¼Œè®¾ç½®ä¸?`"100"`ã€?

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
      --fcc-editable-region--
          <label for="age">Age (optional):</label>

      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

