---
id: 66ad24c7eb8c121000c603a6
title: 步骤 25
challengeType: 0
dashedName: step-25
---

# --description--

要使复选框输入默认为选中状态，可以添加 `checked` 属性。

下面是使用 `checked` 属性的示例：

```html
<input checked type="checkbox" id="checked" name="checked">
```

给 `id` 为 `"reputation"` 的复选框输入添加 `checked` 属性，使其默认为选中状态。

# --hints--

你的 `id` 为 `"reputation"` 的复选框输入应该具有 `checked` 属性。

```js
assert.exists(document.querySelector("input#reputation[checked]"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required>
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input type="radio" id="yes-option" name="hotel-stay">
          <label for="yes-option">Yes</label>
          <input type="radio" id="no-option" name="hotel-stay">
          <label for="no-option">No</label>
        </fieldset>
        <fieldset>
          <legend>Why did you choose to stay at our hotel? (Check all that apply)</legend>
          <input type="checkbox" id="ads" name="ads" value="ads">
          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="recommendation" name="recommendation" value="recommendation">
          <label for="recommendation">Personal Recommendation</label>
          <input type="checkbox" id="location" name="location" value="location">
          <label for="location">Location</label>
          <input type="checkbox" id="reputation" name="reputation" value="reputation" checked>
          <label for="reputation">Reputation</label>
          <input type="checkbox" id="price" name="price" value="price">
          <label for="price">Price</label>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

