---
id: 67a51d3a8a6fe123b77b0c6e
title: æ­¥éª¤15
challengeType: 0
dashedName: step-15
---

# --description--

`input` å…ƒç´ å¯ä»¥æœ‰ä¸€ä¸?`size` å±æ€§ã€?è¯¥å±æ€§å®šä¹‰äº†ç”¨æˆ·åœ¨è¾“å…¥æ—¶åº”æ˜¾ç¤ºçš„å­—ç¬¦æ•°ã€?`size` çš„å€¼åº”ä¸ºå¤§äºé›¶çš„éè´Ÿæ•´æ•°ã€?å¦‚æœæœªæŒ‡å®?`size` æˆ–æŒ‡å®šçš„å€¼æ— æ•ˆï¼Œè¾“å…¥å†…å®¹å°†ä½¿ç”¨æµè§ˆå™¨è®¾ç½®çš„é»˜è®¤å®½åº¦ã€?

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

ä¸ºå§“åå’Œé‚®ç®±è¾“å…¥èµ‹äºˆ `size` å±æ€§ï¼Œå…¶å€¼ä¸º `"20"`ã€?

# --hints--

ä½ åº”å°†å§“å?`input` çš?`size` å±æ€§å€¼è®¾ç½®ä¸º `"20"`ã€?

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

ä½ åº”å°†é‚®ç®?`input` çš?`size` å±æ€§å€¼è®¾ç½®ä¸º `"20"`ã€?

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.MasterPuti.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

