---
id: 67a51d3a8a6fe123b77b0c6e
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

`input` 元素可以有一个 `size` 属性。该属性定义了用户在输入时应显示的字符数。`size` 的值应该为大于零的非负整数。如果未指定 `size` 或指定的值无效，输入内容将使用浏览器设置的默认宽度。

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

为姓名和邮箱输入元素添加 `size` 属性，其值为 `"20"`。

# --hints--

你应该将姓名 `input` 元素的 `size` 属性值设置为 `"20"`。

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

你应该将邮箱 `input` 元素的 `size` 属性值设置为 `"20"`。

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="https://hotel-feedback.MasterPuti.org" method="POST">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">
          <label for="email">Email address (required):</label>
          <input type="email" id="email" name="email" placeholder="example@email.com" required size="20">
          <label for="age">Age (optional):</label>
          <input type="number" id="age" name="age" min="3" max="100">
        </fieldset>
      </form>
    </main>
  </body>
</html>
```