---
id: 5f331e55dfab7a896e53c3a1
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8`. 这告诉浏览器如何为页面编码字符。

# --hints--

你应该有一个`meta`标签。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta`元素是一个空元素，它不应该有结束标签`</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的`meta`标签应该有一个`charset`属性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的`charset`属性应该有一个值`utf-8`。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
