---
id: 5f3313e74582ad9d063e3a38
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在`html`元素内添加一个`head`元素，这样你就可以添加一个`title`元素。`title`元素的文本应该是`Cafe Menu`。

# --hints--

你应该有一个开头的`<head>`标签。

```js
assert.match(code, /<head>/i);
```

你应该有一个结尾的`</head>`标签。

```js
assert.match(code, /<head>/i);
```

你应该有一个开头的`<title>`标签。

```js
assert.match(code, /<title>/i);
```

你应该有一个结尾的`</title>`标签。

```js
assert.match(code, /<\/title>/i);
```

你的`title`元素应该嵌套在你的`head`元素中。

```js
assert.match(code, /<head>\s*<title>.*<\/title>\s*<\/head>/si);
```

你的`title`元素应该有文本`Cafe Menu`。你可能需要检查你的拼写。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
