---
id: 5f3477ae8466a9a3d2cc953c
title: 步骤 15
challengeType: 0
dashedName: step-15
---

# --description--

现在你已经在`styles.css`文件中有CSS了，继续删除`style`元素及其所有内容。一旦删除，居中的文本将移回左侧。

# --hints--

你的代码中不应该有任何`style`标签。

```js
assert.notMatch(code, /style/i);
```

你的HTML文件中不应该有任何CSS选择器。

```js
assert.isFalse(editableContents.includes('style'));
assert.isFalse(editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

