---
id: 5f3477ae34c1239cafe128be
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

你现在有三个具有相同样式的类型选择器。你可以通过创建选择器列表将相同样式组添加到许多元素中。每个选择器用逗号分隔，就像这样：

```css
selector1, selector2 {
  property: value;
}
```

删除现有的三个类型选择器，用一个选择器列表替换它们，该列表可以居中`h1`、`h2`和`p`元素的文本。

# --hints--

你应该为所有三个元素使用单一的类型选择器，`h1, h2, p`。确保使用这个顺序。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

你的`style`元素中应该只有一个选择器。

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

