---
id: 5f356ed63c7807a4f1e6d054
title: 步骤 21
challengeType: 0
dashedName: step-21
---

# --description--

现在的目标是让`div`不占用整个页面的宽度。CSS的`width`属性正好适合这个需求。

你可以使用`id`选择器来定位具有`id`属性的特定元素。

你在之前的课程中学到了如何使用`id`选择器，像这样：

```css 
#cat {
  width: 250px;
}
```

使用`#menu`选择器给你的元素一个`300px`的宽度。

# --hints--

你应该有一个`#menu`选择器。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert.exists(hasDiv);
```

你应该将`width`属性设置为`300px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert.isTrue(hasWidth);
```

你的`div`应该有300px的宽度。

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert.equal(divWidth, '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 300px;
}
--fcc-editable-region--

```

