---
id: 5f35e5c44359872a137bd98f
title: ç¬?28 æ­?
challengeType: 0
dashedName: step-28
---

# --description--

ç”±äºŽå’–å•¡é¦†ä¸»è¦é”€å”®çš„äº§å“æ˜¯å’–å•¡ï¼Œå› æ­¤ä½ å¯ä»¥ä½¿ç”¨å’–å•¡è±†çš„å›¾åƒä½œä¸ºé¡µé¢èƒŒæ™¯ã€?

åˆ é™¤ `body` ç±»åž‹é€‰æ‹©å™¨ä¸­çš„æ³¨é‡ŠåŠå…¶å†…å®¹ã€?ä¹‹åŽï¼Œæ·»åŠ ä¸€ä¸?`background-image` å±žæ€§å¹¶å°†å…¶å€¼è®¾ç½®ä¸º `url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg)`ã€?

# --hints--

ä½ åº”è¯¥åˆ é™¤æ³¨é‡ŠæŽ‰çš?`background-color` å±žæ€§ã€?

```js
assert.notMatch(code, /\/\*\s*background-color:\s*burlywood;?\s*\*\//i);
```

ä½ çš„ `body` é€‰æ‹©å™¨ä¸åº”æœ‰ä»»ä½•æ³¨é‡Šã€?

```js
assert.notMatch(code, /body\s*{\s*\/\*/i);
```

ä½ åº”è¯¥å°† `background-image` å±žæ€§è®¾ç½®ä¸º `url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg)`ã€?

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg")`)
assert.isTrue(hasBackground)
```

ä½ çš„ `body` å…ƒç´ åº”è¯¥æœ‰å’–å•¡è±†çš„èƒŒæ™¯å›¾ç‰‡ã€?

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');

assert.equal(bodyBackground, `url("https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```


