---
id: 5f356ed69db0a491745e2bb6
title: 步骤 27
challengeType: 0
dashedName: step-27
---

# --description--

要将类的样式应用到`div`元素，请删除`id`属性并在`div`元素的开始标签中添加`class`属性。确保将类值设置为`menu`。

# --hints--

你的`div`应该仍然渲染。确保你没有错误格式化`<div>`标签。

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

你的`div`元素应该有`menu`类。

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

你的`div`元素应该不再有`menu` id。

```js
assert.lengthOf(document.querySelectorAll('div#menu'), 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

