---
id: 5f3c866d5414453fc2d7b480
title: 步骤 31
challengeType: 0
dashedName: step-31
---

# --description--

在现有的咖啡/价格对下方，使用 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素添加以下咖啡和价格，每个元素内嵌套两个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素。

```md
焦糖玛奇朵 3.75
南瓜香料 3.50
榛子 4.00
摩卡 4.50
```

和之前一样，第一个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素的文本应包含咖啡口味，第二个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素的文本应包含价格。

# --hints--

你应该有五个 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素。

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

每个 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素应该有两个 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素。

```js
const articles = document.querySelectorAll('article');
assert.lengthOf(articles?.[0]?.children, 2);
assert.lengthOf(articles?.[1]?.children, 2);
assert.lengthOf(articles?.[2]?.children, 2);
assert.lengthOf(articles?.[3]?.children, 2);
assert.lengthOf(articles?.[4]?.children, 2);
```

你的第一个 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素应该有带有文本`法国香草`和`3.00`的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素。

```js
const children = document.querySelector('article')?.children;
assert.match(children?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(children?.[1]?.innerText.trim(), /3\.00/i);
```

你的第二个 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素应该有带有文本`焦糖玛奇朵`和`3.75`的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素。

```js
const children = document.querySelectorAll('article')?.[1]?.children;
assert.match(children?.[0]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText.trim(), /3\.75/i);
```

你的第三个 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素应该有带有文本`南瓜香料`和`3.50`的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素。

```js
const children = document.querySelectorAll('article')?.[2]?.children;
assert.match(children?.[0]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText.trim(), /3\.50/i);
```

你的第四个 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素应该有带有文本`榛子`和`4.00`的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素。

```js
const children = document.querySelectorAll('article')?.[3]?.children;
assert.match(children?.[0]?.innerText.trim(), /Hazelnut/i);
assert.match(children?.[1]?.innerText.trim(), /4\.00/i);
```

你的第五个 [article](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L101-L120) 元素应该有带有文本`摩卡`和`4.50`的 [p](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-cafe-menu/5f3c866d5414453fc2d7b480.md#L102-L103) 元素。

```js
const children = document.querySelectorAll('article')?.[4]?.children;
assert.match(children?.[0]?.innerText.trim(), /Mocha/i);
assert.match(children?.[1]?.innerText.trim(), /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

