---
id: 5f3cade9fa77275d9f4efe62
title: æ­¥éª¤ 40
challengeType: 0
dashedName: step-40
---

# --description--

That worked, but there is still a little space on the right of the price.

ä½ å¯ä»¥ç»§ç»­å°è¯•å„ç§ç™¾åˆ†æ¯”çš„å®½åº¦ã€?ç›¸åï¼Œä½¿ç”¨é€€æ ¼é”®å°†å¸¦æœ?`price` ç±»çš„ `p` å…ƒç´ ç§»åŠ¨åˆ°å¸¦æœ?`flavor` ç±»çš„ `p` å…ƒç´ æ—ï¼Œè¿™æ ·å®ƒä»¬åœ¨ç¼–è¾‘å™¨ä¸­å°±åœ¨åŒä¸€è¡Œäº†ã€?ç¡®ä¿ä¸¤ä¸ªå…ƒç´ ä¹‹é—´æ²¡æœ‰ç©ºæ ¼ã€?

# --hints--

ä½ çš„ `p` å…ƒç´ ä¹‹é—´ä¸åº”æœ‰ä»»ä½•ç©ºæ ¼ã€?

```js
assert.match(code, /Vanilla<\/p><p/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
--fcc-editable-region--
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 49%;
}

.price {
  text-align: right;
  width: 49%;
}
```


