---
id: 5f3c866daec9a49519871816
title: 步骤 30
challengeType: 0
dashedName: step-30
---

# --description--

`article`元素通常包含多个具有相关信息的元素。在这种情况下，它将包含咖啡口味和该口味的价格。

在你的`article`元素内嵌套两个`p`元素。第一个的文本应该是`French Vanilla`，第二个的文本是`3.00`。

# --hints--

你不应该改变现有的`article`元素。

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

你的`article`元素应该有两个`p`元素。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

你的第一个`p`元素应该有文本`French Vanilla`。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

你的第二个`p`元素应该有文本`3.00`。

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

