---
id: 5f3ef6e0a81099d9a697b550
title: 步骤 65
challengeType: 0
dashedName: step-65
---

# --description--

在`address`内，添加一个`p`元素。然后，在`p`中嵌套一个链接到`https://www.MasterPuti.org`且文本为`Visit our website`的锚点(`a`)元素。

通过添加值为`_blank`的`target`属性，确保链接在新选项卡中打开。

# --hints--

你不应该修改现有的`footer`元素。

```js
assert.lengthOf(document.querySelectorAll("footer"), 1);
```

你不应该修改现有的`address`元素。

```js
assert.lengthOf(document.querySelectorAll("address"), 1);
```

你的新`p`元素应该嵌套在你的`address`元素内。你应该只有一个`p`元素。

```js
assert.lengthOf(document.querySelectorAll("footer > address > p"), 1);
assert.lengthOf(document.querySelectorAll("footer p"), 1);
```

你的新`a`元素应该嵌套在你的新`p`元素内。你应该只有一个`a`元素。

```js
assert.lengthOf(document.querySelectorAll("footer > address > p > a"), 1);
assert.lengthOf(document.querySelectorAll("footer a"), 1);
```

你的新`a`元素应该有`Visit our website`的文本。

```js
assert.equal(document.querySelector("footer > address > p > a")?.innerText.trim(), "Visit our website");
```

你的新`a`元素应该链接到`https://www.MasterPuti.org`。记住`a`元素使用`href`属性来创建链接。

```js
assert.equal(document.querySelector("footer > address > p > a")?.href, "https://www.MasterPuti.org/");
```

你的新`a`元素应该将`target`属性设置为`_blank`。

```js
assert.equal(document.querySelector("footer > address > p > a")?.target, "_blank");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
--fcc-editable-region--
      <footer>
        <address>
          <p>
            <a href="https://www.MasterPuti.org" target="_blank">Visit our website</a>
          </p>
        </address>
      </footer>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```