---
id: 5f3ef6e04559b939080db057
title: æ­¥éª¤ 55
challengeType: 0
dashedName: step-55
---

# --description--

ç”±äºŽèœå• `4` ä¸ªè¾¹çš„å†…éƒ¨é—´è·å€¼ç›¸åŒï¼Œè¯·ç§»é™¤åŽŸæœ‰çš„å››æ¡å±žæ€§å£°æ˜Žï¼Œæ”¹ç”¨å•ä¸ªå€¼ä¸º `20px` çš?`padding` å±žæ€§ä»£æ›¿ã€?

# --hints--

ä½ åº”è¯¥åˆ é™?`padding-left` å±žæ€§ã€?

```js
assert.notMatch(code, /padding-left/i);
```

ä½ åº”è¯¥åˆ é™?`padding-right` å±žæ€§ã€?

```js
assert.notMatch(code, /padding-right/i);
```

ä½ åº”è¯¥åˆ é™?`padding-top` å±žæ€§ã€?

```js
assert.notMatch(code, /padding-top/i);
```

ä½ åº”è¯¥åˆ é™?`padding-bottom` å±žæ€§ã€?

```js
assert.notMatch(code, /padding-bottom/i);
```

ä½ åº”è¯¥å°† `padding` å±žæ€§è®¾ç½®ä¸º `20px`ã€?

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['padding'] === '20px');
assert.isTrue(hasPadding);
```

ä½ çš„ `.menu` å…ƒç´ çš?`padding` å€¼åº”ä¸?`20px`ã€?

```js
const menuPadding = new __helpers.CSSHelp(document).getStyle('.menu')?.getPropertyValue('padding');
assert.equal(menuPadding, '20px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding-left: 20px;
  padding-right: 20px;
  padding-top: 20px;
  padding-bottom: 20px;
}
--fcc-editable-region--

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```


