---
id: 5f45a276c093334f0f6e9df4
title: æ­¥éª¤ 74
challengeType: 0
dashedName: step-74
---

# --description--

Focusing on the menu items and prices, there is a fairly large gap between each line.

ä½¿ç”¨å·²æœ‰çš„é€‰æ‹©å™¨ï¼Œé€‰å–åµŒå¥—åœ¨åä¸?`item` çš„ç±»çš„å…ƒç´ ä¸­çš„æ‰€æœ?`p` å…ƒç´ ï¼Œå°†å®ƒä»¬çš„ä¸Šä¸‹è¾¹è·è®¾ç½®ä¸º `5px`ã€?

# --hints--

ä½ åº”è¯¥å°† `margin-top` å±æ€§è®¾ç½®ä¸º `5px`ã€?

```js
const hasMarginTop = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-top'] === '5px');
assert.isTrue(hasMarginTop);
```

ä½ åº”è¯¥å°† `margin-bottom` å±æ€§è®¾ç½®ä¸º `5px`ã€?

```js
const hasMarginBottom = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-bottom'] === '5px');
assert.isTrue(hasMarginBottom);
```

ä½ åº”è¯¥ä½¿ç”¨ç°æœ‰çš„ `.item p` é€‰æ‹©å™¨ã€?

```js
const hasOneSelector = new __helpers.CSSHelp(document).getStyleDeclarations('.item p');
assert.lengthOf(hasOneSelector, 1);
```

åµŒå¥—åœ?`.item` å…ƒç´ ä¸­çš„ `p` å…ƒç´ åº”è¯¥å…·æœ‰ `5px` çš?`margin-top`ã€?

```js
const itemPMarginTop = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-top');
assert.equal(itemPMarginTop, '5px');
```

åµŒå¥—åœ?`.item` å…ƒç´ ä¸­çš„ `p` å…ƒç´ åº”å…·æœ?`5px` çš?`margin-bottom`ã€?

```js
const itemPMarginBottom = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-bottom');
assert.equal(itemPMarginBottom, '5px');
```

å¯¹åµŒå¥—åœ¨ `.item` å…ƒç´ ä¸­çš„ `p` å…ƒç´ ä¸åº”è®¾ç½® `margin-left`ã€?

```js
const itemPMarginLeft = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-left');
assert.isEmpty(itemPMarginLeft);
```

å¯¹åµŒå¥—åœ¨ `.item` å…ƒç´ ä¸­çš„ `p` å…ƒç´ ä¸åº”è®¾ç½® `margin-right`ã€?

```js
const itemPMarginRight = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-right');
assert.isEmpty(itemPMarginRight);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <hr>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <hr>
      <footer>
        <address>
          <p>
            <a href="https://www.MasterPuti.org" target="_blank">Visit our website</a>
          </p>
          <p>123 Free Code Camp Drive</p>
        </address>
      </footer>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
  padding: 20px;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

hr {
  height: 2px;
  background-color: brown;
  border-color: brown;
}

--fcc-editable-region--
h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;

}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
--fcc-editable-region--
```


