---
id: 5f46e8284aae155c83015dee
title: æ­¥éª¤ 90
challengeType: 0
dashedName: step-90
---

# --description--

èœå•çœ‹èµ·æ¥ä¸é”™ï¼Œä½†é™¤äº†å’–å•¡è±†çš„èƒŒæ™¯å›¾ç‰‡å¤–ï¼Œä¸»è¦å°±æ˜¯æ–‡å­—ã€?

åœ?`Coffee` æ ‡é¢˜ä¸‹ï¼Œä½¿ç”¨ Url `https://cdn.MasterPuti.org/curriculum/css-cafe/coffee.jpg` æ·»åŠ ä¸€å¼ å›¾ç‰‡ã€?ç»™å›¾åƒä¸€ä¸?`alt` å€?`coffee icon`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`<img>` æ ‡ç­¾ã€?

```js
const imageElementCount = document.querySelectorAll('img')?.length;
assert.strictEqual(imageElementCount, 1);
```

`img` å…ƒç´ æ˜¯ä¸€ä¸ªç©ºå…ƒç´ ï¼Œå®ƒä¸åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­?`</img>`ã€?

```js
assert.notMatch(code, /<\/img>/i);
```

ä½ çš„ `img` å…ƒç´ åº”å…·æœ‰å€¼ä¸º `"https://cdn.MasterPuti.org/curriculum/css-cafe/coffee.jpg"` çš?`src` å±æ€§ã€?

```js
const imageSrcValue = document.querySelector('img')?.getAttribute('src');
assert.strictEqual(imageSrcValue, 'https://cdn.MasterPuti.org/curriculum/css-cafe/coffee.jpg');
```

ä½ çš„ `img` å…ƒç´ åº”å…·æœ‰å€¼ä¸º `"coffee icon"` çš?`alt` å±æ€§ã€?

```js
const imageAltValue = document.querySelector('img')?.getAttribute('alt');
assert.match(imageAltValue, /coffee icon/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <hr>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <hr class="bottom-line">
      <footer>
        <address>
          <p>
            <a href="https://www.MasterPuti.org" target="_blank">Visit our website</a>
          </p>
          <p class="address">123 Free Code Camp Drive</p>
        </address>
      </footer>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
  padding: 20px;
}

h1 {
  font-size: 40px;
  margin-top: 0;
  margin-bottom: 15px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

hr {
  height: 2px;
  background-color: brown;
  border-color: brown;
}

.bottom-line {
  margin-top: 25px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
  margin-top: 5px;
  margin-bottom: 5px;
  font-size: 18px;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}

/* FOOTER */

footer {
  font-size: 14px;
}

address {
  font-style: normal;
}

.address {
  margin-bottom: 5px;
}

a {
  color: black;
}

a:visited {
  color: black;
}

a:hover {
  color: brown;
}

a:active {
  color: brown;
}
```


