---
id: 5f7692f7c5b3ce22a57788b6
title: æ­¥éª¤ 44
challengeType: 0
dashedName: step-44
---

# --description--

To complete the styling, add the applicable class names `flavor` and `price` to all the remaining `p` elements.

# --hints--

ä½ åº”è¯¥æœ‰äº”ä¸ª `.flavor` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('.flavor'), 5);
```

ä½ åº”è¯¥æœ‰äº”ä¸ª `.price` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('.price'), 5);
```

ä½ çš„ `.flavor` å…ƒç´ åº”è¯¥æ˜¯å¸¦æœ?`French Vanilla`ã€`Caramel Macchiato`ã€`Pumpkin Spice`ã€`Hazelnut` å’?`Mocha` æ–‡æœ¬çš?`p` å…ƒç´ ã€?

```js
const p = document.querySelectorAll('p');
const flavor = document.querySelectorAll('.flavor');
assert.equal(p?.[1], flavor?.[0]);
assert.equal(p?.[3], flavor?.[1]);
assert.equal(p?.[5], flavor?.[2]);
assert.equal(p?.[7], flavor?.[3]);
assert.equal(p?.[9], flavor?.[4]);
```

ä½ çš„ `.price` å…ƒç´ åº”è¯¥æ˜¯å¸¦æœ‰æ–‡æœ?`3.00`ã€`3.75`ã€`3.50`ã€`4.00` å’?`4.50` çš?`p` å…ƒç´ ã€?

```js
const p = document.querySelectorAll('p');
const price = document.querySelectorAll('.price');
assert.equal(p?.[2], price?.[0]);
assert.equal(p?.[4], price?.[1]);
assert.equal(p?.[6], price?.[2]);
assert.equal(p?.[8], price?.[3]);
assert.equal(p?.[10], price?.[4]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article class="item">
            <p>Caramel Macchiato</p><p>3.75</p>
          </article>
          <article class="item">
            <p>Pumpkin Spice</p><p>3.50</p>
          </article>
          <article class="item">
            <p>Hazelnut</p><p>4.00</p>
          </article>
          <article class="item">
            <p>Mocha</p><p>4.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
```


