---
id: 5f76967fad478126d6552b0d
title: æ­¥éª¤ 34
challengeType: 0
dashedName: step-34
---

# --description--

æ¥ç€ï¼Œä½ æƒ³è¦å°†ä»·æ ¼ï¼ˆpriceï¼‰å³å¯¹é½ã€?Add a class named `price` to your `p` element that has `3.00` as its text.

# --hints--

ä½ åº”è¯¥æŠŠ `price` ç±»æ·»åŠ åˆ°ä½ çš„ `p` å…ƒç´ ä¸­ã€?

```js
const el = document.querySelector('p.price');
assert.exists(el);
```

å…·æœ‰ `price` ç±»åº”è¯¥åªæœ‰ä¸€ä¸ªå…ƒç´ ã€?

```js
const elements = document.querySelectorAll('.price');
assert.lengthOf(elements, 1);
```

ä½ çš„ `price` ç±»åº”è¯¥åœ¨ `p` å…ƒç´ ä¸Šï¼Œæ–‡å­—æ˜?`3.00`ã€?

```js
const el = document.querySelector('.price');
assert.equal(el?.innerText.trim(), "3.00");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
--fcc-editable-region--
            <p>3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.MasterPuti.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}
```

