---
id: 66ad9c633fa26d3c1475eae3
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

现在，是时候添加机器人的另一条消息了。

在代码中添加另一个 `console` 语句，记录消息 `"Allow me to introduce myself."`。

# --hints--

你的代码中应该有第三个 `console.log()` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

你的 `console` 语句应该有消息 `"Allow me to introduce myself."`。不要忘记引号。

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```