---
id: 66ad8d150264db3926eccfeb
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在之前的课程中，你学习了 `let` 关键字以及如何在 JavaScript 中声明变量。

这里提醒一下如何使用 `let` 关键字声明变量：

```js
let greeting;
```

太棒了！使用 `let` 关键字来声明一个名为 `bot` 的变量吧。

**小贴士**：你在这里使用 `let` 是因为在这个研讨会的后续内容中，你会改变 `bot` 变量的值哦。

# --hints--

你的代码中应该有 `let` 关键字哦。

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

你应该声明一个名为 `bot` 的变量，加油。

```js
assert.isNotNull(bot);
```

你应该使用 `let` 关键字来声明变量，你做得很好。

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```