---
id: 66ad8ab945d266383f318cbf
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

现在你应该在控制台中看到机器人的第一条消息。

是时候添加机器人的第二条消息了。

添加另一个 `console.log` 语句，将消息 `"I am excited to talk to you."` 输出到控制台。

# --hints--

你的代码中应该有第二个 `console.log()` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

你的 `console` 语句应该有消息 `"I am excited to talk to you."`。不要忘记引号。

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```