---
id: 615f7ad94380408d971d14f6
title: Step 50
challengeType: 0
dashedName: step-50
---

# --description--

[:not](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f7ad94380408d971d14f6.md#L27-L28) 伪选择器可以用来选择所有不匹配给定CSS规则的元素。

```css
div:not(#example) {
  color: red;
}
```

上面的代码选择所有没有 [id](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f5af373a68e744a3c5a76.md#L50-L50) 为 [example](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f7ad94380408d971d14f6.md#L13-L15) 的 [div](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f5af373a68e744a3c5a76.md#L87-L87) 元素。

修改你的 [.daily-value p](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f7ad94380408d971d14f6.md#L177-L179) 选择器以排除 [.no-divider](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f7ad94380408d971d14f6.md#L77-L77) 元素。

# --hints--

你应该有一个 [.daily-value p:not(.no-divider)](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f7ad94380408d971d14f6.md#L27-L28) 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.daily-value p:not(.no-divider)'));
```

你不应该有 [.daily-value p](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f7ad94380408d971d14f6.md#L177-L179) 选择器。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.daily-value p'));
```

你不应该改变 [.daily-value p:not(.no-divider)](file:///f:/code/blocks_Chinese/blocks_english/workshop-nutritional-label/615f7ad94380408d971d14f6.md#L27-L28) 选择器中的属性。

```js
const style = new __helpers.CSSHelp(document).getStyle('.daily-value p:not(.no-divider)');
assert(style?.borderBottom === '1px solid rgb(136, 137, 137)');
```

# --seed--

## --seed-contents--

