---
id: 615f38cabc64e3556f98cc1a
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

边框可以用来分组并突出页面内容要点。

请创建一个 `.label` 选择器，并将其 `border` 属性设为 `2px solid black`。继续加油！

# --hints--

你应该有一个 `.label` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label'));
```

你的 `.label` 选择器应有 `border` 属性，值为 `2px solid black`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.border === '2px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--

--fcc-editable-region--
```
