---
id: 615f2abbe7d18d49a1e0e1c8
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

我们已为你准备好基础的 HTML 模板，接下来一起动手添加内容吧！

请在 `body` 元素内创建一个 `h1` 元素，并把文本设置为 `Nutrition Facts`。

# --hints--

你应该添加一个新的 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

该 `h1` 元素应位于 `body` 元素内。

```js
assert.equal(document.querySelector('h1')?.parentElement?.localName, 'body');
```

该 `h1` 元素的文本应为 `Nutrition Facts`。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
