---
id: 615f3b091162165948e1cb33
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

如果用浏览器的开发者工具检查 `.label` 元素，你可能会发现它实际宽度是 288px 而不是 270px。这是因为默认情况下，浏览器在计算元素尺寸时会把边框和内边距包含在内。

为了解决此问题，请创建一个 `*` 选择器，并将其 `box-sizing` 属性设为 `border-box`，这样可以统一盒模型计算。

# --hints--

你应该创建一个 `*` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

你的 `*` 选择器的 `box-sizing` 属性应设为 `border-box`。

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
