---
id: 615f3949f58e12577dcefb00
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

合理利用空白能够让页面的重要元素更突出，并引导用户的阅读视线。

请为 `.label` 选择器添加 `width` 属性，值设为 `270px`。继续，很棒的进展！

# --hints--

你的 `.label` 选择器应有 `width` 属性，值为 `270px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.width === '270px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--
.label {
  border: 2px solid black;
}
--fcc-editable-region--
```
