---
id: 615f34ecc1091b4fd5a8a484
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

请在 `head` 元素中添加两个 `link` 元素：一个用于导入字体，`rel="stylesheet"` 且 `href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800"`，另一个用于链接你的本地 `styles.css` 文件。

这样可以为页面引入 `Open Sans` 字体（字重 400、700、800），并加载样式文件，帮助页面更好地呈现。

# --hints--

你的代码中应包含两个 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

这两个 `link` 是空元素，不应有闭合标签 `</link>`。

```js
assert.notMatch(code, /<\/link>/);
```

两个 `link` 元素应位于 `head` 中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

两个 `link` 元素的 `rel` 属性应为 `stylesheet`。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

其中一个 `link` 的 `href` 应指向本地样式文件 `styles.css`（或等价属性检验）。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

另一个 `link` 的 `href` 应为 `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`，用于加载字体。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
