---
id: 615f357957e370510f21ea16
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

请为页面添加 `body` 选择器，并将其 `font-family` 设置为 `"Open Sans", sans-serif`（若字体名包含空格，记得用引号包裹）。加油，你做得到！

# --hints--

你应该包含一个 `body` 选择器来设置全局样式。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` 的 `font-family` 应设置为 `"Open Sans", sans-serif`。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
