---
id: 615f2d4150fe0d4cbd0f2628
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 元素下方添加一个 `p` 元素，内容为 `8 servings per container`。继续棒棒的节奏！

# --hints--

你应该添加一个新的 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

该 `p` 元素应位于 `body` 内。

```js
assert.equal(document.querySelector('p')?.parentElement?.localName, 'body');
```

该 `p` 元素应位于 `h1` 之后。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

该 `p` 元素的文本应为 `8 servings per container`。

```js
assert.equal(document.querySelector('p')?.innerText, '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
