---
id: 60f02e7361b68405e27b62a5
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将学习如何设计表单样式，通过设计一个注册表单来实现。

所有的HTML模板代码（`DOCTYPE`、`html`、`head`和`body`）都已经为你提供好了。

在`body`中，添加一个标题来为表单提供上下文，使用`h1`元素并添加文本`Registration Form`。

# --hints--

你应该在`body`中添加`h1`元素。

```js
assert.exists(document.querySelector('body > h1'));
```

你应该给`h1`元素添加文本`Registration Form`。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
  </body>
--fcc-editable-region--
</html>
```