---
id: 60f1a9cbd23023082e149fee
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

这样更好了。现在，通过将body的`background-color`更改为`#1b1b32`来使背景看起来舒适。然后，为了看到文本，将`color`更改为`#f5f6f7`。

# --hints--

你应该在body元素选择器内添加`background-color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

你应该给`background-color`一个值`#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

你应该在body元素选择器内添加`color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

你应该给`color`一个值`#f5f6f7`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
--fcc-editable-region--
```