---
id: 60f1922fcbd2410527b3bd89
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在之前的课程中，你学会了如何使用视口单位。

记住`vh`单位代表视口高度，等于视口`height`的1%。这使得它相对于视口高度。

是时候用一些CSS来美化项目了。首先给[body](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f1922fcbd2410527b3bd89.md#L48-L51)一个`width`为`100%`，`height`为`100vh`。

# --hints--

你应该使用[body](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f1922fcbd2410527b3bd89.md#L48-L51)元素选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你应该给[body](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f1922fcbd2410527b3bd89.md#L48-L51)一个`width`为`100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

你应该给[body](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f1922fcbd2410527b3bd89.md#L48-L51)一个`height`为`100vh`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```