---
id: 60f1a5e2d2c23707a4f9a660
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

在 fieldset 元素内添加一个 label 元素，其文本为 `Enter Your First Name:`。

# --hints--

你应该添加一个 label 元素。

```js
assert.exists(document.querySelector('label'));
```

你的 label 元素应该在 fieldset 元素内。

```js
assert.exists(document.querySelector('label')?.parentElement?.tagName === 'FIELDSET');
```

你的 label 元素的文本应该是 `Enter Your First Name:`。

```js
assert.strictEqual(document.querySelector('label')?.innerText, 'Enter Your First Name:');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form>
      <fieldset>
        <label>Enter Your First Name:</label>
      </fieldset>
    </form>
  </body>
</html>
```

```css

```