---
id: 60f5cb8875ab6a0610f05071
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

由于表单将有三个不同的部分，在form元素内添加三个fieldset元素。

# --hints--

你应该添加三个fieldset元素。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

三个fieldset元素应该在form内。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

三个fieldset元素应该是兄弟元素。

```js
assert.exists(document.querySelector('fieldset + fieldset + fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset></fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```