---
id: 60f5c3e399ff1a05629964e4
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在第一个 input 元素下方添加另一个 label 元素，其文本为 `Enter Your Last Name:`。

# --hints--

你应该添加第二个 label 元素。

```js
assert.exists(document.querySelectorAll('label')?.[1]);
```

你的 label 元素应该在 fieldset 元素内。

```js
assert.exists(document.querySelectorAll('label')?.[1]?.parentElement?.tagName === 'FIELDSET');
```

你的 label 元素的文本应该是 `Enter Your Last Name:`。

```js
assert.strictEqual(document.querySelectorAll('label')?.[1]?.innerText, 'Enter Your Last Name:');
```

你的 label 元素应该在第一个 input 元素之后。

```js
assert.exists(document.querySelectorAll('input')?.[0]?.nextElementSibling?.tagName === 'LABEL');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form>
      <fieldset>
        <label>Enter Your First Name:</label>
        <input />
        <label>Enter Your Last Name:</label>
      </fieldset>
    </form>
  </body>
</html>
```

```css

```