---
id: 60f805f813eaf2049bc2ceea
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

在每个[label](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f5dc35c07ac1078f140916.md#L61-L64)元素内嵌套一个[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)元素。确保将每个[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)放在[label](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f5dc35c07ac1078f140916.md#L61-L64)文本之后，并在冒号后添加一个空格。

# --hints--

你应该在[fieldset](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60ffe9cb47809106eda2f2c9.md#L56-L71)元素中添加四个[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)元素。

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

你应该将[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)元素嵌套在[label](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f5dc35c07ac1078f140916.md#L61-L64)元素内。

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

你应该将第一个[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)放在"label"文本`Enter Your First Name:`之后，并在冒号后添加一个空格。

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

你应该将第二个[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)放在"label"文本`Enter Your Last Name:`之后，并在冒号后添加一个空格。

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

你应该将第三个[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)放在"label"文本`Enter Your Email:`之后，并在冒号后添加一个空格。

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

你应该将第四个[input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60fabf0dd4959805dbae09e6.md#L77-L78)放在"label"文本`Create a New Password:`之后，并在冒号后添加一个空格。

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```