---
id: 60f803d5241e6a0433a523a1
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

在之前的课程中，你学会了如何使用`rem`单位。记住`rem`单位代表根`em`，并且相对于[html](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/%40types/react/index.d.ts#L2881-L2881)元素的字体大小。

由于[label](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f5dc35c07ac1078f140916.md#L61-L64)元素默认是内联的，它们都在同一行上显示，这使得它们的文本难以阅读。

为了让它们显示在单独的行上，向[label](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f5dc35c07ac1078f140916.md#L61-L64)元素添加`display: block`，并添加`0.5rem 0`的[margin](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f1922fcbd2410527b3bd89.md#L100-L100)来将它们彼此分开。

# --hints--

你应该使用[label](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f5dc35c07ac1078f140916.md#L61-L64)选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

你应该添加一个值为`block`的`display`属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

你应该添加一个值为`0.5rem 0`的[margin](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60f1922fcbd2410527b3bd89.md#L100-L100)属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```