---
id: 60f8618d191b940d62038513
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

在第五个 input 元素下方添加另一个 label 元素，其文本为 `How did you hear about us?`。

# --hints--

你应该添加第三个 label 元素。

```js
assert.exists(document.querySelectorAll('label')?.[5]);
```

你的 label 元素应该在第二个 fieldset 元素内。

```js
assert.exists(document.querySelectorAll('label')?.[5]?.parentElement?.tagName === 'FIELDSET');
```

你的 label 元素的文本应该是 `How did you hear about us?`。

```js
assert.strictEqual(document.querySelectorAll('label')?.[5]?.innerText, 'How did you hear about us?');
```

你的 label 元素应该在第五个 input 元素之后。

```js
assert.exists(document.querySelectorAll('input')?.[4]?.nextElementSibling?.tagName === 'LABEL');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form>
      <fieldset>
        <label>Enter Your First Name:</label>
        <input type="text" required />
        <label>Enter Your Last Name:</label>
        <input type="text" required />
        <label>Enter Your Email:</label>
        <input type="email" required />
        <label>Create a New Password:</label>
        <input type="password" required />
        <input type="submit" value="Submit" />
      </fieldset>
      <fieldset>
        <label>Input your age (years):</label>
        <input />
        <label>How did you hear about us?</label>
      </fieldset>
    </form>
  </body>
</html>
```

```css

```