---
id: 60f8604682407e0d017bbf7f
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

在第二个 label 元素下方添加一个 input 元素。

# --hints--

你应该添加第五个 input 元素。

```js
assert.exists(document.querySelectorAll('input')?.[4]);
```

你的 input 元素应该在第二个 fieldset 元素内。

```js
assert.exists(document.querySelectorAll('input')?.[4]?.parentElement?.tagName === 'FIELDSET');
```

你的 input 元素应该在第二个 label 元素之后。

```js
assert.exists(document.querySelectorAll('label')?.[4]?.nextElementSibling?.tagName === 'INPUT');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form>
      <fieldset>
        <label>Enter Your First Name:</label>
        <input type="text" required />
        <label>Enter Your Last Name:</label>
        <input type="text" required />
        <label>Enter Your Email:</label>
        <input type="email" required />
        <label>Create a New Password:</label>
        <input type="password" required />
        <input type="submit" value="Submit" />
      </fieldset>
      <fieldset>
        <label>Input your age (years):</label>
        <input />
      </fieldset>
    </form>
  </body>
</html>
```

```css

```