---
id: 60f80e0081e0f2052ae5b505
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

为 `input` 元素指定 `type` 属性对于浏览器了解它应该期待哪种数据非常重要。如果未指定 `type`，浏览器将默认使用 `text`。

为前两个 `input` 元素设置 `type` 属性值为 `text`，第三个设置为 `email`，第四个设置为 `password`。

`email` 类型只允许包含 `@` 和域名中包含 `.` 的电子邮件地址。`password` 类型会隐藏输入内容，并在网站未使用 HTTPS 时发出警告。

# --hints--

你应该为第一个 `input` 元素设置 `type` 属性值为 `text`。

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

你应该为第二个 `input` 元素设置 `type` 属性值为 `text`。

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

你应该为第三个 `input` 元素设置 `type` 属性值为 `email`。

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

你应该为第四个 `input` 元素设置 `type` 属性值为 `password`。

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}
```