---
id: 60fad1cafcde010995e15306
title: 步骤 38
challengeType: 0
dashedName: step-38
---

# --description--

在提交表单时，为每个可提交元素提供name属性是非常有用的，也是好的实践。该属性用于标识表单提交中的元素。

除了两个radio输入（你已经命名的）外，为每个提交的元素指定唯一的name属性。

# --hints--

你应该给要输入名字的input一个name属性。_提示：我会选择first-name_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

你应该给要输入姓氏的input一个name属性。_提示：我会选择last-name_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

你应该给email一个name属性。_提示：我会选择email_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

你应该给password一个name属性。_提示：我会选择password_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

你应该给checkbox一个name属性。_提示：我会选择terms_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

你应该给file一个name属性。_提示：我会选择file_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

你应该给number一个name属性。_提示：我会选择age_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

你应该给select元素一个name属性。_提示：我会选择referrer_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

你应该给textarea元素一个name属性。_提示：我会选择bio_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

你不该给任何option元素添加name属性。

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

你不该给任何label元素添加name属性。

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

你不该给任何fieldset元素添加name属性。

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
--fcc-editable-region--
        <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" name="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" name="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" name="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" name="password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file" name="file"/></label>
        <label for="age">Input your age (years): <input id="age" type="number" name="age" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer" name="referrer">
            <option value="">(select one)</option>
            <option value="1">MasterPuti News</option>
            <option value="2">MasterPuti YouTube Channel</option>
            <option value="3">MasterPuti Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id="bio" name="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
      <a href="https://www.MasterPuti.org/news/terms-of-service">Read our terms and conditions</a>
      <label for="terms-and-conditions">
        <input name="terms" id="terms-and-conditions" type="checkbox" required />
        I accept the terms and conditions
      </label>
      --fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}
```