---
id: 60fab4a123ce4b04526b082b
title: 步骤 24
challengeType: 0
dashedName: step-24
---

# --description--

目前用户可以在不选择单选按钮的情况下提交表单。虽然你之前使用过 `required` 属性来表示输入是必需的，但在此情况下它不起作用，因为将 `required` 属性同时添加到两个输入框会向用户传递错误的信息。

为了解决这个问题，你可以在第二个 `fieldset` 元素内的 `label` 元素之前添加一个包含文本 `账户类型 (必填)` 的 `legend` 元素，以提供所需的上下文信息。然后为"个人"输入框添加 `checked` 属性，以确保表单提交时包含所需的数据。

# --hints--

你的 `input` 元素应保持 `type` 为 `radio`。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

你应该在第二个 `fieldset` 元素内添加一个 `legend` 元素。

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

你应该在 `legend` 元素中添加 `Account type (required)` 文本。

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

你应该为"个人"输入框添加 `checked` 属性。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <legend>Account type (required)</legend>
        <label><input type="radio" name="account-type" value="personal" checked /> Personal</label>
        <label><input type="radio" name="account-type" value="business" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}
```