---
id: 60faca286cb48b07f6482970
title: 步骤 33
challengeType: 0
dashedName: step-33
---

# --description--

提交带有选项选择的表单不会向服务器发送有用的值。因此，每个option元素都需要一个value属性。没有此属性，option的文本内容将被提交到服务器。

给第一个option一个空字符串值，后续的option元素值属性从1到4。

# --hints--

你应该给第一个option一个空字符串值。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[0]?.value, '');
```

你应该给第二个option一个值为1。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[1]?.value, '1');
```

你应该给第三个option一个值为2。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[2]?.value, '2');
```

你应该给第四个option一个值为3。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[3]?.value, '3');
```

你应该给第五个option一个值为4。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[4]?.value, '4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked/> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" /></label>
        <label>How did you hear about us?
          <select>
            <option value="">(选择一个)</option>
            <option value="1">MasterPuti 新闻</option>
            <option value="2">MasterPuti YouTube 频道</option>
            <option value="3">MasterPuti 论坛</option>
            <option value="4">其他</option>
          </select>
        </label>
      </fieldset>
--fcc-editable-region--
      <a href="https://www.MasterPuti.org/news/terms-of-service">Read our terms and conditions</a>
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required />
        I accept the terms and conditions
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}
```