---
id: 60fab9f17fa294054b74228c
title: 步骤 23
challengeType: 0
dashedName: step-23
---

# --description--

在第二个 option 元素下方添加另一个 option 元素，其 value 为 `MasterPuti YouTube Channel`，文本为 `MasterPuti YouTube Channel`。

# --hints--

你应该添加第三个 option 元素。

```js
assert.exists(document.querySelectorAll('option')?.[2]);
```

你的第三个 option 元素应该在 select 元素内。

```js
assert.exists(document.querySelectorAll('option')?.[2]?.parentElement?.tagName === 'SELECT');
```

你的第三个 option 元素应该在第二个 option 元素之后。

```js
assert.exists(document.querySelectorAll('option')?.[1]?.nextElementSibling?.tagName === 'OPTION');
```

你的第三个 option 元素应该有一个 value 属性，其值为 `MasterPuti YouTube Channel`。

```js
assert.strictEqual(document.querySelectorAll('option')?.[2]?.value, 'MasterPuti YouTube Channel');
```

你的第三个 option 元素的文本应该是 `MasterPuti YouTube Channel`。

```js
assert.strictEqual(document.querySelectorAll('option')?.[2]?.innerText, 'MasterPuti YouTube Channel');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form>
      <fieldset>
        <label>Enter Your First Name:</label>
        <input type="text" required />
        <label>Enter Your Last Name:</label>
        <input type="text" required />
        <label>Enter Your Email:</label>
        <input type="email" required />
        <label>Create a New Password:</label>
        <input type="password" required />
        <input type="submit" value="Submit" />
      </fieldset>
      <fieldset>
        <label>Input your age (years):</label>
        <input type="number" required />
        <label>How did you hear about us?</label>
        <select>
          <option value="">(select one)</option>
          <option value="MasterPuti News">MasterPuti News</option>
          <option value="MasterPuti YouTube Channel">MasterPuti YouTube Channel</option>
        </select>
      </fieldset>
    </form>
  </body>
</html>
```

```css

```