---
id: 60fc22d1e64d1b04cdd4e602
title: 步骤 48
challengeType: 0
dashedName: step-48
---

# --description--

通过给 `.inline` 元素设置 `0.5em` 的右边距，在 `.inline` 元素和 `label` 文本之间添加一些间距。同时，将其他所有边距设置为 `0`。

# --hints--

你应该给 `.inline` 元素设置 `0` 的上边距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.inline')?.marginTop, '0px');
```

你应该给 `.inline` 元素设置 `0.5em` 的右边距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.inline')?.marginRight, '0.5em');
```

你应该给 `.inline` 元素设置 `0` 的下边距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.inline')?.marginBottom, '0px');
```

你应该给 `.inline` 元素设置 `0` 的左边距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.inline')?.marginLeft, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8">
    <title>注册表单</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>注册表单</h1>
    <p>请填写此表格并提供所需信息</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">输入您的名字：<input id="first-name" name="first-name" type="text" required /></label>
        <label for="last-name">输入您的姓氏：<input id="last-name" name="last-name" type="text" required /></label>
        <label for="email">输入您的邮箱：<input id="email" name="email" type="email" required /></label>
        <label for="new-password">创建新密码：<input id="new-password" name="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>账户类型（必填）</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" class="inline" checked /> 个人</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" class="inline" /> 企业</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">上传个人头像：<input id="profile-picture" type="file" name="file" /></label>
        <label for="age">输入您的年龄（岁）：<input id="age" type="number" name="age" min="13" max="120" /></label>
        <label for="referrer">您是如何了解到我们的？
          <select id="referrer" name="referrer">
            <option value="">(选择一个)</option>
            <option value="1">MasterPuti 新闻</option>
            <option value="2">MasterPuti YouTube 频道</option>
            <option value="3">MasterPuti 论坛</option>
            <option value="4">其他</option>
          </select>
        </label>
        <label for="bio">提供一段个人简介：
          <textarea id="bio" name="bio" rows="3" cols="30" placeholder="我喜欢在沙滩上编程..."></textarea>
        </label>
      </fieldset>
      <a href="https://www.MasterPuti.org/news/terms-of-service">阅读我们的服务条款</a>
      <label for="terms-and-conditions">
        <input name="terms" id="terms-and-conditions" type="checkbox" required />
        我接受服务条款
      </label>
      <input type="submit" value="提交" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
  font-family: Tahoma;
  font-size: 16px;
}

h1, p {
  margin: 1em auto;
  text-align: center;
}

form {
  width: 60vw;
  max-width: 500px;
  min-width: 300px;
  margin: 0 auto;
}

fieldset {
  border: none;
  padding: 2rem 0;
  border-bottom: 3px solid #3b3b4f;
}

fieldset:last-of-type {
  border-bottom: none;
}

label {
  display: block;
  margin: 0.5rem 0;
}

input,
textarea,
select {
  margin: 10px 0 0 0;
  width: 100%;
}

--fcc-editable-region--
.inline {
  width: unset;

}
--fcc-editable-region--

```


