---
id: 60ffec2825da1007509ddd06
title: 步骤 58
challengeType: 0
dashedName: step-58
---

# --description--

大多数浏览器会为其不同元素注入自己的默认CSS属性和值。如果你仔细观察，可能会注意到类型为 [file](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60ffec2825da1007509ddd06.md#L56-L56) 的 [input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/62ff8b9dab5ac88e4d3d43a3.md#L82-L82) 比其他文本类型的 [input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/62ff8b9dab5ac88e4d3d43a3.md#L82-L82) 元素更小。默认情况下，可输入的 [input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/62ff8b9dab5ac88e4d3d43a3.md#L82-L82) 元素会被给予 `1px 2px` 的 [padding](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60ffecefac971607ae73c60f.md#L130-L130)。

使用属性选择器，为 [type](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/62ff8e998d3e7eae14d6ae3b.md#L33-L33) 属性值为 [file](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60ffec2825da1007509ddd06.md#L56-L56) 的 [input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/62ff8b9dab5ac88e4d3d43a3.md#L82-L82) 设置与其他 [input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/62ff8b9dab5ac88e4d3d43a3.md#L82-L82) 元素相同的 [padding](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60ffecefac971607ae73c60f.md#L130-L130)。

# --hints--

你应该使用属性选择器来设置 [input](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/62ff8b9dab5ac88e4d3d43a3.md#L82-L82) 元素的样式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="file"]'));
```

你应该给 [input[type="file"]](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60ffec2825da1007509ddd06.md#L152-L152) 一个 `1px 2px` 的 [padding](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/60ffecefac971607ae73c60f.md#L130-L130)。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="file"]')?.padding, '1px 2px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8">
    <title>注册表单</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>注册表单</h1>
    <p>请填写以下必填信息完成注册</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">输入您的名字：<input id="first-name" name="first-name" type="text" required /></label>
        <label for="last-name">输入您的姓氏：<input id="last-name" name="last-name" type="text" required /></label>
        <label for="email">输入您的邮箱：<input id="email" name="email" type="email" required /></label>
        <label for="new-password">创建新密码：<input id="new-password" name="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>账户类型（必选）</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" class="inline" checked /> 个人</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" class="inline" /> 企业</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">上传个人头像：<input id="profile-picture" type="file" name="file" /></label>
        <label for="age">输入您的年龄（岁）：<input id="age" type="number" name="age" min="13" max="120" /></label>
        <label for="referrer">您是如何了解到我们的？
          <select id="referrer" name="referrer">
            <option value="">(请选择一项)</option>
            <option value="1">MasterPuti 新闻</option>
            <option value="2">MasterPuti YouTube 频道</option>
            <option value="3">MasterPuti 论坛</option>
            <option value="4">其他</option>
          </select>
        </label>
        <label for="bio">个人简介：
          <textarea id="bio" name="bio" rows="3" cols="30" placeholder="我喜欢在沙滩上编程..."></textarea>
        </label>
      </fieldset>
      <a href="https://www.MasterPuti.org/news/terms-of-service">阅读我们的服务条款</a>
      <label for="terms-and-conditions">
        <input name="terms" id="terms-and-conditions" type="checkbox" required />
        我接受服务条款
      </label>
      <input type="submit" value="提交" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
  font-family: Tahoma;
  font-size: 16px;
}

h1, p {
  margin: 1em auto;
  text-align: center;
}

form {
  width: 60vw;
  max-width: 500px;
  min-width: 300px;
  margin: 0 auto;
}

fieldset {
  border: none;
  padding: 2rem 0;
  border-bottom: 3px solid #3b3b4f;
}

fieldset:last-of-type {
  border-bottom: none;
}

label {
  display: block;
  margin: 0.5rem 0;
}

input,
textarea,
select {
  margin: 10px 0 0 0;
  width: 100%;
  min-height: 2em;
}

input, textarea {
  background-color: #0a0a23;
  border: 1px solid #0a0a23;
  color: #ffffff;
}

.inline {
  width: unset;
  margin: 0 0.5em 0 0;
  vertical-align: middle;
}

input[type="submit"] {
  display: block;
  width: 60%;
  margin: 1em auto;
  height: 2em;
  font-size: 1.1rem;
  background-color: #3b3b4f;
  border-color: white;
  min-width: 300px;
}

--fcc-editable-region--

--fcc-editable-region--
```