---
id: 62ff8e998d3e7eae14d6ae3b
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

你需要确认用户已阅读条款和条件。

添加一个label元素。在新创建的label元素中添加一个input元素，并将type属性设置为`checkbox`。将这个input元素设置为required，这样用户在不同意条款和条件的情况下无法注册。

在元素中添加id和for属性，值为`terms-and-conditions`以实现无障碍访问。

# --hints--

你应该在第三个fieldset元素之后添加一个label。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

你应该在label元素中添加一个input。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

你应该在input元素中添加值为`checkbox`的type属性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

你应该在input元素中添加一个required属性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

input元素应该有一个`terms-and-conditions`的id。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

label元素应该有一个for属性，值为`terms-and-conditions`。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required />
        I accept the terms and conditions
      </label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}
```