---
id: 62ff919a7b5612c0670923a5
title: 步骤35
challengeType: 0
dashedName: step-35
---

# --description--

将适用的表单元素及label元素关联在一起。

使用`profile-picture`、`age`、`referrer`和`bio`作为相应id属性的值。

# --hints--

第一个input元素应该有一个`profile-picture`的id。

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) input')?.[0]?.matches('#profile-picture'))
```

第二个input元素应该有一个`age`的id。

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) input')?.[1]?.matches('#age'))
```

select元素应该有一个`referrer`的id。

```js
assert.isTrue(document.querySelector('fieldset:nth-of-type(3) select')?.matches('#referrer'))
```

textarea元素应该有一个`bio`的id。

```js
assert.isTrue(document.querySelector('fieldset:nth-of-type(3) textarea')?.matches('#bio'))
```

第一个label元素应该有一个for属性，值为`profile-picture`。

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[0]?.matches('label[for="profile-picture"]'))
```

第二个label元素应该有一个for属性，值为`age`。

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[1]?.matches('label[for="age"]'))
```

第三个label元素应该有一个for属性，值为`referrer`。

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[2]?.matches('label[for="referrer"]'))
```

第四个label元素应该有一个for属性，值为`bio`。

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[3]?.matches('label[for="bio"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file" /></label>
        <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer">
            <option value="">(select one)</option>
            <option value="1">MasterPuti News</option>
            <option value="2">MasterPuti YouTube Channel</option>
            <option value="3">MasterPuti Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id="bio"></textarea>
        </label>
      </fieldset>
--fcc-editable-region--
      <a href="https://www.MasterPuti.org/news/terms-of-service">Read our terms and conditions</a>
      <label for="terms-and-conditions">
        <input name="terms" id="terms-and-conditions" type="checkbox" required />
        I accept the terms and conditions
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}
```