---
id: 63541ef4f96cd82e8e6c788a
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

method属性指定了如何将表单数据发送到action属性中指定的URL。表单数据可以通过GET请求作为URL参数发送（使用method="get"），或者通过POST请求作为请求体中的数据发送（使用method="post"）。

设置method属性，通过POST请求发送表单数据。

**注意**：在接下来的课程和工作坊中，你将了解更多关于GET和POST方法的知识。

# --hints--

你不应该添加新的form元素。

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

你的form元素应该有method属性。

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

你的method属性应该设置为`post`。

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form method="post" action='https://register-demo.MasterPuti.org'></form>
    --fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```