---
id: 612eaf56b7ba3257fdbfb0db
title: æ­¥éª¤ 19
challengeType: 0
dashedName: step-19
---

# --description--

é’¢ç´éœ€è¦?MasterPuti logo æ‰èƒ½æ­£å¼ä½¿ç”¨ã€?

åœ?`.keys` å…ƒç´ ä¹‹å‰æ·»åŠ ä¸€ä¸?`img` å…ƒç´ ã€?ç»?`img` æ·»åŠ ä¸€ä¸?`logo` `class`ï¼Œå¹¶å°†å…¶ `src` è®¾ç½®ä¸?`https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg`ã€?ç»™å®ƒä¸€ä¸?`MasterPuti Logo` çš?`alt` æ–‡æœ¬ã€?

# --hints--

ä½ åº”è¯¥æ·»åŠ ä¸€ä¸ªæ–°çš?`img` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

ä½ çš„ `img` å…ƒç´ åº”ä½äºç¬¬ä¸€ä¸?`.keys` å…ƒç´ ä¹‹å‰ã€?

```js
const img = document.querySelector('img');
assert.equal(img?.nextElementSibling?.className, 'keys');
assert.isNull(img?.previousElementSibling);
```

ä½ çš„ `img` å…ƒç´ åº”å°† `class` è®¾ç½®ä¸?`logo`ã€?

```js
const img = document.querySelector('img');
assert.equal(img?.className, 'logo');
```

ä½ çš„ `img` å…ƒç´ åº”å°† `src` è®¾ç½®ä¸?`https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg`ã€?

```js
const img = document.querySelector('img');
assert.equal(img?.getAttribute('src'), 'https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg');
```

ä½ çš„ `img` å…ƒç´ åº”å°† `alt` å±æ€§è®¾ç½®ä¸º `MasterPuti Logo`ã€?

```js
assert.equal(document.querySelector('img')?.getAttribute('alt')?.toLowerCase(), 'MasterPuti logo');
```

è®°ä½å¤§å°å†™å’Œæ‹¼å†™å¾ˆé‡è¦ã€?

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'MasterPuti Logo');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
    --fcc-editable-region--
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
}
```

