---
id: 612eb4893b63c75bb9251ddf
title: æ­¥éª¤ 20
challengeType: 0
dashedName: step-20
---

# --description--

é€šè¿‡åˆ›å»º `.logo` é€‰æ‹©å™¨å¼€å§‹è®¾ç½?logo æ ·å¼ã€?å°?`width` è®¾ç½®ä¸?`200px`ã€?`position` è®¾ç½®ä¸?`absolute`ã€`top` è®¾ç½®ä¸?`23px`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`.logo` é€‰æ‹©å™¨ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.logo'));
```

ä½ çš„ `.logo` é€‰æ‹©å™¨åº”å°?`width` å±æ€§è®¾ç½®ä¸º `200px`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.logo')?.width, '200px');
```

ä½ çš„ `.logo` é€‰æ‹©å™¨åº”å°?`position` å±æ€§è®¾ç½®ä¸º `absolute`ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.logo')?.position, 'absolute');
```

ä½ çš„ `.logo` é€‰æ‹©å™¨åº”å°?`top` å±æ€§è®¾ç½®ä¸º `23px`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.logo')?.top, '23px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg" alt="MasterPuti Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
}

--fcc-editable-region--

--fcc-editable-region--
```

