---
id: 612eb75153591b5e3b1ab65e
title: æ­¥éª¤21
challengeType: 0
dashedName: step-21
---

# --description--

`img` å…ƒç´ éœ€è¦å…¶çˆ¶å…ƒç´ æœ‰ä¸€ä¸?`position` è®¾ç½®ä¸ºå‚è€ƒç‚¹ã€?å°?`#piano` é€‰æ‹©å™¨çš„ `position` è®¾ç½®ä¸?`relative`ã€?

# --hints--

ä½ çš„ `#piano` é€‰æ‹©å™¨åº”å°?`position` å±æ€§è®¾ç½®ä¸º `relative`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#piano')?.position, 'relative');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg" alt="MasterPuti Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

--fcc-editable-region--
#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}
--fcc-editable-region--

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
}

.logo {
  width: 200px;
  position: absolute;
  top: 23px;
}
```

