---
id: 684cd59ec94b5e063bf776f0
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

如果能在这个浅红色背景周围加上边框就更好了。

要在 Tailwind CSS 中设置 `border-width`，你可以使用 `border-number` 工具类，其中 `number` 表示边框的粗细程度。

下面是应用 `4px` 边框宽度到所有边的示例：

```html
<p class="border-4">这是设置边框宽度的方法。</p>
```

在你的 `div` 开始标签中，添加正确的工具类来设置所有边的 `border-width` 为 `2px`。

# --hints--

你的 `div` 开始标签应该有 `border-2` 类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```