---
id: 684ccbdc1ef89e042b744e09
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

接下来的部分是添加关闭按钮。

在你的 `p` 元素下面，添加一个 `button` 元素，设置 `aria-label` 为 `"Dismiss"`。在该 `button` 元素内部，添加一个带有 `aria-hidden` 属性的 `span` 元素。`span` 元素的文本应该是 `&times;`。

*注意*：由于这个实践项目只专注于使用 Tailwind CSS 样式，按钮不会添加任何类型的交互功能。

# --hints--

你应该有一个 `button` 元素。

```js
assert.exists(document.querySelector("button"));
```

你的 `button` 元素应该有一个 `aria-label` 属性。

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

你的 `button` 元素应该有一个设置为 `"Dismiss"` 的 `aria-label` 属性。

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

你的 `button` 元素内部应该嵌套一个 `span` 元素。

```js
assert.exists(document.querySelector("button span"));
```

你的 `span` 元素应该有一个 `aria-hidden` 属性。

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

你的 `span` 元素应该有文本 `&times;`。

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```