---
id: 684cc59eed2706fa50e8933c
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个实践项目中，你将学习如何通过构建一个错误信息组件来开始使用 Tailwind CSS 的基础知识。

在前几个步骤中，你将专注于构建初始的 HTML 标记。

在 `body` 元素内部，添加一个 `div` 元素。然后在该 `div` 元素内部，添加一个 `p` 元素。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你应该有一个 `p` 元素。

```js
assert.exists(document.querySelector("p"));
```

你的 `p` 元素应该在 `div` 元素内部。

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```