---
id: 684d009994a3840b98c9aa3a
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

现在你看到 `div` 元素的宽度占据了整个页面的宽度。要在 Tailwind CSS 中应用宽度，你有几个选择，比如使用 `w-number` 模式设置固定宽度或使用 `w-fraction` 模式设置百分比。

下面是应用宽度为其父容器 `66%` 的示例：

```html
<p class="w-2/3">这是百分比宽度。</p>
```

在你的 `div` 开始标签中，应用宽度为其父容器 `50%` 的设置。

# --hints--

你的 `div` 开始标签应该有 `w-1/2` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```