---
id: 684d07cbfa43a20d6a002ed7
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

`div` 元素的样式开始成形了。但如果文本周围有一些空间就更好了。

要使用 Tailwind CSS 为元素应用 `padding`，你可以使用以下模式：`p-number`，其中 `number` 表示要在四边应用的像素数。

下面是为元素应用 `8px` `padding` 的示例：

```html
<p class="p-8">这是内边距。</p>
```

在你的 `div` 开始标签中，使用正确的工具类来应用 `4px` 的 `padding`。

# --hints--

你的 `div` 开始标签应该有 `p-4` 工具类。

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>错误信息组件</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```