---
id: 68da6138d70960c197ec5ef0
title: 步骤 6
challengeType: 20
dashedName: step-6
---

# --description--

现在你将实现主要的 BFS 循环。创建一个 [while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) 循环，只要 [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14) 不为空（[while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14) 计算结果为 `True`）就继续循环。

在循环内部，打印 [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14)。

# --hints--

你的 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 函数中应该有一个 [while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) 循环。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

你的 [while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) 循环应该有条件 [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14)。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

你应该在 [while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) 循环内打印 [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14)。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []


    return result
--fcc-editable-region--
```
