---
id: 68da6138d70960c197ec5eef
title: 步骤 5
challengeType: 20
dashedName: step-5
---

# --description--

对于广度优先搜索方法，你将使用队列来跟踪构建括号组合时的不同状态。每个状态将表示为一个包含三个元素的元组：

- 正在构建的当前字符串
- 到目前为止使用的开括号数量
- 到目前为止使用的闭括号数量

创建一个名为 [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14) 的变量并将其初始化为包含一个元组的列表：`('', 0, 0)`。这代表起始状态，字符串为空且未使用括号。

# --hints--

你应该在你的 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 函数中声明一个名为 [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14) 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

你应该将 [queue](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eef.md#L14-L14) 初始化为包含元组 `('', 0, 0)` 的列表。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'

    result = []
    return result
--fcc-editable-region--
```
