---
id: 68da6138d70960c197ec5eec
title: 步骤 2
challengeType: 20
dashedName: step-2
---

# --description--

在实现核心算法之前，你需要验证输入。[pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 参数应该是一个整数，因为它代表要生成的括号对数量。

在函数开头添加一个 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句，检查 [pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 是否不是整数。为此使用 [isinstance()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L17-L17) 函数。

如果条件为真，返回字符串 `The number of pairs should be an integer`。

# --hints--

你的 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 函数中应该有一个 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

你的 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句应该使用 `not isinstance(pairs, int)` 检查 [pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 是否不是整数。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

当输入不是整数时，你应该返回字符串 `The number of pairs should be an integer`。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):

    return []
--fcc-editable-region--
```
