---
id: 68da6138d70960c197ec5eed
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

接下来，你需要验证括号对的数量至少为一，因为不能用零对或负对括号生成括号组合。

添加另一个 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句检查 [pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 是否小于 `1`。如果这个条件为真，返回字符串 `The number of pairs should be at least 1`。

# --hints--

你的 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 函数中应该有第二个 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

你的第二个 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句应该检查 [pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 是否小于 `1`。

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

当 [pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 小于 `1` 时，你应该返回字符串 `The number of pairs should be at least 1`。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'

    return []
--fcc-editable-region--
```
