---
id: 68da604b8b6012a33e2344fb
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个工作坊中，你将使用广度优先搜索（BFS）方法实现一个生成所有有效括号组合的函数。例如，两对括号的有效组合是 `(())` 和 `()()`。

首先创建一个名为 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 的函数，它只有一个参数 [pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9)。现在，从函数中返回一个空列表。

# --hints--

你应该定义一个名为 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 的函数。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

你的 [gen_parentheses](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 函数应该有一个名为 [pairs](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eec.md#L2-L9) 的参数。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

你的函数应该返回一个空列表。

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
