---
id: 68da6138d70960c197ec5ef1
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

现在你需要确定何时构建了一个完整的括号组合。完整组合的长度等于括号对数的两倍（因为每对贡献一个开括号和一个闭括号）。

在你的 [while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) 循环内，添加一个 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句检查 `len(current) == 2 * pairs`。当这个条件为真时，将 [current](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16) 添加到 [result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eed.md#L6-L6) 列表中。

# --hints--

你的 [while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) 循环内应该有一个 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

你的 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句应该检查 `len(current) == 2 * pairs`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

当新 [if](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L45-L47) 语句条件为真时，你应该将 [current](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16) 添加到 [result](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68da6138d70960c197ec5eed.md#L6-L6) 列表中。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'

    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)

--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
