---
id: 68e00c5349cbddfef6dbc053
title: 步骤 7
challengeType: 20
dashedName: step-7
---

# --description--

在你的 [while](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-linked-list-class/68867c3898b09115d8c5140f.md#L49-L55) 循环内，使用 `queue.pop(0)` 从队列中移除并获取第一个元素。这实现了 BFS 特有的先进先出（FIFO）行为。

将这个元组解包为三个变量：[current](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16)、[opens_used](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16) 和 [closes_used](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16)。

# --hints--

你应该使用 `queue.pop(0)` 从队列中获取第一个元素。

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

你应该将结果解包为三个变量：[current](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16)、[opens_used](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16) 和 [closes_used](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-breadth-first-search/68e00c5349cbddfef6dbc053.md#L16-L16)。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'

    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)

--fcc-editable-region--
    return result
```
