---
id: 6859400f8f7f872e50e1238c
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

[User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类需要一个方法来向其他用户发送电子邮件。发送电子邮件时，你将创建一个新的 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 对象并将其添加到接收者的收件箱中。

在你的 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类中创建一个名为 [send_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L46-L46) 的方法。这个方法应该接受 [receiver](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6859400f8f7f872e50e1238c.md#L21-L21)、[subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 和 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 参数。

在方法内部，创建一个新的 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 对象，参数值如下，并将其分配给名为 [email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6859400f8f7f872e50e1238c.md#L21-L21) 的变量：

- [sender](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21)：[self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40)（当前用户）
- [receiver](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21)：[receiver](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6859400f8f7f872e50e1238c.md#L21-L21)（接收邮件的用户）
- [subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21)：[subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21)（邮件的主题）
- [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21)：[body](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21)（邮件的正文）

# --hints--

你应该在 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类中定义一个 [send_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L46-L46) 方法。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

[send_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L46-L46) 方法应该有参数：[self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40)、[receiver](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6859400f8f7f872e50e1238c.md#L21-L21)、[subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 和 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21)。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

该方法应该创建一个名为 [email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6859400f8f7f872e50e1238c.md#L21-L21) 的 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 对象，并带有正确的参数。

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--