---
id: 68593ea22eac7e228ed585ce
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

电子邮件应该能够跟踪它们是否已被阅读。在 [__init__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 方法中添加一个 [read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68593ea22eac7e228ed585ce.md#L21-L21) 属性，并将其默认设置为 [False](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68593ea22eac7e228ed585ce.md#L21-L21)，因为新邮件开始时是未读的。

# --hints--

你应该在 [__init__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 方法中将 [False](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68593ea22eac7e228ed585ce.md#L21-L21) 分配给 [self.read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68593ea22eac7e228ed585ce.md#L21-L21)。

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
