---
id: 688754d6d59578df1313cb3c
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

现在是时候连接 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 和 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 类，以便电子邮件能够真正送达！

用户需要有适当的 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 对象而不是简单的列表，[send_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L46-L46) 方法应该将电子邮件发送到接收者的收件箱。

更新 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类的 [__init__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L22-L22) 方法，使每个用户获得一个实际的 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 对象，而不是之前创建的空列表。

# --hints--

你应该更新 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类以初始化 [self.inbox = Inbox()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L39-L39)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
