---
id: 68b930ee0e7b8fb81895e59a
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

现在你将添加一个将电子邮件标记为已读的方法。当有人打开一封电子邮件时，你想要将其状态从未读更改为已读。这个方法很简单 - 它只需要将 [read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L22-L22) 属性设置为 [True](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L48-L48)。

在你的 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中添加一个名为 [mark_as_read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25) 的方法。这个方法应该只接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 作为参数，并将 [read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L22-L22) 属性设置为 [True](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L48-L48)。

# --hints--

你应该在 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中定义一个名为 [mark_as_read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25) 的方法，该方法应该只接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

该方法应该将 [self.read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L22-L22) 设置为 [True](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L48-L48)。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
