---
id: 68b96d6a290d52d5e57a7485
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

移除 [email_obj](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24) 和以下的打印语句。

# --hints--

你的代码中不应该有创建 [email_obj](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24) 的那一行。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

你的代码中不应该有 [print(email_obj.sender)](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

你的代码中不应该有 [print(email_obj.subject)](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

你的代码中不应该有 [print(email_obj.read)](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

你的代码中不应该有 [email_obj.mark_as_read()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
