---
id: 68b9a011daf902ff1eba40a8
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

现在电子邮件已经接收，移除对象创建、方法调用和打印语句。

# --hints--

你的代码中不应该有 [alice = User("Alice")](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L22-L22)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("alice = User(\\"Alice\\")")`));
    }
})
```

你的代码中不应该有 [bob = User("Bob")](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L22-L22)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("bob = User(\\"Bob\\")")`));
    }
})
```

你的代码中不应该有 [alice.send_email(bob, "Hello", "Hi Bob, how are you?")](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L46-L46)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("alice.send_email(bob, \\"Hello\\", \\"Hi Bob, how are you?\\")")`));
    }
})
```

你的代码中不应该有 [print(len(bob.inbox.emails))](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L39-L39)。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(len(bob.inbox.emails))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

alice.send_email(bob, "Hello", "Hi Bob, how are you?")

print(len(bob.inbox.emails))
--fcc-editable-region--
```
