---
id: 61fd6b7c83dbf54a08cf0498
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在你刚创建的 `div` 元素中，添加三个 `span`，每个 `span` 的 `class` 设置为 `year`，并按顺序填入文本：`2019`、`2020`、`2021`。干得漂亮，继续保持节奏！

# --hints--

提示：`div` 中应包含三个 `span` 子元素。

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

提示：每个 `span` 的 `class` 应为 `year`，这样便于统一样式。

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

提示：第一个 `span` 的文本应为 `2019`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

提示：第二个 `span` 的文本应为 `2020`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

提示：第三个 `span` 的文本应为 `2021`。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
