---
id: 61fd7a160ed17960e971f28b
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

在第 4 个 `tr` 中添加一个 `th`，文本为 `Total Assets`。请把其中的单词 `Assets` 用 `span` 包裹并设置 `class="sr-only"`，这样有助于无障碍体验。

然后添加 3 个 `td`，按顺序填入 `$579`、`$736`、`$809`，并把第三个 `td` 的 `class` 设为 `current`。完成得很棒，继续前进！

# --hints--

提示：第 4 个 `tr` 中应该包含一个 `th`。

```js
assert.exists(
  document
    .querySelector('tbody')
    ?.querySelectorAll('tr')?.[3]
    ?.querySelector('th')
);
```

提示：`th` 的文本应为 `Total Assets`（保持原样以通过测试）。

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableHeader = tableRow.querySelector('th');
assert.strictEqual(tableHeader.innerText.trim(), 'Total Assets');
```

提示：把 `Assets` 包裹在 `span` 中，并设置 `class="sr-only"`，以优化无障碍体验。

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.strictEqual(span.textContent.trim(), 'Assets');
```

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.isTrue(span?.classList?.contains('sr-only'));
```

提示：该行应包含 3 个 `td` 元素。

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableDatas = tableRow.querySelectorAll('td');
assert.lengthOf(tableDatas, 3);
```

提示：第一个 `td` 的文本应为 `$579`。

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '$579');
```

提示：第二个 `td` 的文本应为 `$736`。

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[1];
assert.strictEqual(tableData?.textContent.trim(), '$736');
```

提示：第三个 `td` 的文本应为 `$809`，并且 `class` 应为 `current`。

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[2];
assert.strictEqual(tableData.textContent.trim(), '$809');
assert.isTrue(tableData?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
--fcc-editable-region--
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
