---
id: 61fd7b3fcaa5406257abc5d1
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

现在开始处理第二个表格吧：先为它添加 `caption`（文本为 `Liabilities`），然后加入 `thead` 和 `tbody`。你已经做得很好了，继续加油！

# --hints--

提示：第二个 `table` 应包含 `caption` 元素。

```js
assert.strictEqual(
  document.querySelectorAll('table')?.[1]?.children?.[0]?.localName,
  'caption'
);
```

提示：`caption` 的文本应为 `Liabilities`（请保持该文本以便测试通过）。

```js
assert.strictEqual(
  document.querySelectorAll('table')?.[1]?.querySelector('caption')
    ?.textContent.trim(),
  'Liabilities'
);
```

提示：第二个 `table` 应包含 `thead`。

```js
assert.isNotNull(
  document.querySelectorAll('table')?.[1]?.querySelector('thead')
);
```

提示：第二个 `table` 应包含 `tbody`。

```js
assert.isNotNull(
  document.querySelectorAll('table')?.[1]?.querySelector('tbody')
);
```

提示：`thead` 应紧跟在 `caption` 之后，`tbody` 应紧跟在 `thead` 之后。

```js
const table = document.querySelectorAll('table')?.[1];
assert.strictEqual(
  table?.querySelector('caption')?.nextElementSibling?.localName,
  'thead'
);
assert.strictEqual(
  table?.querySelector('thead')?.nextElementSibling?.localName,
  'tbody'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
