---
id: 61fd6cc9475a784b7776233e
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

在现有 `div` 元素下方创建一个新的 `div`，并把它的 `class` 设置为 `table-wrap`，它将作为表格的容器。

在这个容器中添加三个 `table` 元素。我们会用 CSS 把它们在视觉上组合成一个表格，但保持独立的 `table` 元素有助于屏幕阅读器按合理的文档流阅读内容。做得很棒，继续！

# --hints--

提示：创建一个新的 `div` 容器来包裹你的表格。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

提示：这个容器应有 `class="table-wrap"`。

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

提示：`.table-wrap` 应位于之前的 `div` 之后。

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

提示：在 `.table-wrap` 中添加三个 `table` 元素，每个都是独立的表结构，便于辅助工具理解。

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
