---
id: 61fd70336ebb3e4f62ee81ba
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML 表格使用 `caption` 元素来说明表格的用途。`caption` 应该是 `table` 的第一个子元素，不过可以通过 CSS 的 `caption-side` 属性来调整显示位置。

请在第一个 `table` 中添加一个 `caption`，并设置其文本为 `Assets`。你做得很棒，继续加油！

# --hints--

提示：第一个 `table` 的第一个子元素应该是 `caption`。

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

提示：`caption` 的文本应为 `Assets`（请保持该文本以便测试通过）。

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
