---
id: 61fd6ab779390f49148773bb
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `h1` 元素下面创建一个 `div`，并设置它的 `id` 为 `years`。这个元素不需要被屏幕阅读器读取，请把 `aria-hidden` 属性设为 `true` 来隐藏它 — 做得很好，继续！

# --hints--

提示：在 `h1` 后面创建一个新的 `div` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling?.localName,
  'div'
);
```

提示：新 `div` 的 `id` 应为 `years`（请保持该值以通过测试）。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('id'), 'years');
```

提示：将 `aria-hidden` 设置为 `true`，表示此元素对屏幕阅读器不可见。

```js
assert.strictEqual(
  document.querySelector('div')?.getAttribute('aria-hidden'),
  'true'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
