---
id: 61fd77f7ad2aeb5ae34d07d6
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

在第二个 `tr` 中，添加一个 `th` 元素，文本为 `Checking Our primary transactional account.`。请把除 `Checking` 之外的文字包裹在 `class="description"` 的 `span` 中。

随后添加三个 `td`，按顺序填入 `$54`、`$56`、`$53`，并将第三个 `td` 的 `class` 设为 `current`。你做得很棒，继续完成后续行！

# --hints--

提示：第二个 `tr` 中应包含一个 `th` 元素。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isNotNull(tableRow?.querySelector('th'));
```

提示：`th` 的文本应为 `Checking Our primary transactional account.`（请保留该文本以便测试通过）。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Checking Our primary transactional account.'
);
```

提示：把 `Our primary transactional account.` 包裹在 `span` 中。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Our primary transactional account.'
);
```

提示：该 `span` 的 `class` 应为 `description`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

提示：确保此行包含 3 个 `td` 元素。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

提示：第一个 `td` 的文本应为 `$54`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$54');
```

提示：第二个 `td` 的文本应为 `$56`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$56');
```

提示：第三个 `td` 的文本应为 `$53`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$53');
```

提示：第三个 `td` 的 `class` 应为 `current`。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data"></tr>
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
