---
id: 61fd66c687e610436494c6f1
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `section` 元素里，添加一个 `h1`，并在 `h1` 内嵌入一个 `span` 元素 — 很好的一步，继续保持！

# --hints--

提示：确认 `section` 中包含一个 `h1` 元素，这能帮助页面结构更清晰。

```js
assert.strictEqual(
  document.querySelector('section')?.children?.[0]?.localName,
  'h1'
);
```

提示：在 `h1` 内加入 `span`，方便后续单独样式化标题的一部分。

```js
assert.strictEqual(
  document.querySelector('h1')?.children?.[0]?.localName,
  'span'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
