---
id: 61fd5a93fd62bb35968adeab
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

从在 `head` 元素里添加一个样式表的 `link` 标签开始吧，这会让你的项目马上变得更有样式 — 加油，你可以的！

# --hints--

提示：请确保你的 HTML 中包含一个 `link` 标签，用来引入样式表。

```js
assert.match(code, /<link/);
```

提示：把 `link` 标签放在 `head` 元素内部，这样样式表才能正确加载。

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

提示：`link` 标签需要 `rel="stylesheet"` 属性来标识它是样式表。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

提示：`href` 属性应指向你的样式文件 `styles.css`（可带或不带相对路径 `./`）。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
