---
id: 61fd67a656743144844941cb
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

屏幕阅读器会按照文档流逐个读取 HTML 元素。我们希望页面视觉上显示的顺序与为屏幕阅读器优化的语义顺序都能兼顾。

请把现有的 `span` 元素的 `class` 设置为 `flex`，并在这个 `span` 内添加两个新的 `span` 元素：第一个的文本为 `AcmeWidgetCorp`，第二个的文本为 `Balance Sheet`。我们会用 CSS 在视觉上颠倒它们的显示顺序，但保持 HTML 的语义顺序能让辅助工具更友好 — 你做得很好！

# --hints--

提示：把现有 `span` 的 `class` 设置为 `flex`，方便我们用 CSS 调整视觉顺序。

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

提示：在 `.flex` 中添加两个 `span` 子元素，分别作为公司名和标题的容器。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

提示：这两个新 `span` 不需要 `class`，保持简单即可。

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

提示：第一个新 `span` 应包含文本 `AcmeWidgetCorp`（保持原样以便测试通过）。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

提示：第二个新 `span` 应包含文本 `Balance Sheet`（保持原样以便测试通过）。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
