---
id: acda2fb1324d9b0fa741e6b5
title: Build a Confirm the Ending Tool
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

在这个实验中，你将实现一个函数，检查一个字符串是否以给定的目标字符串结尾。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该创建一个名为 `confirmEnding` 的函数，它接受两个参数：要检查的字符串和要检查的目标字符串。
2. 如果第一个字符串以第二个字符串结尾，函数应该返回 `true`，否则返回 `false`。
3. 你不应该使用 `.endsWith()` 方法；而是使用其中一个JavaScript子字符串方法来实现这个功能。

# --hints--

你应该创建一个名为 `confirmEnding` 的函数。

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` 应该接受2个参数。

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` 应该返回 `true`。

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` 应该返回 `true`。

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` 应该返回 `false`。

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` 应该返回 `false`。

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` 应该返回 `true`。

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` 应该返回 `true`。

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` 应该返回 `false`。

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` 应该返回 `false`。

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` 应该返回 `false`。

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` 应该返回 `true`。

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

你的代码不应该使用内置方法 `.endsWith()` 来解决这个实验。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
