---
id: 657bdc55a322aae1eac3838f
title: 创建一个回文检测器
challengeType: 25
dashedName: build-a-palindrome-checker
demoType: onClick
---

# --description--

<dfn>回文</dfn> 是指忽略标点符号、大小写和间距，正向和反向读法相同的单词或短语。

**注意：** 检查回文时，你需要先去除 **所有非字母数字的字符**（标点、空格和符号），并将所有字母都转换成大写或都转换成小写。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `id` 为 `"text-input"` 的 `input` 元素。
2. 你应该有一个 `id` 为 `"check-btn"` 的 `button` 元素。
3. 你应该有一个 `div`、`span` 或 `p` 元素，其 `id` 为 `"result"`。
4. 当你点击 `#check-btn` 元素但没有给 `#text-input` 元素输入一个值时，会出现提示 `Please input a value`。
5. 当 `#text-input` 元素只包含字母 `A`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `A is a palindrome`。
6. 当 `#text-input` 元素包含文本 `eye`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `eye is a palindrome`。
7. 当 `#text-input` 元素包含文本 `_eye`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `_eye is a palindrome`。
8. 当 `#text-input` 元素包含文本 `race car`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `race car is a palindrome`。
9. 当 `#text-input` 元素包含文本 `not a palindrome`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `not a palindrome is not a palindrome`。
10. 当 `#text-input` 元素包含文本 `A man, a plan, a canal Panama` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `A man, a plan, a canal. Panama is a palindrome`。
11. 当 `#text-input` 元素包含文本 `never odd or even`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `never odd or even is a palindrome`。
12. 当 `#text-input` 元素包含文本 `nope`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `nope is not a palindrome`。
13. 当 `#text-input` 元素包含文本 `almostomla`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `almostomla is not a palindrome`。
14. 当 `#text-input` 元素包含文本 `My age is 0, 0 si ega ym.`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `My age is 0, 0 si ega ym. is a palindrome`。
15. 当 `#text-input` 元素包含文本 `1 eye for of 1 eye.` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `1 eye for of 1 eye. is not a palindrome`。
16. 当 `#text-input` 元素包含文本 `0_0 (: /-\ :) 0-0`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `0_0 (: /-\ :) 0-0 is a palindrome`。
17. 当 `#text-input` 元素包含文本 `five|\_/|four`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `five|\_/|four is not a palindrome`。

# --hints--

你应该有一个 `id` 为 `"text-input"` 的 `input` 元素。

```js
const el = document.getElementById('text-input');
assert.strictEqual(el?.nodeName?.toLowerCase(), 'input');
```

你应该有一个 `id` 为 `"check-btn"` 的 `button` 元素。

```js
const el = document.getElementById('check-btn');
assert.strictEqual(el?.nodeName?.toLowerCase(), 'button');
```

你应该有一个 `div`、`span` 或 `p` 元素，其 `id` 为 `"result"`。

```js
const el = document.getElementById('result');
assert.oneOf(el?.nodeName?.toLowerCase(), ['div', 'span', 'p']);
```

当你点击 `#check-btn` 元素但没有给 `#text-input` 元素输入一个值时，会出现提示 `Please input a value`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
let alertMessage;
window.alert = (message) => alertMessage = message; // Override alert and store message

inputEl.value = '';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(alertMessage?.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'please input a value');
```

当 `#text-input` 元素只包含字母 `A` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `A is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'A';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'a is a palindrome');
```

当 `#text-input` 元素包含文本 `eye` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `eye is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'eye';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'eye is a palindrome');
```

当 `#text-input` 元素包含文本 `_eye` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `_eye is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = '_eye';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), '_eye is a palindrome');
```

当 `#text-input` 元素包含文本 `race car` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `race car is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'race car';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'race car is a palindrome');
```

当 `#text-input` 元素包含文本 `not a palindrome` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `not a palindrome is not a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'not a palindrome';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'not a palindrome is not a palindrome');
```

当 `#text-input` 元素包含文本 `A man, a plan, a canal. Panama` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `A man, a plan, a canal. Panama is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'A man, a plan, a canal. Panama';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'a man, a plan, a canal. panama is a palindrome');
```

当 `#text-input` 元素包含文本 `never odd or even` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `never odd or even is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'never odd or even';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'never odd or even is a palindrome');
```

当 `#text-input` 元素包含文本 `nope` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `nope is not a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'nope';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'nope is not a palindrome');
```

当 `#text-input` 元素包含文本 `almostomla` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `almostomla is not a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'almostomla';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'almostomla is not a palindrome');
```

当 `#text-input` 元素包含文本 `My age is 0, 0 si ega ym.` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `My age is 0, 0 si ega ym. is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'My age is 0, 0 si ega ym.';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'my age is 0, 0 si ega ym. is a palindrome');
```

当 `#text-input` 元素包含文本 `1 eye for of 1 eye.` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `1 eye for of 1 eye. is not a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = '1 eye for of 1 eye.';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), '1 eye for of 1 eye. is not a palindrome');
```

当 `#text-input` 元素包含文本 `0_0 (: /-\ :) 0-0` 并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `0_0 (: /-\ :) 0-0 is a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = '0_0 (: /-\ :) 0-0';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), '0_0 (: /-\ :) 0-0 is a palindrome');
```

当 `#text-input` 元素包含文本 `five|\_/|four`，并点击 `#check-btn` 元素时，`#result` 元素应该包含文本 `five|\_/|four is not a palindrome`。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

inputEl.value = 'five|\_/|four';
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();
assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), 'five|\_/|four is not a palindrome');
```

当 `#text-input` 元素包含一个字母数字回文时，`#result` 元素应该正确识别它是一个回文。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

const characters = 'abcdefghijklmnopqrstuvwxyz0123456789';
const charactersLength = characters.length;

const firstLetter = characters.charAt(Math.floor(Math.random() * charactersLength));
const secondLetter = characters.charAt(Math.floor(Math.random() * charactersLength));
const thirdLetter = characters.charAt(Math.floor(Math.random() * charactersLength));
const fourthLetter = characters.charAt(Math.floor(Math.random() * charactersLength));

const phrase = firstLetter + secondLetter + thirdLetter + fourthLetter + fourthLetter + thirdLetter + secondLetter + firstLetter;

inputEl.value = phrase;
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();

assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), phrase + ' is a palindrome');
```

当 `#text-input` 元素包含一个随机的非回文字母数字字符序列时，`#result` 元素应该说它不是回文。

```js
const inputEl = document.getElementById('text-input');
const checkBtn = document.getElementById('check-btn');
const resultEl = document.getElementById('result');

let characters = 'abcdefghijklmnopqrstuvwxyz0123456789';
let charactersLength = characters.length;

const firstIndex = Math.floor(Math.random() * charactersLength);
const firstLetter = characters.charAt(firstIndex);
characters = characters.slice(0,firstIndex) + characters.slice(firstIndex + 1);
charactersLength--; 


const secondIndex = Math.floor(Math.random() * charactersLength); 
const secondLetter = characters.charAt(secondIndex);
characters = characters.slice(0,secondIndex) + characters.slice(secondIndex + 1);
charactersLength--; 

const thirdIndex = Math.floor(Math.random() * charactersLength); 
const thirdLetter = characters.charAt(thirdIndex);
characters = characters.slice(0,thirdIndex) + characters.slice(thirdIndex + 1);
charactersLength--; 

const fourthIndex = Math.floor(Math.random() * charactersLength); 
const fourthLetter = characters.charAt(fourthIndex);
characters = characters.slice(0,fourthIndex) + characters.slice(fourthIndex + 1);
charactersLength--; 

const phrase = firstLetter + secondLetter + thirdLetter + fourthLetter;

inputEl.value = phrase;
inputEl.dispatchEvent(new Event('change'))
checkBtn.click();

assert.strictEqual(resultEl.innerText.trim().replace(/[.,?!]+$/g, '').toLowerCase(), phrase + ' is not a palindrome');
```

