---
id: a105e963526e7de52b219be9
title: 实现一个唯一的有序并集
challengeType: 26
dashedName: implement-a-unique-sorted-union
---

# --description--

**目标：** 完成下面的用户故事并通过所有测试来完成这个实验。

**用户故事：**

1. 你应该有一个名为 `uniteUnique` 的函数。
1. `uniteUnique` 函数应该接受两个或更多的数组作为参数。
1. 函数应该返回一个新的数组，该数组包含来自参数数组的唯一值，按照它们在参数中首次出现的顺序排列。例如，输入 `[1, 2, 4], [2, 3, 5]` 的输出应该是 `[1, 2, 4, 3, 5]`。

# --hints--

你应该有一个 `uniteUnique` 函数。

```js
assert.isFunction(uniteUnique);
```

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` 应该返回 `[1, 3, 2, 5, 4]`。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` 应该返回 `[1, 2, 3, 5]`。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` 应该返回 `[1, 2, 3, 5, 4, 6, 7, 8]`。

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` 应该返回 `[1, 3, 2, 5, 4, 6]`。

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` 应该返回 `[1, 3, 2, 5, 4]`。

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const uniteUnique = (...arrays) => {
  return arrays.reduce((a, b) => {
    return [
      ...a,
      ...b.filter((e, currentIndex) => b.indexOf(e) === currentIndex && a.indexOf(e) === -1)
    ];
  }, []);
};
```