---
id: 688c90634eb5ae69845ac35d
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你将创建一个表情反应器应用程序。HTML模板和所有CSS样式都已经为你提供好了。

首先创建一个 `main` 元素。

在 `main` 元素内部创建一个 `h1` 元素，它有一个 `class` 为 `title`，文本内容为 `你今天感觉如何？`。

# --hints--

你应该有一个 `main` 元素。

```js
assert.exists(document.querySelector("main"));
```

你应该有一个作为 `main` 子元素的 `h1`。

```js
assert.exists(document.querySelector("main > h1"));
```

`h1` 元素应该有一个 `title` 类。

```js
assert.equal(document.querySelector("main > h1")?.className, "title");
```

`h1` 元素应该包含文本 `你今天感觉如何？`。

```js
assert.equal(document.querySelector("main > h1")?.textContent.trim(), "你今天感觉如何？");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```