---
id: 6899b0d1825ff24e9fe8d747
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `h1` 元素下方，添加一个带有 `description` 类的 `p` 元素。

在 `p` 元素内部写入文本 `点击下方按钮来评价你的情绪。`。

# --hints--

你应该在 `h1` 元素下方创建一个新的 `p` 元素。

```js
assert.exists(document.querySelector("h1 + p"));
```

新的 `p` 元素应该有一个 `description` 类。

```js
assert.equal(document.querySelector("h1 + p")?.className, "description");
```

新的 `p` 元素应该有文本 `点击下方按钮来评价你的情绪。`。

```js
assert.equal(document.querySelector("h1 + p")?.textContent.trim(), "点击下方按钮来评价你的情绪。");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Emoji Reactor</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <main>
--fcc-editable-region--
    <h1 class="title">你今天感觉如何？</h1>
    <p class="description">点击下方按钮来评价你的情绪。</p>
--fcc-editable-region--
  </main>
</body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #efefef;
  --white: #ffffff;
  --very-dark-blue: #0a0a23;
  --light-purple: #a78bfa;
  --very-light-purple: #c4b5fd;
  --purple: #8b5cf6;
}

body {
  background-color: var(--very-dark-blue);
  color: var(--light-grey);
  font-family: sans-serif;
}

main {
  text-align: center;
  padding: 10px;
}

.btn-container,
button {
  display: flex;
  justify-content: space-evenly;
  align-items: center;
}

.btn-container {
  flex-direction: column;
}

.emoji-btn {
  width: 70%;
  cursor: pointer;
  color: var(--white);
  background-color: var(--light-purple);
  background-image: linear-gradient(
    to bottom,
    var(--very-light-purple),
    var(--light-purple)
  );
  border: 3px solid var(--purple);
  border-radius: 8px;
  padding: 10px;
  font-size: 1.5rem;
  margin: 10px 0;
  transition: background-color 0.3s ease, transform 0.2s ease;
}

@media (min-width: 768px) {
  .emoji-btn {
    width: 30%;
  }
}

.emoji-btn:hover {
  background-color: var(--purple);
  background-image: none;
}

.title {
  margin-top: 15px;
  font-size: 2rem;
}

.description {
  font-size: 1.4rem;
  margin: 20px 0;
}
```