---
id: 60f030d388cb74067cf291c3
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `form` element below the `p` element.

# --hints--

You should add a `form` element.

```js
assert.exists(document.querySelector('form'));
```

Your `form` element should be within the `body` element.

```js
assert.exists(document.querySelector('form')?.parentElement?.tagName === 'BODY');
```

Your `form` element should come after the `p` element.

```js
assert.exists(document.querySelector('p')?.nextElementSibling?.tagName === 'FORM');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
</html>

```

```css

```