---
id: 60f1a9cbd23023082e149fee
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

That's better. Now make the background a little more comfortable by changing the `body`'s `background-color` to `#1b1b32`. Then to see the text, set the `color` to `#f5f6f7`.

# --hints--

You should add `background-color` within the `body` element selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

You should give `background-color` a value of `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

You should add `color` within the `body` element selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

You should give `color` a value of `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
--fcc-editable-region--
```