---
id: 60f5dc35c07ac1078f140916
title: æ­¥éª¤ 10
challengeType: 0
dashedName: step-10
---

# --description--

åœ?`label` å…ƒç´ ä¸­æ·»åŠ ä»¥ä¸‹æ–‡æœ¬ï¼š

- `Enter Your First Name:`
- `Enter Your Last Name:`
- `Enter Your Email:`
- `Create a New Password:`

# --hints--

ç¬¬ä¸€ä¸?`label` çš„æ–‡æœ¬åº”è¯¥ä¸º `Enter Your First Name:`ã€?

```js
assert.match(document.querySelector('label')?.innerHTML, /Enter Your First Name:/i);
```

ç¬¬äºŒä¸?`label` çš„æ–‡æœ¬åº”è¯¥ä¸º `Enter Your Last Name:`ã€?

```js
assert.match(document.querySelector('fieldset > label:nth-child(2)')?.innerHTML, /Enter Your Last Name:/i);
```

ç¬¬ä¸‰ä¸?`label` çš„æ–‡æœ¬åº”è¯¥ä¸º `Enter Your Email:`ã€?

```js
assert.match(document.querySelector('fieldset > label:nth-child(3)')?.innerHTML, /Enter Your Email:/i);
```

ç¬¬å››ä¸?`label` çš„æ–‡æœ¬åº”è¯¥ä¸º `Create a New Password:`ã€?

```js
assert.match(document.querySelector('fieldset > label:nth-child(4)')?.innerHTML, /Create a New Password:/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
--fcc-editable-region--
      <fieldset>
        <label></label>
        <label></label>
        <label></label>
        <label></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```


