---
id: 60f5c3e399ff1a05629964e4
title: æ­¥éª¤6
challengeType: 0
dashedName: step-6
---

# --description--

æ­£å¦‚æ ‡é¢˜æ‰€è¿°ï¼Œä½ æ­£åœ¨åˆ›å»ºè¡¨å•ã€?å› æ­¤ï¼Œåœ¨ `p` å…ƒç´ ä¹‹åï¼Œæ’å…¥ä¸€ä¸?`form` å…ƒç´ ï¼Œå…¶ `action` å±æ€§å€¼ä¸º `https://register-demo.MasterPuti.org`ã€?

# --hints--

ä½ åº”è¯¥åœ¨ `p` å…ƒç´ ä¹‹åæ·»åŠ ä¸€ä¸?`form` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('p + form'));
```

ä½ åº”è¯¥ç»™ `form` å…ƒç´ æ·»åŠ ä¸€ä¸?`action` å±æ€§ã€?

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

ä½ åº”è¯¥å°† `action` çš„å€¼è®¾ä¸?`https://register-demo.MasterPuti.org`ã€?

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.MasterPuti.org/');
```

ä½ çš„ `form` å…ƒç´ åº”è¯¥æ²¡æœ‰ç»“æŸæ ‡ç­¾ `</form>`ã€?

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```


