---
id: 60f5d2776c854e069560fbe6
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

第一个 `fieldset` 将包含姓名、电子邮件和密码字段。首先在第一个 `fieldset` 中添加四个 `label` 元素。

# --hints--

总共应该有三个 `fieldset` 元素。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

`fieldset` 元素都应该是 `form` 元素的直接子元素。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

四个 `label` 元素都应该是第一个 `fieldset` 元素的直接子元素。确保你关闭了 `label` 元素。

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

最后两个 `fieldset` 元素应该是空的。

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

确保你关闭了 `label` 元素。

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```