---
id: 60f81167d0d4910809f88945
title: æ­¥éª¤15
challengeType: 0
dashedName: step-15
---

# --description--

`type` ä¸?`submit` çš„ç¬¬ä¸€ä¸?`input` å…ƒç´ ä¼šè‡ªåŠ¨è®¾ç½®ä¸ºæäº¤å…¶æœ€è¿‘çš„ `form` çˆ¶å…ƒç´ ã€?

ä¸ºäº†å¤„ç†è¡¨æ ¼æäº¤ï¼Œåœ¨æœ€åä¸€ä¸?`fieldset` å…ƒç´ ä¹‹åæ·»åŠ ä¸€ä¸?`input` å…ƒç´ ï¼Œ`type` å±æ€§è®¾ç½®ä¸º `submit`ï¼Œ`value` å±æ€§è®¾ç½®ä¸º `Submit`ã€?

# --hints--

ä½ åº”è¯¥åœ¨æœ€åä¸€ä¸?`fieldset` å…ƒç´ åæ·»åŠ?`input` å…ƒç´ ã€?

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

ä½ åº”è¯¥ç»™ `input` å…ƒç´ ä¸€ä¸?`type` å±æ€§ï¼Œå€¼ä¸º `submit`ã€?

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

ä½ åº”è¯¥ç»™ `input` å…ƒç´ ä¸€ä¸?`value` å±æ€§ï¼Œå€¼ä¸º `Submit`ã€?

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>

--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```


