---
id: 60f84ec41116b209c280ba91
title: 步骤 18
challengeType: 0
dashedName: step-18
---

# --description--

当 `type="password"` 时，你可以使用 `pattern` 属性来定义一个正则表达式，密码必须匹配该正则表达式才能被视为有效。

为密码 `input` 元素添加一个 `pattern` 属性，要求输入的内容匹配：`[a-z0-5]{8,}`

以上是一个正则表达式，它匹配八个或更多小写字母或数字 `0` 到 `5`。然后，删除 `minlength` 属性，并尝试一下。

# --hints--

你应该为密码 `input` 元素添加一个 `pattern` 属性。

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

你应该将 `pattern` 属性的值设置为 `[a-z0-5]{8,}`。

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

你应该从密码 `input` 元素中移除 `minlength` 属性。

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>注册表单</h1>
    <p>请填写此表格并提供所需信息</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">输入您的名字：<input id="first-name" type="text" required /></label>
        <label for="last-name">输入您的姓氏：<input id="last-name" type="text" required /></label>
        <label for="email">输入您的邮箱：<input id="email" type="email" required /></label>
        <label for="new-password">创建一个新密码：<input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="提交" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```


