---
id: 60f8618d191b940d62038513
title: æ­¥éª¤ 22
challengeType: 0
dashedName: step-22
---

# --description--

ä½ æƒ³æ¯æ¬¡åªèƒ½é€‰æ‹©ä¸€ä¸ªå•é€‰è¾“å…¥ã€?ä½†æ˜¯ï¼Œè¡¨å•å¹¶ä¸çŸ¥é“è¿™äº›å•é€‰è¾“å…¥æ˜¯ç›¸å…³çš„ã€?

ç»™æ¯ä¸€ä¸ªå•é€‰æŒ‰é’®æ·»åŠ ç›¸åŒçš„ `name` å±æ€§å¹¶è®¾ç½®å…¶å€¼ä¸º `account-type`ï¼Œä»¥å…³è”å•é€‰æŒ‰é’®ã€?ç°åœ¨ï¼Œæ— æ³•åŒæ—¶é€‰æ‹©ä¸¤ä¸ªå•é€‰è¾“å…¥äº†ã€?

# --hints--

ä½ åº”è¯¥ä¸ºç¬¬ä¸€ä¸ªå•é€‰è¾“å…¥æ·»åŠ?`name` å±æ€§ï¼Œå€¼ä¸º `account-type`ã€?

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(1) input[type="radio"]')?.name, 'account-type');
```

ä½ åº”è¯¥ä¸ºç¬¬äºŒä¸ªå•é€‰è¾“å…¥æ·»åŠ?`name` å±æ€§ï¼Œå€¼ä¸º `account-type`ã€?

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(2) input[type="radio"]')?.name, 'account-type');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label><input type="radio" /> Personal</label>
        <label><input type="radio" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```


