---
id: 60f803d5241e6a0433a523a1
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在第三个 [input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素下方添加另一个 [label](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素，其文本为 `Create a New Password:`。

# --hints--

你应该添加第四个 [label](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素。

```js
assert.exists(document.querySelectorAll('label')?.[3]);
```

你的 [label](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该在 [fieldset](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素内。

```js
assert.exists(document.querySelectorAll('label')?.[3]?.parentElement?.tagName === 'FIELDSET');
```

你的 [label](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素的文本应该是 `Create a New Password:`。

```js
assert.strictEqual(document.querySelectorAll('label')?.[3]?.innerText, 'Create a New Password:');
```

你的 [label](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素应该在第三个 [input](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/@types/node/stream.d.ts#L53-L53) 元素之后。

```js
assert.exists(document.querySelectorAll('input')?.[2]?.nextElementSibling?.tagName === 'LABEL');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form>
      <fieldset>
        <label>Enter Your First Name:</label>
        <input />
        <label>Enter Your Last Name:</label>
        <input />
        <label>Enter Your Email:</label>
        <input />

      </fieldset>
    </form>
  </body>
</html>

```

```css

```
