---
id: 60f805f813eaf2049bc2ceea
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

在每个 `label` 中嵌套一个 `input` 元素。确保在 `label` 文本后添加每个 `input`，并在冒号后包含一个空格。

# --hints--

你应该在 `fieldset` 元素中添加四个 `input` 元素。

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

你应该在 `label` 元素中嵌套 `input` 元素。

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

你应该在 `label` 文本 `Enter Your First Name:` 后添加第一个 `input`，并在冒号后包含一个空格。

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

你应该在 `label` 文本 `Enter Your Last Name:` 后添加第二个 `input`，并在冒号后包含一个空格。

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

你应该在 `label` 文本 `Enter Your Email:` 后添加第三个 `input`，并在冒号后包含一个空格。

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

你应该在 `label` 文本 `Create a New Password:` 后添加第四个 `input`，并在冒号后包含一个空格。

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
  --fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```