---
id: 60fabf0dd4959805dbae09e6
title: 步骤 24
challengeType: 0
dashedName: step-24
---

# --description--

在第一个 `label` 元素后添加另一个 `label` 元素，其文本为 `输入您的年龄（岁）:`。然后，在其中嵌套一个 `type` 为 `number` 的 `input`。

接下来，为 `input` 添加一个值为 `13` 的 `min` 属性，因为 13 岁以下的用户不应注册。同时，用户的年龄可能不会超过 120 岁；添加一个值为 `120` 的 `max` 属性。

现在，如果有人尝试提交超出此范围的表单，将会出现警告，并且表单不会被提交。试试看吧。

# --hints--

你应该在第三个 `fieldset` 中添加一个 `label`，位于现有的 `label` 之后。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

你应该给 `label` 设置文本 `输入您的年龄（岁）:`。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), '输入您的年龄（岁）:');
```

你应该给 `label` 一个 `type` 为 `number` 的 `input`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

你应该给 `input` 一个值为 `13` 的 `min` 属性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

你应该给 `input` 一个值为 `120` 的 `max` 属性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="UTF-8">
    <title>注册表单</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>注册表单</h1>
    <p>请填写以下必填信息完成注册</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">输入您的名字: <input id="first-name" type="text" required /></label>
        <label for="last-name">输入您的姓氏: <input id="last-name" type="text" required /></label>
        <label for="email">输入您的邮箱: <input id="email" type="email" required /></label>
        <label for="new-password">创建新密码: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>账户类型 (必选)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> 个人</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> 企业</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>上传个人头像: <input type="file" /></label>
        <label>输入您的年龄（岁）: <input type="number" min="13" max="120" /></label>
      </fieldset>
--fcc-editable-region--
      <a href="https://www.MasterPuti.org/news/terms-of-service">阅读我们的服务条款</a>
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required />
        我接受服务条款
      </label>
      <input type="submit" value="提交" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```


