---
id: 60ffe4f4ec18cd04dc470c56
title: ç¬?52 æ­?
challengeType: 0
dashedName: step-52
---

# --description--

ä½ å¸Œæœ?`select` å…ƒç´ ä¿æŒç™½è‰²èƒŒæ™¯ï¼Œä½†æ˜¯ç°åœ¨å®ƒçš?`min-height` å±æ€§å’Œ `input` ä»¥åŠ `textarea` å…ƒç´ çš„ä¸ä¸€è‡´ã€?

ç§»åŠ¨ `min-height` å±æ€§å’Œå€¼ï¼Œä½¿è¿™ä¸‰ä¸ªå…ƒç´ å…·æœ‰ç›¸åŒçš?`min-height` å€¼ï¼Œä¿æŒ `select` å…ƒç´ èƒŒæ™¯ä¸ºç™½è‰²ã€?

# --hints--

ä½ åº”è¯¥ç§»åŠ?`input, textarea, select` é€‰æ‹©å™¨çš„ `min-height`ã€?

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input, textarea, select')?.minHeight);
```

ä½ åº”è¯¥è®¾ç½?`input, textarea, select` é€‰æ‹©å™¨çš„ `min-height` ä¸?`2em`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input, textarea, select')?.minHeight, '2em');
```

ä½ åº”è¯¥ç§»é™?`input, textarea` é€‰æ‹©å™¨çš„ `min-height` å£°æ˜ã€?

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('input, textarea')?.minHeight);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" name="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" name="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" name="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" name="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" class="inline" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" class="inline" /> Business</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file" name="file" /></label>
        <label for="age">Input your age (years): <input id="age" type="number" name="age" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer" name="referrer">
            <option value="">(select one)</option>
            <option value="1">MasterPuti News</option>
            <option value="2">MasterPuti YouTube Channel</option>
            <option value="3">MasterPuti Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id="bio" name="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
      <a href="https://www.MasterPuti.org/news/terms-of-service">Read our terms and conditions</a>
      <label for="terms-and-conditions">
        <input name="terms" id="terms-and-conditions" type="checkbox" required />
        I accept the terms and conditions
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
  font-family: Tahoma;
  font-size: 16px;
}

h1, p {
  margin: 1em auto;
  text-align: center;
}

form {
  width: 60vw;
  max-width: 500px;
  min-width: 300px;
  margin: 0 auto;
}

fieldset {
  border: none;
  padding: 2rem 0;
  border-bottom: 3px solid #3b3b4f;
}

fieldset:last-of-type {
  border-bottom: none;
}

label {
  display: block;
  margin: 0.5rem 0;
}

input,
textarea,
select {
  margin: 10px 0 0 0;
  width: 100%;
}

input, textarea {
  background-color: #0a0a23;
  border: 1px solid #0a0a23;
  color: #ffffff;
  min-height: 2em;
}

.inline {
  width: unset;
  margin: 0 0.5em 0 0;
  vertical-align: middle;
}
--fcc-editable-region--

```


