---
id: 62ff8b9dab5ac88e4d3d43a3
title: æ­¥éª¤ 13
challengeType: 0
dashedName: step-13
---

# --description--

æŒ‰ç…§å¯è®¿é—®æ€§æœ€ä½³å®è·µï¼Œä½¿ç”¨ `for` å±æ€§å°† `input` å…ƒç´ å’?`label` å…ƒç´ å…³è”åœ¨ä¸€èµ·ã€?

ä½¿ç”¨ `first-name`ã€`last-name`ã€`email` å’?`new-password` ä½œä¸ºç›¸åº”çš?`id` å±æ€§çš„å€¼ã€?

# --hints--

ç¬¬ä¸€ä¸?`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`first-name` `id`ã€?

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

ç¬¬äºŒä¸?`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`last-name` `id`ã€?

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

ç¬¬ä¸‰ä¸?`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`email` `id`ã€?

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

ç¬¬å››ä¸?`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`new-password` `id`ã€?

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

ç¬¬ä¸€ä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `first-name`ã€?

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

ç¬¬äºŒä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `last-name`ã€?

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

ç¬¬ä¸‰ä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `email`ã€?

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

ç¬¬å››ä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `new-password`ã€?

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```


