---
id: 62ff919a7b5612c0670923a5
title: æ­¥éª¤35
challengeType: 0
dashedName: step-35
---

# --description--

å°†é€‚ç”¨çš„è¡¨å•å…ƒç´ åŠå…?`label` å…ƒç´ å…³è”åœ¨ä¸€èµ·ã€?

ä½¿ç”¨ `profile-picture`ã€`age`ã€`referrer` å’?`bio` ä½œä¸ºç›¸åº”çš?`id` å±æ€§çš„å€¼ã€?

# --hints--

ç¬¬ä¸€ä¸?`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`profile-picture` `id`ã€?

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) input')?.[0]?.matches('#profile-picture'))
```

ç¬¬äºŒä¸?`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`age` `id`ã€?

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) input')?.[1]?.matches('#age'))
```

`select` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`referrer` `id`ã€?

```js
assert.isTrue(document.querySelector('fieldset:nth-of-type(3) select')?.matches('#referrer'))
```

`textarea` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`bio` `id`ã€?

```js
assert.isTrue(document.querySelector('fieldset:nth-of-type(3) textarea')?.matches('#bio'))
```

ç¬¬ä¸€ä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `profile-picture`ã€?

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[0]?.matches('label[for="profile-picture"]'))
```

ç¬¬äºŒä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `age`ã€?

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[1]?.matches('label[for="age"]'))
```

ç¬¬ä¸‰ä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `referrer`ã€?

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[2]?.matches('label[for="referrer"]'))
```

ç¬¬å››ä¸?`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±æ€§ï¼Œå€¼ä¸º `bio`ã€?

```js
assert.isTrue(document.querySelectorAll('fieldset:nth-of-type(3) label')?.[3]?.matches('label[for="bio"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" /></label>
        <label>How did you hear about us?
          <select>
            <option value="">(select one)</option>
            <option value="1">MasterPuti News</option>
            <option value="2">MasterPuti YouTube Channel</option>
            <option value="3">MasterPuti Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label>Provide a bio:
          <textarea></textarea>
        </label>
      </fieldset>
--fcc-editable-region--
      <a href="https://www.MasterPuti.org/news/terms-of-service">Read our terms and conditions</a>
      <label for="terms-and-conditions">
        <input name="terms" id="terms-and-conditions" type="checkbox" required />
        I accept the terms and conditions
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```


