---
id: 62ff8e998d3e7eae14d6ae3b
title: æ­¥éª¤ 26
challengeType: 0
dashedName: step-26
---

# --description--

ä½ éœ€è¦ç¡®è®¤ç”¨æˆ·å·²é˜…è¯»æ¡æ¬¾å’Œæ¡ä»¶ã€?

æ·»åŠ ä¸€ä¸?`label` å…ƒç´ ã€?åœ¨æ–°åˆ›å»ºçš?`label` å…ƒç´ ä¸­æ·»åŠ ä¸€ä¸?`input` å…ƒç´ ï¼Œå¹¶å°?`type` å±žæ€§è®¾ç½®ä¸º `checkbox`ã€?å°†è¿™ä¸?`input` å…ƒç´ è®¾ç½®ä¸?`required`ï¼Œè¿™æ ·ç”¨æˆ·åœ¨ä¸åŒæ„æ¡æ¬¾å’Œæ¡ä»¶çš„æƒ…å†µä¸‹æ— æ³•æ³¨å†Œã€?

åœ¨å…ƒç´ ä¸­æ·»åŠ  `id` å’?`for` å±žæ€§ï¼Œå€¼ä¸º `terms-and-conditions` ä»¥å®žçŽ°æ— éšœç¢è®¿é—®ã€?

# --hints--

ä½ åº”è¯¥åœ¨ç¬¬ä¸‰ä¸?`fieldset` å…ƒç´ ä¹‹åŽæ·»åŠ ä¸€ä¸?`label`ã€?

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

ä½ åº”åœ?`label` å…ƒç´ ä¸­æ·»åŠ ä¸€ä¸?`input` ã€?

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

ä½ åº”åœ?`input` å…ƒç´ ä¸­æ·»åŠ å€¼ä¸º `checkbox` çš?`type` å±žæ€§ã€?

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

ä½ åº”åœ?`input` å…ƒç´ ä¸­æ·»åŠ ä¸€ä¸?`required` å±žæ€§ã€?

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

`input` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`terms-and-conditions` `id`ã€?

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

`label` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`for` å±žæ€§ï¼Œå€¼ä¸º `terms-and-conditions`ã€?

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.MasterPuti.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```


