---
id: 63541ef4f96cd82e8e6c788a
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

[method](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L52-L52)属性指定了如何将表单数据发送到[action](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L52-L52)属性中指定的URL。表单数据可以通过[GET](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L10-L10)请求作为URL参数发送（使用[method="get"](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L10-L10)），或者通过[POST](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L12-L12)请求作为请求体中的数据发送（使用[method="post"](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L12-L12)）。

设置[method](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L52-L52)属性，通过[POST](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L12-L12)请求发送表单数据。

**注意**：在接下来的课程和工作坊中，你将了解更多关于[GET](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L10-L10)和[POST](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L12-L12)方法的知识。

# --hints--

你不应该添加新的[form](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L52-L52)元素。

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

你的[form](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L52-L52)元素应该有[method](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L52-L52)属性。

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

你的[method](file:///E:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-registration-form/63541ef4f96cd82e8e6c788a.md#L52-L52)属性应该设置为`post`。

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.MasterPuti.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```