---
id: 68f890b73bbfc9efa7379e7b
title: 调试有缺陷应用程序中的类型强制错误
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

你刚刚加入了一家本地网络开发公司，你的第一个任务是清理前开发人员留下的有缺陷的代码。

他们试图进行简单的算术运算，但出了问题，结果没有意义。你的任务是审查、调试和纠正代码，使其按预期执行并更易于阅读。

**目标：** 完成以下用户故事并让所有测试通过以完成实验。

**用户故事：**

1. 你应该有一个名为`firstResult`的变量，它正确地将数字`5`和`10`相加得到值`15`。
2. 你应该有一个名为`secondResult`的变量，它正确地从`8`中减去`5`得到值`3`。
3. 你应该有一个名为`thirdResult`的变量，通过相加两个数字产生值`6`。用数字替换表达式中当前使用的布尔值。
4. 你应该有一个名为`fourthResult`的变量，通过相加两个数字产生值`8`。用数字替换表达式中当前使用的布尔值。
5. 你应该有一个名为`fifthResult`的变量，正确地将两个数字相乘得到乘积`20`。用数字替换表达式中当前使用的字符串。
6. 你应该有一个名为`sixthResult`的变量，正确地将两个数字相加得到值`22`。用数字替换表达式中当前使用的`null`。

# --hints--

你应该有一个名为`firstResult`的变量。

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult`应该有数值`15`。

```js
assert.strictEqual(firstResult, 15);
```

你应该有一个名为`secondResult`的变量。

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult`应该有数值`3`。

```js
assert.strictEqual(secondResult, 3);
```

你应该有一个名为`thirdResult`的变量。

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult`应该有数值`6`。

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult`不应该包含布尔值（`true`或`false`）。

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

你应该有一个名为`fourthResult`的变量。

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult`应该有数值`8`。

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult`不应该包含布尔值（`true`或`false`）。

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

你应该有一个名为`fifthResult`的变量。

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult`应该有数值`20`。

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult`不应该包含字符串乘法（例如 "10" * 2）。

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

你应该有一个名为`sixthResult`的变量。

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult`应该有数值`22`。

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult`不应该在其表达式中包含`null`。

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```