---
id: 673b3d6b7ef7318eef926d5a
title: Build a Mood Board
challengeType: 25
dashedName: build-a-mood-board
demoType: onClick
---
# --description--

情绪板是图像和文本的拼贴，传达关于特定主题的一般想法、目标或感觉。

在这个实验中，您将使用可重用的React组件创建一个情绪板。CSS已经为您提供。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该创建并导出一个[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件，该组件接受三个属性：[color](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-color-picker/66c0701d3475cd92421b9ac3.md#L35-L53)、[image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)和[description](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)。
2. 您的[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件应返回一个带有[mood-board-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L190-L201)类的div作为其顶级元素。
3. 您应该使用内联样式将.mood-board-item元素的背景颜色设置为[color](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-color-picker/66c0701d3475cd92421b9ac3.md#L35-L53)属性的值。
4. 您应该在.mood-board-item元素内渲染一个带有[mood-board-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L203-L208)类的img元素，并将其src属性设置为[image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)属性的值。
5. 您应该在.mood-board-item元素内渲染一个带有[mood-board-text](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L210-L213)类的h3元素，并将其文本设置为[description](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)属性的值。
6. 您应该创建并导出一个[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)。
7. 您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应返回一个div作为其顶级元素。
8. 您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应渲染一个带有[mood-board-heading](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L174-L179)类的h1元素，文本为"Destination Mood Board"。
9. 您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应渲染一个带有[mood-board](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L181-L188)类的div。
10. 您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应在.mood-board元素内渲染至少三个[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件，每个组件都应传递具有有效值的[color](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-color-picker/66c0701d3475cd92421b9ac3.md#L35-L53)、[image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)和[description](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)属性。

如果您愿意，可以在情绪板中使用以下图像：

- `https://cdn.MasterPuti.org/curriculum/labs/pathway.jpg`
- `https://cdn.MasterPuti.org/curriculum/labs/shore.jpg`
- `https://cdn.MasterPuti.org/curriculum/labs/grass.jpg`
- `https://cdn.MasterPuti.org/curriculum/labs/ship.jpg`
- `https://cdn.MasterPuti.org/curriculum/labs/santorini.jpg`
- `https://cdn.MasterPuti.org/curriculum/labs/pigeon.jpg`

# --hints--

您应该导出一个[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件。

```js
assert.isFunction(window.index.MoodBoardItem, 1);
```

您的[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件应返回一个带有[mood-board-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L190-L201)类的div作为其顶级元素。

```js
const item = document.getElementsByClassName('mood-board-item')[0];
assert.exists(item);
assert.equal(item.tagName, 'DIV');
```

.mood-board-item元素的背景颜色应使用内联样式设置为[color](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-color-picker/66c0701d3475cd92421b9ac3.md#L35-L53)属性的值。

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { color: "red" });
  const moodBoardItem = container.querySelector(".mood-board-item");
  assert.equal(moodBoardItem.style.backgroundColor, "red");
```

您的[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件应渲染一个带有[mood-board-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L203-L208)类的img元素，并将其src设置为[image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)属性的值。

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { image: "https://cdn.MasterPuti.org/curriculum/labs/pathway.jpg" });
  const img = container.querySelector(".mood-board-image");
  assert.exists(img);
  assert.equal(img.tagName, "IMG");
  assert.equal(img.src, "https://cdn.MasterPuti.org/curriculum/labs/pathway.jpg");
```

您的[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件应渲染一个带有[mood-board-text](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L210-L213)类的h3元素，并将其文本设置为[description](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)属性的值。

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { description: "Caribbean" });
  const text = container.querySelector(".mood-board-text");
  assert.exists(text);
  assert.equal(text.tagName, "H3");
  assert.equal(text.textContent, "Caribbean");
```

您应该导出一个[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件。

```js
assert.isFunction(window.index.MoodBoard, 1);
```

您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应返回一个div作为其顶级元素。

```js
const root = document.getElementById('root');
assert.equal(root.children[0].tagName, 'DIV');
```

您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应渲染一个带有[mood-board-heading](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L174-L179)类的h1元素，文本为"Destination Mood Board"。

```js
const heading = document.getElementsByClassName('mood-board-heading')[0];
assert.exists(heading);
assert.equal(heading.tagName, 'H1');
assert.equal(heading.textContent, 'Destination Mood Board');
```

您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应渲染至少三个[MoodBoardItem](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/workshop-todo-app/66b8957bad1d0a11a980f0d7.md#L37-L47)组件，每个组件都应传递具有有效值的[color](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-color-picker/66c0701d3475cd92421b9ac3.md#L35-L53)、[image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)和[description](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-blog-post-card/66b8957bad1d0a11a980f0d8.md#L45-L64)属性。

```js
const moodboard = document.getElementsByClassName('mood-board')[0];
const items = moodboard.querySelectorAll(".mood-board-item");
assert.isAtLeast(items.length, 3);     

items.forEach((item) => {
  const color = item.style.backgroundColor;
  const img = item.querySelector("img");
  const text = item.querySelector("h3");

  assert.isAtLeast(color.length, 1);
  assert.isAtLeast(img.src.length, 1);
  assert.isAtLeast(text.textContent.length, 1);
});
```

您的[MoodBoard](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-mood-board/673b3d6b7ef7318eef926d5a.md#L313-L351)组件应渲染到页面的#root元素。

```js
const root = document.getElementById('root');
const moodBoard = root.getElementsByClassName('mood-board')[0];
assert.exists(moodBoard);
```

# --seed--

## --seed-contents--```
