---
id: a77dbc43c33f39daa4429b4f
title: Build a Boolean Check Function
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

在本练习中，你需要实现一个函数，用来判断某个值是否属于布尔原始值。

布尔原始值是 `true` 和 `false`。

**目标：**实现下面的用户需求并通过所有测试，完成该练习。

**用户需求：**

1. 请实现名为 `booWho` 的函数，它接收一个参数。
1. 如果传入的参数是布尔原始值，函数应返回 `true`。
1. 如果参数是其他任意值，函数应返回 `false`。

# --hints--

你应该实现 `booWho` 这个函数。

```js
assert.isFunction(booWho);
```

`booWho(true)` 应该返回 `true`。

```js
assert.isTrue(booWho(true));
```

`booWho(false)` 应该返回 `true`。

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` 应该返回 `false`。

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` 应该返回 `false`。

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` 应该返回 `false`。

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` 应该返回 `false`。

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` 应该返回 `false`。

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` 应该返回 `false`。

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` 应该返回 `false`。

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` 应该返回 `false`。

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
