---
id: 66f17db06803d11a1bd19a20
title: 构建一个测验游戏
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该创建一个名为 `questions` 的数组。
2. `questions` 数组应该包含至少五个对象，每个对象都有 `category`、`question`、`choices` 和 `answer` 键。
3. `category` 键应该有一个字符串值，表示问题类别。
4. `question` 键应该有一个字符串值，表示一个问题。
5. `choices` 键应该有一个数组值，包含三个字符串，这些是问题的备选答案。
6. `answer` 键应该有一个字符串值，表示问题的正确答案。此外，`answer` 的值应该包含在 `choices` 数组中。
7. 您应该有一个名为 `getRandomQuestion` 的函数，该函数接受一个问题数组作为参数，并从数组中返回一个随机问题对象。
8. 您应该有一个名为 `getRandomComputerChoice` 的函数，该函数接受可用选择的数组作为参数，并返回对所选问题的随机答案。
9. 您应该有一个名为 `getResults` 的函数，该函数接受问题对象作为第一个参数，计算机的选择作为第二个参数。如果答案正确，函数应该返回 `The computer's choice is correct!`。否则，它返回 `The computer's choice is wrong. The correct answer is: <correct-answer>`，其中 `<correct-answer>` 是所选问题的正确答案值。

# --hints--

您应该创建一个名为 `questions` 的数组。

```js
assert.isArray(questions);
```

`questions` 数组应该包含至少五个对象，每个对象都有 `category`、`question`、`choices` 和 `answer` 键。

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

`category` 键应该有一个字符串值，表示问题类别。

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => assert.isString(category));
```

`question` 键应该有一个字符串值，表示一个问题。

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
})
```

`choices` 键应该有一个包含三个字符串的数组值。

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    choices.forEach(choice => assert.isString(choice));
})
```

`answer` 键应该有一个字符串值。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => assert.isString(answer));
```

`answer` 的值应该包含在 `choices` 数组中。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

您应该有一个名为 `getRandomQuestion` 的函数，该函数接受一个问题数组作为参数，并从数组中返回一个随机问题对象。

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

您应该有一个名为 `getRandomComputerChoice` 的函数，该函数接受可用选择的数组作为参数，并返回对所选问题的随机答案。

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

您应该有一个名为 `getResults` 的函数。

```js
assert.isFunction(getResults);
```

您的 `getResults` 函数应该接受问题对象作为第一个参数，计算机的选择作为第二个参数。

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
assert.notEqual(getResults(questions[0].answer, questions[0]), `The computer's choice is correct!`);
```

如果计算机选择与答案匹配，`getResults` 应该返回 `The computer's choice is correct!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

如果计算机选择与答案不匹配，`getResults` 应该返回 `The computer's choice is wrong. The correct answer is: <correct-answer>`，其中 `<correct-answer>` 是所选问题的正确答案值。

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);
```