---
id: 619665c9abd72906f3ad30f9
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你将要创建一个快乐的 Flappy 企鹅，并在此过程中进一步练习 CSS 变换和动画。

首先将你的样式表链接到页面上。

# --hints--

你的代码应该有一个 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/619665c9abd72906f3ad30f9.md#L52-L56) 元素。

```js
assert.exists(document.querySelector('link'));
```

你的 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/619665c9abd72906f3ad30f9.md#L52-L56) 元素应该在你的 [head](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/619665c9abd72906f3ad30f9.md#L52-L56) 元素内。

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

你的 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/619665c9abd72906f3ad30f9.md#L52-L56) 元素应该有一个 [rel](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61a8fe15a6a31306e60d1e89.md#L64-L71) 属性，值为 [stylesheet](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/619665c9abd72906f3ad30f9.md#L52-L56)。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 [link](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/619665c9abd72906f3ad30f9.md#L52-L56) 元素应该有一个 [href](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61a8fe15a6a31306e60d1e89.md#L64-L71) 属性，值为 [styles.css](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/619665c9abd72906f3ad30f9.md#L52-L56)。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```