---
id: 61969aa6acef5b12200f672e
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

为了在背景中创建一些风景，你将添加两座山。

在 `.penguin` 元素上方，添加一个 `class` 为 `left-mountain` 的 `div`。

# --hints--

你应该在 `body` 中添加一个新的 `div`。期望看到3个 `div` 元素，但找到了--fcc-actual--个。

```js
assert.lengthOf(document.querySelectorAll('body > div'), 3);
```

你应该给 `div` 设置 `class` 为 `left-mountain`。

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

你应该将 `left-mountain` 放置在 `.penguin` 之前。

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```