---
id: 680b3ef395479b0e449ecb6e
title: 实现选择排序算法
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

选择排序是大多数计算机科学课程中教授的另一种流行的排序算法。

该算法通过重复查找列表未排序部分中的最小元素并将其与第一个未排序元素交换来工作。它首先选择整个列表中的最小值并将其与第一个元素交换。然后它移动到第二个位置，查找剩余未排序元素中的最小值，并将其与第二个元素交换。这个过程继续进行，一次移动一个元素，直到整个列表排序完成。

选择排序在最好、平均和最坏情况下都会产生二次时间复杂度。空间复杂度将是常数 `O(1)`，因为排序是就地进行的，无论列表大小如何，使用的内存量都是恒定的。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该定义一个名为 [selection_sort](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 的函数。
1. 你的 [selection_sort](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 函数应该有一个参数，代表项目列表。
1. 你的 [selection_sort](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 函数应该接受一个列表并将项目从小到大就地排序。
1. 你的 [selection_sort](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 函数不应使用内置的 [sort()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L23-L23) 方法或 [sorted()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 函数。

# --hints--

你应该有一个名为 [selection_sort](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 的函数。

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

你的 [selection_sort](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 函数应该有一个参数。

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

你不应在代码中使用内置的 [sort()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.d.ts#L23-L23) 方法或 [sorted()](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 函数。

```js
(
    {
        test: () => runPython(`
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

[selection_sort([33, 1, 89, 2, 67, 245])](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 应返回 [1, 2, 33, 89, 67, 245](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141)。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]

        `)
    }
)
```

[selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 应返回 [3, 5, 12, 15, 16, 23, 72, 99, 567](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141)。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]

        `)
    }
)
```

[selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 应返回 [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141)。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]

        `)
    }
)
```

你的 [selection_sort](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-selection-sort/680b3ef395479b0e449ecb6e.md#L130-L141) 函数应该能正确排序任何数字列表。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]

        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
