---
id: 66fe4f33a2cc9b33f4d5cd9b
title: 构建一个排列生成器
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

在这个实验中，您将构建一个排列生成器，它将接受一个字符串并返回字符串中字符的所有可能排列。例如，字符串 `cat` 的可能排列是 `cat`、`cta`、`act`、`atc`、`tac` 和 `tca`。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该创建一个名为 `permuteString` 的函数。
2. `permuteString` 函数应该接受三个参数 - 一个字符串、一个前缀值和一个用于存储和返回结果的空数组。前缀值将累积字符以形成排列。
3. 在函数内部，您应该检查传入字符串的长度是否为 `0`。如果是，则将当前前缀推入结果中并返回结果。
4. 遍历输入字符串中的每个字符，对于每次迭代，从字符串中移除当前字符，并递归调用 `permuteString` 函数，使用更新后的参数来构建剩余的排列。
5. 您应该返回最终的结果数组。
6. 您应该通过删除重复项来确保排列是唯一的。

# --hints--

您应该有一个 `permuteString` 函数。

```js
assert.isFunction(permuteString);
```

您应该在 `permuteString` 函数中使用递归。

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` 应该返回 `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`。

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` 应该返回 `[ "fcc", "cfc", "ccf" ]`。

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` 应该返回 `[ "p" ]`。

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` 应该返回 `[""]`。

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);
```

`permuteString("walk")` 应该返回 `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`。

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` 应该返回正确的结果。

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```