# --description--

设置 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61967e74a8e3690ab6292daa.md#L47-L49) 元素的 [background](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/6196d0cda039d026f7f78d1e.md#L45-L52) 为一个线性渐变，角度为顺时针45度，从 [rgb(118, 201, 255)](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61967e74a8e3690ab6292daa.md#L53-L55) 开始，到 [rgb(247, 255, 222)](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61967e74a8e3690ab6292daa.md#L53-L55) 结束。

# --hints--

你应该使用 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61967e74a8e3690ab6292daa.md#L47-L49) 元素选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你应该在 [body](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61967e74a8e3690ab6292daa.md#L47-L49) 选择器中使用 [background](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/6196d0cda039d026f7f78d1e.md#L45-L52) 属性。

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

你应该将 [background](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/6196d0cda039d026f7f78d1e.md#L45-L52) 设置为 [linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))](file:///F:/code/blocks_Chinese/blocks_english/workshop-flappy-penguin/61967e74a8e3690ab6292daa.md#L53-L55)。

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
}
--fcc-editable-region--
```