---
id: 619699c10a0f6e11591d73c4
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

使用[margin]属性水平居中[.penguin]元素，并将[margin-top]设置为[75px]。

# --hints--

你应该给[.penguin]设置[margin]为[auto]。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.marginRight, 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.marginLeft, 'auto');
```

你应该给[.penguin]设置[margin-top]为[75px]。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.marginTop, '75px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;

}
.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}