---
id: 61a8fe15a6a31306e60d1e89
title: 第 4 步
challengeType: 0
dashedName: step-4
---

# --description--

通过将 `width` 设置为 `100%`，`height` 设置为 `100vh` 来标准化你的页面。

# --hints--

你应该给 `body` 一个 `--fcc-expected--` 的 `width`，但是找到了 `--fcc-actual--`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

你应该给 `body` 一个 `--fcc-expected--` 的 `height`，但是找到了 `--fcc-actual--`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;

}
--fcc-editable-region--
```