---
id: 6858eadd3144c89d45e02bc7
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

如果你注意到的话，外部的 `div` 太靠近顶部，并且其中的内容太靠近左侧。所以，你现在应该努力留出一些空间。

将类 `mt-8` 添加到 `div` 以将其向下推一点，并添加 `p-4` 以将文本从左侧移开。

# --hints--

你的外部 `div` 元素应该是实用类 `mt-8`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

你的外部 `div` 元素应该是实用类 `p-4`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
       <div>
        <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```