---
id: 6858ebf2552046ab3de9f17c
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

此时，你会注意到 `div` 中内容右侧的空间太多。你可以通过增加宽度来最小化这种情况。

将类 `md:w-1/2` 添加到 `div` 以将中型和大型屏幕上的宽度减小到 50%。也就是说，屏幕分辨率为 768 像素及以上。

要查看效果，请展开预览选项卡。

# --hints--

你的外部 `div` 元素应该具有实用类 `md:w-1/2`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
       <div>
        <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```