---
id: 6858ee73ed8106ccb214942a
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

现在所有内容都紧贴左侧，你需要通过在x轴上添加自动边距来居中整个内容。

将相应的Tailwind类添加到外部 `div` 元素。要看到效果，请调整预览选项卡，直到整个内容居中。

# --hints--

你的外部 `div` 元素应该具有实用类 `mx-auto`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
       <div>
        <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```