---
id: 685973ca82469e940afcc422
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

为了完成按钮样式，为每个 `a` 元素添加实用类 `font-semibold`，使它们的文本更粗。同时，添加类 `rounded` 来移除按钮的尖锐边缘。

# --hints--

你的第一个锚点元素应该具有实用类 `font-semibold`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("font-semibold"));
```

你的第一个锚点元素应该具有实用类 `rounded`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("rounded"));
```

你的第二个锚点元素应该具有实用类 `font-semibold`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("font-semibold"));
```

你的第二个锚点元素应该具有实用类 `rounded`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("rounded"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
      <div>
--fcc-editable-region--
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600 px-4 py-2"
        >
          了解更多
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white px-4 py-2"
         >
          开始聆听
        </a>
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```