---
id: 685972d14c6d8182fbae71f1
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

为了让锚点标签看起来更像按钮，为每个锚点添加类 `px-4` 和 `py-2`。`px-4` 将增加左右内边距 `1rem`，`py-2` 将增加上下内边距 `0.5rem`。

# --hints--

你的第一个锚点元素应该具有实用类 `px-4`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("px-4"));
```

你的第一个锚点元素应该具有实用类 `py-2`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("py-2"));
```

你的第二个锚点元素应该具有实用类 `px-4`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("px-4"));
```

你的第二个锚点元素应该具有实用类 `py-2`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("py-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
      <div>
        --fcc-editable-region--
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600"
        >
          了解更多
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white"
         >
          开始聆听
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```