---
id: 68591755e8cd86d27699d282
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

最后，要移除外层 `div` 元素的尖锐边缘，添加类 `rounded-md`。

# --hints--

你的外层 `div` 元素应该具有实用类 `rounded-md`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
      <div>
        <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```