---
id: 646ce9d790d2a44de5f99e04
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

创建两个 `div` 元素，第一个在 `.cat-left-ear` 元素内，类为 `cat-left-inner-ear`，第二个在 `.cat-right-ear` 元素内，类为 `cat-right-inner-ear`。

# --hints--

你不应该更改现有的带有 `cat-left-ear` 类的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div.cat-left-ear'), 1);
```

你不应该更改现有的带有 `cat-right-ear` 类的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('div.cat-right-ear'), 1);
```

你应该在 `.cat-left-ear` 元素内有一个 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.cat-left-ear div'), 1);
```

你应该在 `.cat-right-ear` 元素内有一个 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.cat-right-ear div'), 1);
```

在 `.cat-left-ear` 内的新 `div` 元素应该有 `cat-left-inner-ear` 类。

```js
assert.isTrue(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

在 `.cat-right-ear` 内的新 `div` 元素应该有 `cat-right-inner-ear` 类。

```js
assert.isTrue(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```