---
id: 6476f7a4827bcc61682f2347
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

当你为`position`属性使用`absolute`值时，元素会脱离文档的正常流，然后通过`top`、`right`、`bottom`和`left`属性来确定其位置。

将`.cat-head`元素的`position`属性设置为`absolute`，然后将`top`和`left`属性设置为任意正像素值。

<!-- **注意**: 你可以在这里试验`top`、`left`、`bottom`和`right`属性，但测试只会在`top`为`300px`和`left`为`400px`时通过。 -->

# --hints--

你的`.cat-head`选择器应该有一个`position`属性设置为`absolute`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

你的`.cat-head`选择器应该有一个`top`属性设置为任意像素值。确保你添加了分号。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

你的`.cat-head`选择器应该有一个`left`属性设置为任意像素值。确保你添加了分号。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```