---
id: 6477062778c85972eb648030
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

你现在应该居中猫头。

给`.cat-head`元素一个`position`属性设置为`absolute`。将`right`、`left`、`top`、`bottom`属性的值都设置为`0`，然后将其所有边的`margin`属性设置为`auto`。这是使用CSS定位垂直和水平居中元素的一种方法。

# --hints--

你的`.cat-head`选择器应该有一个`position`属性设置为`absolute`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

你的`.cat-head`选择器应该有一个`top`属性设置为`0`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

你的`.cat-head`选择器应该有一个`left`属性设置为`0`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '0px');
```

你的`.cat-head`选择器应该有一个`right`属性设置为`0`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right, '0px');
```

你的`.cat-head`选择器应该有一个`bottom`属性设置为`0`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom, '0px');
```

你的`.cat-head`选择器应该有一个`margin`属性设置为`auto`。确保你添加了分号。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```