---
id: 6676a8b01e56ec1a1e07c202
title: 步骤21
challengeType: 0
dashedName: step-21
---

# --description--

注意你现在有一个厚边框的白色方块。边框有对角线，可用于创建三角形。你需要用它们来制作猫耳朵。

在同一个类选择器中，将[height]和[width]调整为[0]。将左侧、右侧和顶部边框设置为[transparent]。

删除[background-color]属性，你应该看到一个三角形。

# --hints--

你应该有[.cat-right-ear]选择器。

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

你的[.cat-right-ear]选择器应将[height]属性设置为[0]。不要忘记加上分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '0px')
```

你的[.cat-right-ear]选择器应将[width]属性设置为[0]。不要忘记加上分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '0px')
```

你的[.cat-right-ear]选择器应将[border-left]属性设置为[35px solid transparent]。不要忘记加上分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

你的[.cat-right-ear]选择器应将[border-right]属性为[35px solid transparent]。不要忘记加上分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

你的[.cat-right-ear]选择器应将[border-top]属性设置为[35px solid transparent]。不要忘记加上分号。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid transparent')
```

你的[.cat-right-ear]选择器应该没有[background-color]属性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
        
        <div class="cat-eyes">
          <div class="cat-left-eye">
            <div class="cat-left-inner-eye"></div>
          </div>
          <div class="cat-right-eye">
            <div class="cat-right-inner-eye"></div>
          </div>
        </div>
        
        <div class="cat-nose"></div>
        
        <div class="cat-mouth">
          <div class="cat-mouth-line-left"></div>
          <div class="cat-mouth-line-right"></div>
        </div>
        
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
  transform: rotate(180deg);
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
  transform: rotate(180deg);
}

.cat-left-eye {
  position: absolute;
  top: 54px;
  left: 39px;
  border-radius: 60%;
  transform: rotate(25deg);
  width: 30px;
  height: 40px;
  background-color: #000;
}

.cat-right-eye {
  position: absolute;
  top: 54px;
  left: 134px;
  border-radius: 60%;
  transform: rotate(-25deg);
  width: 30px;
  height: 40px;
  background-color: #000;
}

.cat-left-inner-eye {
  position: absolute;
  top: 8px;
  left: 2px;
  width: 10px;
  height: 20px;
  transform: rotate(10deg);
  background-color: #fff;
  border-radius: 60%;
}

.cat-right-inner-eye {
  position: absolute;
  top: 8px;
  left: 18px;
  width: 10px;
  height: 20px;
  transform: rotate(-5deg);
  background-color: #fff;
  border-radius: 60%;
}

.cat-nose {
  position: absolute;
  top: 108px;
  left: 85px;
  border-top-left-radius: 50%;
  border-bottom-right-radius: 50%;
  border-bottom-left-radius: 50%;
  transform: rotate(180deg);
  border-left: 15px solid transparent;
  border-right: 15px solid transparent;
  border-bottom: 20px solid #442c2c;
}

.cat-mouth div {
  width: 30px;
  height: 51px;
  border: 2px solid #000;
  border-radius: 190px 190px 0 0;
  border-color: black transparent transparent transparent;
}

.cat-mouth-line-left {
  position: absolute;
  top: 88px;
  left: 74px;
  transform: rotate(170deg);
}

.cat-mouth-line-right {
  position: absolute;
  top: 88px;
  left: 91px;
  transform: rotate(165deg);
}
```