---
id: a202eed8fc186c8434cb6d61
title: 创建一个字符串反转器
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

在本实验室中，你将创建一个简单的字符串反转器，它可以反转给定字符串的字符。

例如， `"hello"` 应该变成 `"olleh"`。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应创建一个名为 `reverseString` 的函数，该函数以字符串为参数。
2. 函数应返回反转的字符串。

# --hints--

你应该有一个名为 `reverseString` 的函数。

```js
assert.isFunction(reverseString);
```

`reverseString` 应使用字符串作为参数。

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

[reverseString("hello")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-reverse-a-string/a202eed8fc186c8434cb6d61.md#L73-L75) 应返回一个字符串。

```js
assert.isString(reverseString('hello'));
```

[reverseString("hello")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-reverse-a-string/a202eed8fc186c8434cb6d61.md#L73-L75) 应返回 [olleh](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-reverse-a-string/a202eed8fc186c8434cb6d61.md#L73-L75)。

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

[reverseString("Howdy")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-reverse-a-string/a202eed8fc186c8434cb6d61.md#L73-L75) 应返回 [ydwoH](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-reverse-a-string/a202eed8fc186c8434cb6d61.md#L73-L75)。

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

[reverseString("Greetings from Earth")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-reverse-a-string/a202eed8fc186c8434cb6d61.md#L73-L75) 应返回 [htraE morf sgniteerG](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-reverse-a-string/a202eed8fc186c8434cb6d61.md#L73-L75)。

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
