---
id: 5dc24165f86c76b9248c6ebc
title: æ­¥éª¤ 9
challengeType: 0
dashedName: step-9
---

# --description--

æ‰€æœ‰çš„ `img` å…ƒç´ éƒ½åº”è¯¥æœ‰ä¸€ä¸?`alt` å±æ€§ã€?`alt` å±æ€§çš„æ–‡æœ¬ç”¨äºå±å¹•é˜…è¯»å™¨ä»¥æ”¹è¿›å¯è®¿é—®æ€§ï¼Œä»¥åŠåœ¨å›¾åƒåŠ è½½å¤±è´¥æ—¶æ˜¾ç¤ºã€?

è¿™æ˜¯ä¸€ä¸ªå¸¦æœ?`alt` å±æ€§çš„ `img` å…ƒç´ çš„ä¾‹å­ï¼š

```html
<img src="cat.jpg" alt="A cat">
```

åœ?`img` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`alt` å±æ€§ï¼Œæ–‡æœ¬ä¸ºï¼š

`A cute orange cat lying on its back`

# --hints--

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`img` å…ƒç´ ã€?ä½ åˆ é™¤äº†å‰ä¸€æ­¥éª¤ä¸­çš„ `img` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('img'));
```

ä½ çš„ `img` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`alt` å±æ€§ã€?æ£€æŸ¥å¼€å§‹æ ‡ç­¾çš„åç§°åé¢æ˜¯å¦æœ‰ç©ºæ ¼å’Œ/æˆ–æ‰€æœ‰å±æ€§åç§°ä¹‹å‰æ˜¯å¦æœ‰ç©ºæ ¼ã€?

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

ä½ çš„ `img` å…ƒç´ çš?`alt` å±æ€§å€¼è¢«è®¾ç½®æˆäº†å…¶ä»–å†…å®¹ï¼Œè€Œé `A cute orange cat lying on its back`ã€?ç¡®ä¿å°?`alt` å±æ€§çš„å€¼æ”¾åœ¨å¼•å·ä¸­ã€?

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

