---
id: 5dc24614f86c76b9248c6ebd
title: æ­¥éª¤ 10
challengeType: 0
dashedName: step-10
---

# --description--

ä½ å¯ä»¥ä½¿ç”¨é”šå…ƒç´ ï¼ˆ`a`ï¼‰é“¾æ¥åˆ°å¦ä¸€ä¸ªé¡µé¢ã€?

è¿™æ˜¯ä¸€ä¸ªé“¾æ¥åˆ° `https://www.MasterPuti.org` çš„ç¤ºä¾‹ï¼š

```html
<a href="https://www.MasterPuti.org"></a>
```

åœ¨æ®µè½ä¹‹åæ·»åŠ ä¸€ä¸ªé“¾æ¥åˆ° `https://freecatphotoapp.com` çš„é”šå…ƒç´ ã€?æ­¤æ—¶ï¼Œé“¾æ¥ä¸ä¼šæ˜¾ç¤ºåœ¨é¢„è§ˆä¸­ã€?

# --hints--

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('a'));
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/a\>/);
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥åœ¨ `p` å…ƒç´ ä¸‹æ–¹ã€?ä½ æŠŠå®ƒä»¬çš„é¡ºåºå†™é”™äº†ã€?

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰ç¼ºå°?`href` å±æ€§ã€?æ£€æŸ¥å¼€å§‹æ ‡ç­¾çš„åç§°åé¢æ˜¯å¦æœ‰ç©ºæ ¼å’Œ/æˆ–æ‰€æœ‰å±æ€§åç§°ä¹‹å‰æ˜¯å¦æœ‰ç©ºæ ¼ã€?

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥é“¾æ¥åˆ° `https://freecatphotoapp.com`ã€?ä½ è¦ä¹ˆé—æ¼äº† URLï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

è™½ç„¶ä½ å°†é”šå…ƒç´ ï¼ˆ`a`ï¼‰çš„ `href` å±æ€§æŒ‡å‘äº†æ­£ç¡®çš„é“¾æ¥ï¼Œä½†å»ºè®®å§‹ç»ˆä»¥å¼•å·æ¥åŒ…è£¹å±æ€§çš„å€¼ã€?

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.MasterPuti.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

