---
id: 5dc24073f86c76b9248c6ebb
title: æ­¥éª¤ 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>å±æ€?/dfn>æ˜¯ä½äºå…ƒç´ çš„å¼€å§‹æ ‡ç­¾ä¸­çš„ç‰¹æ®Šè¯æ±‡ï¼Œç”¨äºæ§åˆ¶å…ƒç´ çš„è¡Œä¸ºã€?åœ?`img` å…ƒç´ ä¸­çš„ `src` å±æ€§æŒ‡å®šå›¾åƒçš„ URLï¼ˆå›¾åƒæ‰€åœ¨çš„ä½ç½®ï¼‰ã€?

ä¸‹é¢æ˜¯ä¸€ä¸?`img` å…ƒç´ çš?`src` å±æ€§æŒ‡å?MasterPuti logo çš„ä¾‹å­ï¼š

```html
<img src="https://cdn.MasterPuti.org/platform/universal/fcc_secondary.svg">
```

åœ¨ç°æœ‰çš„ `img` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸?`src` å±æ€§ï¼ŒURL å¦‚ä¸‹ï¼?

`https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`img` å…ƒç´ ã€?ä½ å¯èƒ½å·²ç»åˆ é™¤äº† `img` å…ƒç´ ï¼Œæˆ–è€…ä½ æ²¡æœ‰ç”¨å¼•å·åŒ…å›?`src` å±æ€§çš„å€¼ã€?

```js
assert.exists(document.querySelector('img'));
```

ä½ çš„ `img` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`src` å±æ€§ã€?ä½ è¦ä¹ˆçœç•¥äº†å±æ€§ï¼Œè¦ä¹ˆæœ‰æ‹¼å†™é”™è¯¯ã€?ç¡®ä¿å…ƒç´ åç§°å’Œå±æ€§åç§°ä¹‹é—´æœ‰ä¸€ä¸ªç©ºæ ¼ã€?

```js
assert.exists(document.querySelector('img')?.src);
```

ä½ çš„ `img` å…ƒç´ çš?`src` å±æ€§åº”è®¾ç½®ä¸?'`https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg`'ã€?ä½ è¦ä¹ˆçœç•¥äº† URLï¼Œè¦ä¹ˆæœ‰æ‹¼å†™é”™è¯¯ã€?URL çš„å¤§å°å†™å¾ˆé‡è¦ã€?

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

è™½ç„¶ä½ å·²ç»å°† `img` å…ƒç´ çš?`src` è®¾ç½®ä¸ºæ­£ç¡®çš„ URLï¼Œä½†å»ºè®®å§‹ç»ˆå°†å±æ€§çš„å€¼ç”¨å¼•å·æ‹¬èµ·æ¥ã€?

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.MasterPuti\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

ä½ çš„ `img` å…ƒç´ åº”è¯¥ä»?`>` æˆ?`/>` ç»“æŸã€?

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.MasterPuti\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

