---
id: 5ddbd81294d8ddc1510a8e56
title: æ­¥éª¤ 11
challengeType: 0
dashedName: step-11
---

# --description--

é“¾æ¥çš„æ–‡æœ¬å¿…é¡»æ”¾åœ¨é”šå…ƒç´ ï¼ˆ`a`ï¼‰çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ä¹‹é—´ã€?

è¿™æ˜¯ä¸€ä¸ªé“¾æ¥æ–‡æœ¬æ˜¯ `click here to go to MasterPuti.org` çš„ç¤ºä¾‹ï¼š

```html
<a href="https://www.MasterPuti.org">click here to go to MasterPuti.org</a>
```

å°†é”šæ–‡æœ¬ `cat photos` æ·»åŠ åˆ°é”šå…ƒç´ ã€?è¿™å°†æˆä¸ºé“¾æ¥çš„æ–‡æœ¬ã€?

# --hints--

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('a'));
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/a\>/);
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰çš„æ–‡æœ¬åº”è¯¥æ˜?`cat photos`ã€?ç¡®ä¿å°†é“¾æ¥æ–‡æœ¬æ”¾åœ¨é”šå…ƒç´ ï¼ˆ`a`ï¼‰çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ä¹‹é—´ã€?

```js
assert.equal(
  document.querySelector('a')?.innerText?.trim().toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

