---
id: 5dfa3589eacea3f48c6300ae
title: æ­¥éª¤ 20
challengeType: 0
dashedName: step-20
---

# --description--

åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ ä¸­ï¼Œæ·»åŠ ä¸€ä¸ªå¸¦æœ‰æ–‡æœ?`Cat Lists` çš„æ–° `h2` å…ƒç´ ã€?

# --hints--

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

ä½ çš„ `h2` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

ä½ çš„ `h2` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

ä½ çš„ç¬¬äºŒä¸?`h2` å…ƒç´ åº”è¯¥åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾çš„ä¸Šæ–¹ã€?å®ƒçš„ä½ç½®ä¸å¯¹ã€?

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

ç¬¬äºŒä¸?`h2` å…ƒç´ åº”è¯¥æœ‰æ–‡æœ?`Cat Lists`ã€?ä½ è¦ä¹ˆé—æ¼äº†æ–‡æœ¬ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

