---
id: 5dfa2407b521be39a3de7be1
title: æ­¥éª¤ 15
challengeType: 0
dashedName: step-15
---

# --description--

è¦åœ¨æ–°æ ‡ç­¾é¡µä¸­æ‰“å¼€é“¾æ¥ï¼Œä½ å¯ä»¥åœ¨é”šå…ƒç´ ï¼ˆ`a`ï¼‰ä¸Šä½¿ç”¨ `target` å±æ€§ã€?

`target` å±æ€§æŒ‡å®šæ‰“å¼€é“¾æ¥æ–‡æ¡£çš„ä½ç½®ã€?`target="_blank"` æ˜¯åœ¨æ–°æ ‡ç­¾é¡µæˆ–çª—å£ä¸­æ‰“å¼€é“¾æ¥çš„æ–‡æ¡£ã€?

è¿™æ˜¯ä¸€ä¸ªå¸¦æœ?`target` å±æ€§çš„ `a` å…ƒç´ çš„åŸºæœ¬è¯­æ³•ï¼š

```html
<a href="https://www.MasterPuti.org" target="_blank">MasterPuti</a>
```

å‘é”šå…ƒç´ ï¼ˆ`a`ï¼‰çš„å¼€å§‹æ ‡ç­¾æ·»åŠ ä¸€ä¸ªå€¼ä¸º `_blank` çš?`target` å±æ€§ï¼Œç”¨ä»¥åœ¨æ–°çš„æ ‡ç­¾é¡µä¸­æ‰“å¼€é“¾æ¥ã€?

# --hints--

ä½ çš„ `p` å…ƒç´ åº”è¯¥åµŒå¥—ä¸€ä¸ªæ–‡æœ¬ä¸º `cat photos` çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰ã€?ä½ å¯èƒ½åˆ é™¤äº†å®ƒæˆ–è€…æ‹¼å†™æœ‰è¯¯ã€?

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰æ²¡æœ?`target` å±æ€§ã€?æ£€æŸ¥å¼€å§‹æ ‡ç­¾çš„åç§°åé¢æ˜¯å¦æœ‰ç©ºæ ¼å’Œ/æˆ–æ‰€æœ‰å±æ€§åç§°ä¹‹å‰æ˜¯å¦æœ‰ç©ºæ ¼ã€?

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

`target` å±æ€§çš„å€¼åº”ä¸?`_blank`ã€?ä½ è¦ä¹ˆé—æ¼äº†å±æ€§å€¼ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?è®°ä½ï¼Œå±æ€§å€¼åº”è¯¥æ”¾åœ¨å¼•å·ä¸­ã€?

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

