---
id: 5dfa37b9eacea3f48c6300b0
title: æ­¥éª¤ 22
challengeType: 0
dashedName: step-22
---

# --description--

è¦åˆ›å»ºä¸€ä¸ªæ— åºé¡¹ç›®åˆ—è¡¨ï¼Œä½ å¯ä»¥ä½¿ç”?`ul` å…ƒç´ ã€?

åœ¨å¸¦æœ?`Things cats love:` æ–‡æœ¬çš?`h3` å…ƒç´ ä¹‹åï¼Œæ·»åŠ ä¸€ä¸ªæ— åºåˆ—è¡¨ï¼ˆ`ul`ï¼‰å…ƒç´ ã€?æ³¨æ„ï¼Œæ­¤æ—¶ä¸ä¼šæ˜¾ç¤ºä»»ä½•å†…å®¹ã€?

# --hints--

ä½ çš„ `ul` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('ul'));
```

ä½ çš„ `ul` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/ul>/);
```

`ul` å…ƒç´ åº”ä½äºç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¹‹ä¸Šã€?

```js
const secondSectionLastElemNode =
  document.querySelectorAll('main > section')[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

