---
id: 5dfa371beacea3f48c6300af
title: æ­¥éª¤ 21
challengeType: 0
dashedName: step-21
---

# --description--

å½“ä½ å‘é¡µé¢æ·»åŠ è¾ƒä½çº§åˆ«æ ‡é¢˜å…ƒç´ æ—¶ï¼Œè¿™æ„å‘³ç€ä½ æ­£åœ¨å¼€å§‹ä¸€ä¸ªæ–°çš„å­éƒ¨åˆ†ã€?

åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ çš„æœ€åä¸€ä¸?`h2` å…ƒç´ ä¹‹åï¼Œæ·»åŠ ä¸€ä¸?`h3` å…ƒç´ ï¼Œæ–‡æœ¬ä¸ºï¼?

`Things cats love:`

# --hints--

ç¬¬äºŒä¸?`section` å…ƒç´ ç¼ºå¤±æˆ–æ²¡æœ‰å¼€å§‹å’Œç»“æŸæ ‡ç­¾ã€?

```js
assert.exists(document.querySelectorAll('main > section')[1]);
assert.lengthOf(code.match(/\<\/section>/g), 2);
```

åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¸Šæ–¹åº”è¯¥æœ‰ä¸€ä¸?`h3` å…ƒç´ ã€?

```js
assert.equal(
  document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName,
    'H3'
);
```

ä½ çš„ `h3` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.lengthOf(code.match(/<\/h3>/g), 1);
```

ç¬¬äºŒä¸?`section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¸Šæ–¹çš„ `h3` å…ƒç´ åº”è¯¥åŒ…å«æ–‡æœ¬ `Things cats love:`ã€?ç¡®ä¿åœ¨æ–‡æœ¬æœ«å°¾åŒ…å«å†’å·ã€?

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild.innerText?.trim().toLowerCase()
    .replace(/\s+/g, ' '), 'things cats love:'
);
```

åµŒå¥—åœ¨æœ€åä¸€ä¸?`section` å…ƒç´ ä¸­çš„æœ€åä¸€ä¸?`h3` å…ƒç´ ä¹‹ä¸Šï¼Œåº”è¯¥æœ‰ä¸€ä¸?`h2` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`Cat Lists`ã€?ä½ å¯èƒ½ä¸å°å¿ƒåˆ é™¤äº?`h2` å…ƒç´ ã€?

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  ?.lastElementChild;
assert.equal( secondSectionLastElemNode?.nodeName, 'H3');
assert.equal(
 secondSectionLastElemNode?.previousElementSibling.innerText?.trim()
      .toLowerCase()
      .replace(/\s+/g, ' '), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

