---
id: 5dfa30b9eacea3f48c6300ad
title: æ­¥éª¤ 17
challengeType: 0
dashedName: step-17
---

# --description--

åœ¨ä¹‹å‰çš„æ­¥éª¤ä¸­ï¼Œä½ ä½¿ç”¨é”šå…ƒç´ å°†æ–‡æœ¬è½¬æ¢ä¸ºé“¾æ¥ã€?å…¶ä»–ç±»å‹çš„å†…å®¹ä¹Ÿå¯ä»¥é€šè¿‡å°†å…¶åŒ…è£¹åœ¨é”šæ ‡ç­¾ä¸­æ¥è½¬æ¢ä¸ºé“¾æ¥ã€?

è¿™æ˜¯ä¸€ä¸ªå°†å›¾åƒè½¬æ¢ä¸ºé“¾æ¥çš„ç¤ºä¾‹ï¼?

```html
<a href="example-link">
  <img src="image-link.jpg" alt="A photo of a cat.">
</a>
```

é€šè¿‡ä½¿ç”¨å¿…è¦çš„å…ƒç´ æ ‡ç­¾å°†å›¾åƒè½¬æ¢ä¸ºé“¾æ¥ã€?ä½¿ç”¨ `https://freecatphotoapp.com` ä½œä¸ºé”šå…ƒç´ çš„ `href` å±æ€§çš„å€¼ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`img` å…ƒç´ ï¼Œå…¶ `src` å€¼ä¸º `https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg`ã€?ä½ å¯èƒ½ä¸å°å¿ƒåˆ é™¤äº†å®ƒã€?

```js
assert.equal(
  document.querySelector('img')?.getAttribute('src'),
      'https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg'
);
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.isAtLeast(document.querySelectorAll('a').length, 3);
```

ä½ åœ¨å›¾åƒä¹‹åç¼ºå°‘ä¸€ä¸ªç»“æŸï¼ˆ`a`ï¼‰æ ‡ç­¾ã€?

```js
assert.lengthOf(document.querySelectorAll('a'), 3);
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.isAtLeast(code.match(/<\/a>/g)?.length, 3);
```

ä½ åº”è¯¥åªæ·»åŠ ä¸€ä¸ªç»“æŸé”šï¼ˆ`a`ï¼‰æ ‡ç­¾ã€?è¯·åˆ é™¤ä»»ä½•å¤šä½™çš„ã€?

```js
assert.lengthOf(code.match(/<\/a>/g), 3);
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰ç¼ºå°?`href` å±æ€§ã€?æ£€æŸ¥å¼€å§‹æ ‡ç­¾çš„åç§°åé¢æ˜¯å¦æœ‰ç©ºæ ¼å’Œ/æˆ–æ‰€æœ‰å±æ€§åç§°ä¹‹å‰æ˜¯å¦æœ‰ç©ºæ ¼ã€?

```js
assert.isTrue(document.querySelectorAll('a')[2]?.hasAttribute('href'));
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥é“¾æ¥åˆ° `https://freecatphotoapp.com`ã€?ä½ è¦ä¹ˆé—æ¼äº† URLï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.equal(
  document.querySelectorAll('a')[2]?.getAttribute('href').trim(),
    'https://freecatphotoapp.com'
);
```

ä½ çš„ `img` å…ƒç´ åº”è¯¥è¢«åµŒå¥—åœ¨é”šå…ƒç´ ï¼ˆ`a`ï¼‰ä¹‹ä¸­ã€?æ•´ä¸ª `img` å…ƒç´ åº”è¯¥ä½äºé”šå…ƒç´ ï¼ˆ`a`ï¼‰çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾å†…ã€?

```js
assert.equal(document.querySelector('img')?.parentNode.nodeName, 'A');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
--fcc-editable-region--
    </main>
  </body>
</html>
```

