---
id: 5dfb6a35eacea3f48c6300b4
title: æ­¥éª¤ 26
challengeType: 0
dashedName: step-26
---

# --description--

å›¾é¢˜ï¼ˆ`figcaption`ï¼‰å…ƒç´ ç”¨äºæ·»åŠ æ ‡é¢˜æ¥æè¿° `figure` å…ƒç´ ä¸­åŒ…å«çš„å›¾åƒã€?

è¿™æ˜¯ä¸€ä¸ªæ ‡é¢˜ä¸º `A cute cat` çš?`figcaption` çš„å…ƒç´ ç¤ºä¾‹ï¼š

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

åœ?`figure` å…ƒç´ ä¸­çš„åµŒå¥—çš„å›¾åƒä¹‹åï¼Œæ·»åŠ ä¸€ä¸?`figcaption` å…ƒç´ ï¼Œæ–‡æœ¬ä¸ºï¼?

`Cats love lasagna.`

# --hints--

ä½ çš„ `figcaption` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('figcaption'));
```

ä½ çš„ `figcaption` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/figcaption\>/);
```

`figcaption` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`figure` å…ƒç´ ä¸­ã€?

```js
assert.exists(document.querySelector('figure > figcaption'));
```

åƒå±‚é?`img` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`figure` å…ƒç´ ä¸­ã€?

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

åµŒå¥—åœ?`figure` å…ƒç´ ä¸­çš„ `figcaption` å…ƒç´ åº”ä½äº?`img` å…ƒç´ ä¸‹æ–¹ã€?ä½ æŠŠå®ƒä»¬çš„é¡ºåºå†™é”™äº†ã€?

```js
assert.equal(
  document.querySelector('figcaption')?.previousElementSibling.nodeName, 'IMG'
);
```

ä½ çš„ `figcaption` å…ƒç´ ä¸­çš„æ–‡æœ¬åº”è¯¥æ˜?`Cats love lasagna.`ã€‚ä½ è¦ä¹ˆé—æ¼äº†æ–‡æœ¬ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.match(
  document.querySelector('figcaption')?.innerText?.trim(), /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

