---
id: 5dfb5ecbeacea3f48c6300b1
title: æ­¥éª¤ 23
challengeType: 0
dashedName: step-23
---

# --description--

`li` å…ƒç´ ç”¨äºåœ¨æœ‰åºæˆ–æ— åºåˆ—è¡¨ä¸­åˆ›å»ºåˆ—è¡¨é¡¹ã€?

è¿™æ˜¯ä¸€ä¸ªæ— åºåˆ—è¡¨ä¸­åˆ—è¡¨é¡¹çš„ç¤ºä¾‹ï¼?

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

åœ?`ul` å…ƒç´ ä¸­åµŒå¥—ä¸‰ä¸ªåˆ—è¡¨é¡¹æ¥æ˜¾ç¤ºçŒ«å–œæ¬¢çš„ä¸‰æ ·ä¸œè¥¿ï¼š

`catnip`

`laser pointers`

`lasagna`

# --hints--

ä½ åº”è¯¥æœ‰ä¸‰ä¸ª `li` å…ƒç´ ã€?æ¯ä¸ª `li` å…ƒç´ éƒ½åº”è¯¥æœ‰è‡ªå·±çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ã€?

```js
assert.lengthOf(document.querySelectorAll('li'), 3);
assert.lengthOf(code.match(/<\/li\>/g), 3);
```

ä½ åº”è¯¥æœ‰ä¸‰ä¸ª `li` å…ƒç´ ï¼Œå¯æŒ‰ä»»æ„é¡ºåºåŒ…å«æ–‡æœ?`catnip`ã€`laser pointers` å’?`lasagna`ã€?ä½ è¦ä¹ˆé—æ¼äº†æŸäº›æ–‡æœ¬ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .map((item) => item.innerText?.trim().replace(/\s+/g, ' ').toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['catnip', 'lasagna', 'laser pointers']
);
```

ä¸‰ä¸ª `li` å…ƒç´ åº”ä½äº?`ul` å…ƒç´ çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ä¹‹é—´ã€?

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'UL'
  ), 3
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>

        </ul>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

