---
id: 5dfb6250eacea3f48c6300b2
title: æ­¥éª¤ 24
challengeType: 0
dashedName: step-24
---

# --description--

åœ¨æ— åºåˆ—è¡¨ä¹‹åï¼Œæ·»åŠ ä¸€ä¸ªæ–°å›¾åƒï¼Œå®ƒå…·æœ‰ `src` å±æ€§ï¼Œå€¼ä¸ºï¼?

`https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg`

è€Œå®ƒçš?`alt` å±æ€§å€¼ä¸ºï¼?

`A slice of lasagna on a plate.`

# --hints--

åœ¨ç»“æ?`</ul>` æ ‡ç­¾ä¹‹ååº”è¯¥æœ‰ä¸€ä¸?`img` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

æ–°å›¾åƒæ²¡æœ?`alt` å±æ€§ã€?æ£€æŸ¥å¼€å§‹æ ‡è®°çš„åç§°åé¢æ˜¯å¦æœ‰ç©ºæ ¼å’Œ/æˆ–æ‰€æœ‰å±æ€§åç§°ä¹‹å‰æ˜¯å¦æœ‰ç©ºæ ¼ã€?

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

æ–°å›¾åƒçš„ `alt` å€¼åº”ä¸?`A slice of lasagna on a plate.`ã€‚ç¡®ä¿å°† `alt` å±æ€§çš„å€¼æ”¾åœ¨å¼•å·ä¸­ã€?

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

æ–°å›¾åƒæ²¡æœ?`src` å±æ€§ã€?æ£€æŸ¥å¼€å§‹æ ‡ç­¾çš„åç§°åé¢æ˜¯å¦æœ‰ç©ºæ ¼å’Œ/æˆ–æ‰€æœ‰å±æ€§åç§°ä¹‹å‰æ˜¯å¦æœ‰ç©ºæ ¼ã€?

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

æ–°å›¾åƒçš„ `src` å€¼åº”ä¸?`https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg`ã€?ç¡®ä¿å°?`src` å±æ€§çš„å€¼æ”¾åœ¨å¼•å·ä¸­ã€?

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

è™½ç„¶ä½ å·²å°†æ–°å›¾åƒçš?`src` è®¾ç½®ä¸ºæ­£ç¡®çš„ URLï¼Œä½†å»ºè®®å§‹ç»ˆå°†å±æ€§çš„å€¼ç”¨å¼•å·æ‹¬èµ·æ¥ã€?

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.MasterPuti\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

