---
id: 5dfb655eeacea3f48c6300b3
title: æ­¥éª¤ 25
challengeType: 0
dashedName: step-25
---

# --description--

`figure` å…ƒç´ ä»£è¡¨è‡ªåŒ…å«çš„å†…å®¹ï¼Œå…è®¸ä½ å°†å›¾åƒä¸æ ‡é¢˜ç›¸å…³è”ã€?

å°†åˆšåˆšæ·»åŠ çš„å›¾åƒåµŒå¥—åœ¨ä¸€ä¸?`figure` å…ƒç´ ä¸­ã€?

# --hints--

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('figure'));
```

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/figure\>/);
```

åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¸Šæ–¹åº”è¯¥æœ‰ä¸€ä¸?`figure` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

åƒå±‚é?`img` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`figure` å…ƒç´ ä¸­ã€?

```js
assert.exists(document.querySelector('figure > img'));
assert.equal(
    document.querySelector('figure > img').getAttribute('src').toLowerCase(),
      'https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

