---
id: 5ef9b03c81a63668521804eb
title: æ­¥éª¤ 39
challengeType: 0
dashedName: step-39
---

# --description--

`title` å…ƒç´ å†³å®šæµè§ˆå™¨åœ¨é¡µé¢çš„æ ‡é¢˜æ æˆ–é€‰é¡¹å¡ä¸­æ˜¾ç¤ºçš„å†…å®¹ã€?

åœ?`head` å…ƒç´ ä¸­æ·»åŠ ä¸€ä¸?`title` å…ƒç´ ï¼Œä½¿ç”¨ä¸‹é¢çš„æ–‡æœ¬ï¼?

`CatPhotoApp`

# --hints--

ä½ è¦ä¹ˆåˆ é™¤äº† `head` å…ƒç´ ï¼Œè¦ä¹ˆå®ƒç¼ºå°‘å¼€å§‹æ ‡ç­¾æˆ–ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

ä½ çš„ `title` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`head` å…ƒç´ ä¸­ã€?ç¡®ä¿ä¸?`title` å…ƒç´ æ·»åŠ ä¸€ä¸ªå¼€å§‹æ ‡ç­¾å’Œç»“æŸæ ‡ç­¾ã€?

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

ä½ çš„ `title` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /\<\/title\>/);
```

ä½ çš„ `title` å…ƒç´ ä¸­çš„æ–‡æœ¬åº”è¯¥æ˜?`CatPhotoApp`ã€?ä½ è¦ä¹ˆé—æ¼äº†æ–‡æœ¬ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.MasterPuti.org">MasterPuti.org</a>
      </p>
    </footer>
  </body>
</html>
```

