---
id: 5ef9b03c81a63668521804d2
title: æ­¥éª¤ 29
challengeType: 0
dashedName: step-29
---

# --description--

æœ‰åºåˆ—è¡¨ï¼ˆ`ol`ï¼‰çš„ä»£ç ç±»ä¼¼äºæ— åºåˆ—è¡¨ï¼Œä½†æœ‰åºåˆ—è¡¨ä¸­çš„åˆ—è¡¨é¡¹åœ¨æ˜¾ç¤ºæ—¶æ˜¯ç¼–å·çš„ã€?

åœ?`h3` å…ƒç´ ä¸‹é¢ï¼Œæ·»åŠ ä¸€ä¸ªå¸¦æœ‰è¿™ä¸‰ä¸ªåˆ—è¡¨é¡¹ç›®çš„æœ‰åºåˆ—è¡¨ï¼š

`flea treatment` `thunder` `other cats`

# --hints--

ä½ çš„ `ol` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('ol'));
```

ä½ çš„ `ol` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/ol>/);
```

`ol` å…ƒç´ åº”ä½äºç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¹‹ä¸Šã€?ä½ æŠŠå®ƒä»¬çš„é¡ºåºå†™é”™äº†ã€?

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'OL');
```

ä¸‰ä¸ª `li` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`ol` å…ƒç´ ä¸­ã€?

```js
assert.lengthOf(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ), 3
);
```

ä½ åº”è¯¥æœ‰ä¸‰ä¸ª `li` å…ƒç´ ï¼Œå¯æŒ‰ä»»æ„é¡ºåºåŒ…å«æ–‡æœ?`flea treatment`ã€`thunder` å’?`other cats`ã€?

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText?.trim().replace(/\s+/g, ' ').toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

ä½ åº”è¯¥åªæœ‰ä¸€ä¸?`ol` å…ƒç´ ã€?

```js 
assert.lengthOf([...document.querySelectorAll('ol')], 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

