---
id: 5ef9b03c81a63668521804e7
title: æ­¥éª¤ 35
challengeType: 0
dashedName: step-35
---

# --description--

`footer` å…ƒç´ ç”¨äºå®šä¹‰æ–‡æ¡£æˆ–ç« èŠ‚çš„é¡µè„šã€?é¡µè„šé€šå¸¸åŒ…å«æ–‡æ¡£ä½œè€…ä¿¡æ¯ã€ç‰ˆæƒæ•°æ®ã€ä½¿ç”¨æ¡æ¬¾é“¾æ¥ã€è”ç³»ä¿¡æ¯ç­‰ã€?

åœ?`main` å…ƒç´ ä¹‹åï¼Œæ·»åŠ ä¸€ä¸?`footer` å…ƒç´ ã€?

# --hints--

ä½ è¦ä¹ˆåˆ é™¤äº† `main` å…ƒç´ ï¼Œè¦ä¹ˆå®ƒç¼ºå°‘å¼€å§‹æ ‡ç­¾æˆ–ç»“æŸæ ‡ç­¾ã€?

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

ä½ çš„ `footer` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('footer'));
```

ä½ çš„ `footer` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/footer\>/);
```

ä½ çš„ `footer` å…ƒç´ åº”ä½äº?`main` å…ƒç´ ç»“æŸæ ‡ç­¾çš„ä¸‹æ–¹ã€?ä½ æŠŠå®ƒæ”¾åœ¨åˆ«çš„åœ°æ–¹äº†ã€?

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
--fcc-editable-region--
    </main>

  </body>
--fcc-editable-region--
</html>
```

