---
id: 5ef9b03c81a63668521804d4
title: æ­¥éª¤ 34
challengeType: 0
dashedName: step-34
---

# --description--

`strong` å…ƒç´ ç”¨äºè¡¨ç¤ºæŸäº›æ–‡æœ¬éå¸¸é‡è¦æˆ–ç´§æ€¥ã€?

åœ¨ä½ åˆšåˆšæ·»åŠ çš?`figcaption` ä¸­ï¼Œé€šè¿‡å°?`hate` åŒ…è£¹åœ?`strong` å…ƒç´ ä¸­æ¥è¡¨æ˜å®ƒéå¸¸é‡è¦ã€?

# --hints--

ä½ çš„ `strong` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('strong'));
```

ä½ çš„ `strong` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/strong\>/);
```

ä½ çš„ `strong` å…ƒç´ åº”åœ¨æ–‡æœ¬ `Cats hate other cats.` ä¸­å°† `hate` ä¸€è¯åŒ…è£¹èµ·æ¥ã€‚ä½ è¦ä¹ˆé—æ¼äº†æ–‡æœ¬ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.equal(
  document
    .querySelectorAll('figcaption')[1]
    ?.querySelector('strong')
    ?.innerText?.trim().toLowerCase(), 'hate'
);
```

`figcaption` çš„æ–‡æœ¬åº”è¯¥æ˜¯ `Cats hate other cats.`ï¼Œæ£€æŸ¥æ˜¯å¦æœ‰æ‹¼å†™é”™è¯¯ä»¥åŠåœ?`strong` å…ƒç´ çš„å¼€å¤´å’Œç»“å°¾æ ‡ç­¾å‘¨å›´æ˜¯å¦æœ‰å¿…è¦çš„ç©ºæ ¼ã€?

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert.match(
   secondFigCaption?.innerText
      .replace(/\s+/gi, ' ')
      .trim(),
      /cats hate other cats\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

