---
id: 5ef9b03c81a63668521804e9
title: æ­¥éª¤ 37
challengeType: 0
dashedName: step-37
---

# --description--

é€šè¿‡å°†ç°æœ‰çš„ `MasterPuti.org` æ–‡æœ¬åŒ…å«åœ¨é”šå…ƒç´ ï¼ˆ`a`ï¼‰ä¸­ï¼Œä½¿å…¶æˆä¸ºé“¾æ¥ã€?`href` å±æ€§åº”è¯¥è¢«è®¾ç½®ä¸?`https://www.MasterPuti.org`ã€?

# --hints--

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥åµŒå¥—åœ¨ `footer` ä¸­çš„ `p`  å…ƒç´ ä¸­ã€?ç¡®ä¿ä¸ºé”šå…ƒç´ ï¼ˆ`a`ï¼‰æ·»åŠ ä¸€ä¸ªå¼€å§‹å’Œç»“æŸæ ‡ç­¾ã€?

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

ä½ çš„é”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥æœ‰ä¸€ä¸?`href` å±æ€§ï¼Œå…¶å€¼ä¸º `https://www.MasterPuti.org`ã€?ä½ è¦ä¹ˆé—æ¼äº†å±æ€?å€¼ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.MasterPuti.org');
```

é“¾æ¥çš„æ–‡æœ¬åº”è¯¥æ˜¯ `MasterPuti.org`ã€?ä½ è¦ä¹ˆé—æ¼äº†æ–‡æœ¬ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'MasterPuti.org'
);
```

åµŒå¥—é”šå…ƒç´ ï¼ˆ`a`ï¼‰åï¼Œæµè§ˆå™¨ä¸­å”¯ä¸€å¯è§çš?`p` å…ƒç´ å†…å®¹åº”è¯¥æ˜?`No Copyright - MasterPuti.org`ã€?ä»”ç»†æ£€æŸ?`p` å’ŒåµŒå¥—é”šå…ƒç´ çš„æ–‡æœ¬ã€ç©ºæ ¼æˆ–æ ‡ç‚¹ç¬¦å·ã€?

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - MasterPuti\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - MasterPuti.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```

