---
id: 5ef9b03c81a63668521804ec
title: æ­¥éª¤ 40
challengeType: 0
dashedName: step-40
---

# --description--

æ³¨æ„ï¼Œé¡µé¢çš„å…¨éƒ¨å†…å®¹éƒ½åµŒå¥—åœ¨ `html` å…ƒç´ ä¸­ã€?`html` å…ƒç´ æ˜?HTML é¡µé¢çš„æ ¹å…ƒç´ ï¼ŒåŒ…å«é¡µé¢ä¸Šçš„æ‰€æœ‰å†…å®¹ã€?

ä½ è¿˜å¯ä»¥é€šè¿‡å?`html` å…ƒç´ æ·»åŠ  `lang` å±æ€§æ¥æŒ‡å®šé¡µé¢çš„è¯­è¨€ã€?

å°†å€¼ä¸º `en` çš?`lang` å±æ€§æ·»åŠ åˆ° `html` å¼€å§‹æ ‡ç­¾ï¼ŒæŒ‡å®šé¡µé¢çš„è¯­è¨€ä¸ºè‹±è¯­ã€?

# --hints--

ä½ è¦ä¹ˆåˆ é™¤äº† `html` å…ƒç´ ï¼Œè¦ä¹ˆå®ƒç¼ºå°‘å¼€å§‹æ ‡ç­¾æˆ–ç»“æŸæ ‡ç­¾ã€?

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

ä½ çš„ `html` å…ƒç´ åº”è¯¥æœ‰å€¼ä¸º `en` çš?`lang` å±æ€§ã€?ä½ è¦ä¹ˆé—æ¼äº†å±æ€?å€¼ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.MasterPuti.org">MasterPuti.org</a>
      </p>
    </footer>
  </body>
</html>
```

