---
id: 5ef9b03c81a63668521804ee
title: æ­¥éª¤ 41
challengeType: 0
dashedName: step-41
---

# --description--

æ‰€æœ‰é¡µé¢éƒ½åº”è¯¥ä»?`<!DOCTYPE html>` å¼€å¤´ã€?è¿™ä¸ªç‰¹æ®Šçš„å­—ç¬¦ä¸²è¢«ç§°ä¸?dfn>å£°æ˜</dfn>ï¼Œç¡®ä¿æµè§ˆå™¨å°è¯•æ»¡è¶³è¡Œä¸šè§„èŒƒã€?

`<!DOCTYPE html>` å‘Šè¯‰æµè§ˆå™¨è¯¥æ–‡æ¡£æ˜¯ä¸€ä¸?HTML5 æ–‡æ¡£ï¼Œå³ HTML çš„æœ€æ–°ç‰ˆæœ¬ã€?

å°†æ­¤å£°æ˜æ·»åŠ ä¸ºä»£ç çš„ç¬¬ä¸€è¡Œã€?

# --hints--

ä½ çš„ä»£ç åº”ä»¥ `<!DOCTYPE html>` å£°æ˜å¼€å¤´ã€?ä½ è¦ä¹ˆé—æ¼äº†å£°æ˜ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ï¼Œè¦ä¹ˆå®ƒä¸æ˜¯ç¬¬ä¸€è¡Œä»£ç ã€?

```js
assert.match(code, /\<\s*!DOCTYPE\s+html\s*\>/);
```

ä½ çš„ `<!DOCTYPE html>` å¿…é¡»ä½äºæ–‡æ¡£çš„é¡¶éƒ¨ã€?

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /^\<\!DOCTYPEhtml\>\<html/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

<html lang="en">
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.MasterPuti.org">MasterPuti.org</a>
      </p>
    </footer>
  </body>
</html>
```

