---
id: 5ef9b03c81a63668521804d1
title: æ­¥éª¤ 28
challengeType: 0
dashedName: step-28
---

# --description--

åœ?`figure` å…ƒç´ ä¹‹åŽï¼Œæ·»åŠ å¦ä¸€ä¸?`h3` å…ƒç´ ï¼Œæ–‡æœ¬ä¸ºï¼?

`Top 3 things cats hate:`

# --hints--

åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¸Šæ–¹åº”è¯¥æœ‰ä¸€ä¸?`h3` å…ƒç´ ã€?ç¡®ä¿å®ƒæœ‰ä¸€ä¸ªå¼€å§‹å’Œç»“æŸæ ‡ç­¾ã€?

```js
assert.equal(
  document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName,
    'H3'
);
assert.lengthOf(code.match(/<\/h3\>/g),  2);
```

æ–°çš„ `h3` å…ƒç´ åº”è¯¥åŒ…å«æ–‡æœ¬ `Top 3 things cats hate:`ã€?ç¡®ä¿åœ¨æ–‡æœ¬æœ«å°¾åŒ…å«å†’å·ã€?

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild.innerText?.trim().toLowerCase()
    .replace(/\s+/g, ' '), 'top 3 things cats hate:'
);
```

æ–°çš„ `h3` å…ƒç´ ä¸Šæ–¹åº”è¯¥æœ‰ä¸€ä¸?`figure`ã€?ä½ å¯èƒ½ä¸å°å¿ƒåˆ é™¤äº?`figure` å…ƒç´ ã€?

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  ?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'H3');
assert.equal(secondSectionLastElemNode.previousElementSibling.nodeName, 'FIGURE');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

