---
id: 5ef9b03c81a63668521804ea
title: æ­¥éª¤ 38
challengeType: 0
dashedName: step-38
---

# --description--

æ³¨æ„ï¼Œåˆ°ç›®å‰ä¸ºæ­¢æ·»åŠ åˆ°é¡µé¢ä¸­çš„æ‰€æœ‰å†…å®¹éƒ½ä½äº`body`å…ƒç´ ä¸­ã€?æ‰€æœ‰åº”è¯¥å‘ˆç°åˆ°é¡µé¢çš„é¡µé¢å†…å®¹å…ƒç´ éƒ½æ”¾åœ¨ `body` å…ƒç´ ä¸­ã€?ä½†æ˜¯ï¼Œå…¶ä»–é‡è¦ä¿¡æ¯åœ¨ `head` å…ƒç´ ä¸­ã€?

`head` å…ƒç´ ç”¨äºåŒ…å«å…³äºæ–‡æ¡£çš„å…ƒæ•°æ®ï¼Œå¦‚å®ƒçš„æ ‡é¢˜ã€æ ·å¼è¡¨é“¾æ¥å’Œè„šæœ¬ã€?å…ƒæ•°æ®æ˜¯ä¸ç›´æ¥æ˜¾ç¤ºåœ¨é¡µé¢ä¸Šçš„å…³äºé¡µé¢çš„ä¿¡æ¯ã€?

åœ?`body` å…ƒç´ ä¸Šæ–¹æ·»åŠ ä¸€ä¸?`head` å…ƒç´ ã€?

# --hints--

ä½ è¦ä¹ˆåˆ é™¤äº† `body` å…ƒç´ ï¼Œè¦ä¹ˆå®ƒç¼ºå°‘å¼€å§‹æ ‡ç­¾æˆ–ç»“æŸæ ‡ç­¾ã€?

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

ä½ çš„ `head` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.match(code, /\<head\>/);
```

ä½ çš„ `head` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /\<\/head\>/);
```

ä½ çš„ `head` å…ƒç´ åº”ä½äº?`body` å…ƒç´ å¼€å§‹æ ‡ç­¾çš„ä¸Šæ–¹ã€?ä½ æŠŠå®ƒæ”¾åœ¨åˆ«çš„åœ°æ–¹äº†ã€?

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.MasterPuti.org">MasterPuti.org</a>
      </p>
    </footer>
  </body>
</html>
```

