---
id: 5ef9b03c81a63668521804d0
title: æ­¥éª¤ 27
challengeType: 0
dashedName: step-27
---

# --description--

è¦å¼ºè°ƒä¸€ä¸ªç‰¹å®šçš„å•è¯æˆ–çŸ­è¯­ï¼Œä½ å¯ä»¥ä½¿ç”?`em` å…ƒç´ ã€?

é€šè¿‡å°?`figcaption` å…ƒç´ ä¸­çš„å•è¯ `love` åŒ…è£¹åœ¨å¼ºè°?`em` å…ƒç´ ä¸­æ¥å¼ºè°ƒå®ƒã€?

# --hints--

ä½ çš„å¼ºè°ƒ `em` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('em'));
```

ä½ çš„å¼ºè°ƒ `em` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/em\>/);
```

ä½ è¦ä¹ˆåˆ é™¤äº† `figcaption` å…ƒç´ ï¼Œè¦ä¹ˆå®ƒç¼ºå°‘å¼€å§‹æˆ–ç»“æŸæ ‡ç­¾ã€?

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

ä½ çš„å¼ºè°ƒ `em` å…ƒç´ åº”åŒ…è£¹æ–‡æœ?`love`ã€?ä½ è¦ä¹ˆçœç•¥äº†æ–‡æœ¬ï¼Œè¦ä¹ˆæœ‰æ‹¼å†™é”™è¯¯ã€?

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

`figcaption` çš„æ–‡æœ¬åº”è¯¥æ˜¯ `Cats love lasagna`ã€?æ£€æŸ¥æ˜¯å¦æœ‰æ‹¼å†™é”™è¯¯ï¼Œä»¥å?`em` å…ƒç´ çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾å‘¨å›´æ˜¯å¦æœ‰å¿…è¦çš„ç©ºæ ¼ã€?

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

