---
id: 5ef9b03c81a63668521804e8
title: æ­¥éª¤ 36
challengeType: 0
dashedName: step-36
---

# --description--

åœ?`footer` å…ƒç´ ä¸­åµŒå¥—ä¸€ä¸?`p` å…ƒç´ ï¼Œå…¶æ–‡æœ¬ä¸?`No Copyright - MasterPuti.org`ã€?

# --hints--

ä½ è¦ä¹ˆåˆ é™¤äº† `footer` å…ƒç´ ï¼Œè¦ä¹ˆå®ƒç¼ºå°‘å¼€å§‹æ ‡ç­¾æˆ–ç»“æŸæ ‡ç­¾ã€?

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

ä½ çš„ `footer` å…ƒç´ åº”è¯¥åŒ…å«ä¸€ä¸?`p` å…ƒç´ ã€?ç¡®ä¿ä¸?`p` å…ƒç´ æ·»åŠ ä¸€ä¸ªå¼€å§‹æ ‡ç­¾å’Œç»“æŸæ ‡ç­¾ã€?

```js
assert.exists(document.querySelector('footer > p'));
```

ä½ çš„ `footer` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 3);
```

ä½ çš„ `p` å…ƒç´ çš„æ–‡æœ¬åº”è¯¥æ˜¯ `No Copyright - MasterPuti.org`ã€?ä½ è¦ä¹ˆé—æ¼äº†æ–‡æœ¬ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ï¼Œè¦ä¹ˆæ²¡æœ‰æŠŠæ–‡æœ¬æ”¾åœ?`p` å…ƒç´ çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ä¹‹é—´ã€?

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0]?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /^No Copyright - MasterPuti\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    <footer>

    </footer>
--fcc-editable-region--
  </body>
</html>
```

