---
id: 5efada803cbd2bbdab94e332
title: æ­¥éª¤ 31
challengeType: 0
dashedName: step-31
---

# --description--

åœ¨ä½ åˆšåˆšæ·»åŠ çš?`figure` å…ƒç´ ä¸­ï¼ŒåµŒå¥—ä¸€ä¸?`img` å…ƒç´ ï¼?`src` å±æ€§è®¾ç½®ä¸º `https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg`ã€?

# --hints--

ä½ çš„ç¬¬äºŒä¸?`figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 2);
```

ä½ çš„ç¬¬äºŒä¸?`figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.isAtLeast(code.match(/<\/figure>/g)?.length, 2);
```

åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¸Šæ–¹åº”è¯¥æœ‰ä¸€ä¸?`figure` å…ƒç´ ã€?ä½ æŠŠå®ƒä»¬çš„é¡ºåºå†™é”™äº†ã€?

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

ä½ åº”è¯¥åœ¨ `figure` å…ƒç´ ä¸­åµŒå¥—ç¬¬ä¸‰ä¸ª `img` å…ƒç´ ã€?

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.exists(
  catsImg
);
```

ç¬¬ä¸‰å¼ å›¾åƒåº”è¯¥æœ‰ä¸€ä¸ªè®¾ç½®ä¸º `https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg` çš?`src` å±æ€§ã€?

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.equal(
    catsImg?.getAttribute('src')?.toLowerCase(), 'https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg'
);
```

è™½ç„¶ä½ å·²å°†æ–°å›¾åƒçš?`src` è®¾ç½®ä¸ºæ­£ç¡®çš„ URLï¼Œä½†å»ºè®®å§‹ç»ˆå°†å±æ€§çš„å€¼ç”¨å¼•å·æ‹¬èµ·æ¥ã€?

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.MasterPuti\.org\/curriculum\/cat-photo-app\/cats\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
--fcc-editable-region--
        <figure>

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

