---
id: 5efae0543cbd2bbdab94e333
title: æ­¥éª¤ 32
challengeType: 0
dashedName: step-32
---

# --description--

ä¸ºäº†æé«˜ä½ åˆšåˆšæ·»åŠ çš„å›¾åƒçš„å¯è®¿é—®æ€§ï¼Œæ·»åŠ å…·æœ‰ä»¥ä¸‹æ–‡æœ¬çš?`alt` å±æ€§ï¼š

`Two tabby kittens sleeping together on a couch.`

# --hints--

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.lengthOf(document.querySelectorAll('figure'), 2);
```

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.lengthOf(code.match(/<\/figure>/g), 2);
```

æœ€åä¸€ä¸?`section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾çš„ä¸Šæ–¹åº”è¯¥æœ‰ä¸€ä¸?`figure` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

çŒ?`img` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`figure` å…ƒç´ ä¸­ã€?

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.exists(catsImg);
```

ç¬¬ä¸‰å¼ å›¾åƒåº”è¯¥æœ‰ä¸€ä¸ªè®¾ç½®ä¸º `https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg` çš?`src` å±æ€§ã€?

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.strictEqual(
  catsImg?.src?.toLowerCase(), 'https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg'
);
```

çŒ?`img` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸?`alt` å±æ€§ï¼Œå€¼ä¸º `Five cats looking around a field.`

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.match(
  catsImg
    ?.getAttribute('alt')
    ?.replace(/\s+/g, ' '),
    /^Two tabby kittens sleeping together on a couch..?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
--fcc-editable-region--
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg">
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

