---
id: 5efae16e3cbd2bbdab94e334
title: æ­¥éª¤ 33
challengeType: 0
dashedName: step-33
---

# --description--

åœ¨æœ€åä¸€ä¸?`img` å…ƒç´ ä¹‹åï¼Œæ·»åŠ ä¸€ä¸?`figcaption` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Cats hate other cats.`

# --hints--

ä½ çš„ `figcaption` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.lengthOf(document.querySelectorAll('figcaption'), 2);
```

ä½ çš„ `figcaption` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.lengthOf(code.match(/<\/figcaption\>/g), 2);
```

åœ¨ç¬¬äºŒä¸ª `section` å…ƒç´ çš„ç»“æŸæ ‡ç­¾ä¸Šæ–¹åº”è¯¥æœ‰ä¸€ä¸?`figure` å…ƒç´ ã€?

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

æœ€åä¸€ä¸?`img` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`figure` å…ƒç´ ä¸­ã€?

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.equal(
    catsImg?.getAttribute('src')?.toLowerCase(), 'https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg'
);
```

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.lengthOf(document.querySelectorAll('figure'), 2);
```

ä½ çš„ `figure` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.lengthOf(code.match(/<\/figure\>/g), 2);
```

`figcaption` å…ƒç´ åº”è¯¥åµŒå¥—åœ?`figure` å…ƒç´ ä¸­ã€?

```js
assert.lengthOf(document.querySelectorAll('figure > figcaption'), 2);
```

åµŒå¥—åœ?`figure` å…ƒç´ ä¸­çš„ `figcaption` å…ƒç´ åº”ä½äº?`img` å…ƒç´ ä¸‹æ–¹ã€?ä½ çš„ `img` å…ƒç´ å’?`figcaption` å…ƒç´ çš„é¡ºåºé”™äº†ã€?

```js
assert.equal(
  document.querySelectorAll('figcaption')[1]?.previousElementSibling.nodeName,
    'IMG'
);
```

`figcaption` å…ƒç´ åº”åŒ…å«æ–‡æœ?`Cats hate other cats.`ã€‚ä½ è¦ä¹ˆå¿½ç•¥äº†å•è¯ï¼Œè¦ä¹ˆæ‹¼å†™æœ‰è¯¯ã€?

```js
assert.match(
  document
    .querySelectorAll('figcaption')[1]
    ?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),
    /^Cats hate other cats\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

