---
id: 5f07be6ef7412fbad0c5626b
title: æ­¥éª¤ 18
challengeType: 0
dashedName: step-18
---

# --description--

åœ¨æ·»åŠ ä»»ä½•æ–°å†…å®¹ä¹‹å‰ï¼Œä½ åº”è¯¥ä½¿ç”¨ `section` å…ƒç´ å°†çŒ«å’ªç…§ç‰‡å†…å®¹ä¸æœªæ¥çš„å†…å®¹åˆ†å¼€ã€?

`section` å…ƒç´ ç”¨äºåœ¨æ–‡æ¡£ä¸­å®šä¹‰å„éƒ¨åˆ†ï¼Œå¦‚ç« èŠ‚ã€é¡µçœ‰ã€é¡µè„šæˆ–æ–‡æ¡£çš„ä»»ä½•å…¶ä»–éƒ¨åˆ†ã€?å®ƒæ˜¯ä¸€ä¸ªå¯¹ SEO å’Œæ— éšœç¢æœ‰å¸®åŠ©çš„è¯­ä¹‰åŒ–å…ƒç´ ã€?

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

å°†ä½ çš?`h2` å…ƒç´ ã€ä¸¤ä¸?`p` å…ƒç´ å’Œé”šå…ƒç´ ï¼ˆ`a`ï¼‰åµŒå¥—åœ¨ä¸€ä¸?`section` å…ƒç´ ä¸­ã€?

# --hints--

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.exists(document.querySelector('section'));
```

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.match(code, /<\/section\s*>/i);
```

æ•´ä¸ª `section` å…ƒç´ åº”è¯¥åœ?`main` å…ƒç´ çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ä¹‹é—´ã€?

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

ç°æœ‰çš?`h2` å…ƒç´ ã€`p` å…ƒç´ å’Œé”šå…ƒç´ ï¼ˆ`a`ï¼‰åº”è¯¥åœ¨ `section` å…ƒç´ çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ä¹‹é—´ã€?

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

`h1` å…ƒç´ ä¸åº”è¯¥è¢«åµŒå¥—åœ?`section` å…ƒç´ ä¸­ã€?

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

