---
id: 5f07c98cdb9413cbd4b16750
title: æ­¥éª¤ 19
challengeType: 0
dashedName: step-19
---

# --description--

æ˜¯æ—¶å€™æ·»åŠ ä¸€ä¸ªæ–°éƒ¨åˆ†äº†ã€?åœ¨ç°æœ?`section` å…ƒç´ ä¸‹æ–¹æ·»åŠ ç¬¬äºŒä¸?`section` å…ƒç´ ã€?

# --hints--

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªå¼€å§‹æ ‡ç­¾ã€?å¼€å§‹æ ‡ç­¾çš„è¯­æ³•ä¸ºï¼š`<elementName>`ã€?

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

ä½ åº”è¯¥åªæ·»åŠ ä¸€ä¸?`section` å¼€å§‹æ ‡ç­¾ã€?è¯·åˆ é™¤ä»»ä½•å¤šä½™çš„ã€?

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

ä½ çš„ `section` å…ƒç´ åº”è¯¥æœ‰ä¸€ä¸ªç»“æŸæ ‡ç­¾ã€?ç»“æŸæ ‡ç­¾åœ?`<` å­—ç¬¦ä¹‹åæœ‰ä¸€ä¸?`/`ã€?

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

ä½ åº”è¯¥åªæ·»åŠ ä¸€ä¸?`section` ç»“æŸæ ‡ç­¾ã€?è¯·åˆ é™¤ä»»ä½•å¤šä½™çš„ã€?

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

ç¬¬äºŒä¸?`section` å…ƒç´ ä¸åº”åµŒå¥—åœ¨ç¬¬ä¸€ä¸?`section` å…ƒç´ ä¸­ã€?

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

ä¸¤ä¸ª `section` å…ƒç´ éƒ½åº”è¯¥åœ¨ `main` å…ƒç´ çš„å¼€å§‹å’Œç»“æŸæ ‡ç­¾ä¹‹é—´ã€?

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

