---
id: 62bb4009e3458a128ff57d5d
title: æ­¥éª¤ 42
challengeType: 0
dashedName: step-42
---

# --description--

ä½ å¯ä»¥é€šè¿‡åœ?`head` ä¸­æ·»åŠ?`meta` å…ƒç´ æ¥è®¾ç½®æµè§ˆå™¨è¡Œä¸ºã€?è¿™æ˜¯ä¸€ä¸ªç¤ºä¾‹ï¼š

```html
<meta attribute="value">
```

åœ?`head` å…ƒç´ å†…ï¼ŒåµŒå¥—ä¸€ä¸?`meta` å…ƒç´ ï¼Œå…¶ `charset` å±æ€§çš„å€¼è®¾ç½®ä¸º `utf-8`ã€?è¿™å‘Šè¯‰æµè§ˆå™¨å¦‚ä½•ä¸ºé¡µé¢ç¼–ç å­—ç¬¦ã€?

æ³¨æ„ `meta` å…ƒç´ æ˜¯ä¸€ä¸ªç©ºå…ƒç´ ã€?

é€šè¿‡æœ€åçš„æ›´æ”¹ï¼Œä½ å°±å®Œæˆäº†çŒ«å’ªç›¸å†Œåº”ç”¨å·¥ä½œåŠã€?æ­å–œï¼?

# --hints--

ä½ åº”è¯¥åœ¨ `head` å…ƒç´ ä¸­åˆ›å»ºä¸€ä¸?`meta` å…ƒç´ ã€?

```js
assert.exists(document.querySelector('head > meta'));
```

ä½ åº”è¯¥ç»™ `meta` å…ƒç´ æ·»åŠ ä¸€ä¸ªå€¼ä¸º `UTF-8` çš?`charset`ã€?

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

ä½ çš„ `meta` å…ƒç´ åº”è¯¥æ˜¯ä¸€ä¸ªç©ºå…ƒç´ ï¼Œä½ ä¸éœ€è¦æ·»åŠ?`</meta>`ã€?

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>

<html lang="en">
--fcc-editable-region--
  <head>

    <title>CatPhotoApp</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.MasterPuti.org">MasterPuti.org</a>
      </p>
    </footer>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.MasterPuti.org">MasterPuti.org</a>
      </p>
    </footer>
  </body>
</html>
```


