---
id: 671144cdcc01d73f7dd79dc9
title: æ­¥éª¤ 16
challengeType: 0
dashedName: step-16
---

# --description--

ç°åœ¨ä½ å·²ç»æ·»åŠ äº†é“¾æ¥ï¼Œå¯ä»¥åˆ é™¤æ³¨é‡Šäº†ã€?

# --hints--

ä½ çš„ä»£ç ä¸­ä¸åº”è¯¥æœ‰ä»»ä½•æ³¨é‡Šã€?

```js
assert.notMatch(code, /<!--/);
assert.notMatch(code, /-->/);
```

æ–‡æœ¬ `TODO: Add link to cat photos` ä¸åº”å†å‡ºç°ã€?

```js
assert.notMatch(code, /TODO:?/i)
assert.notMatch(code, /Add\s*link\s*to\s*cat\s*photos/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
--fcc-editable-region--
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

