---
id: 671144cdcc01d73f7dd79dc9
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

现在你已经添加了链接，可以删除注释了。

太棒了！你已经成功地为猫咪照片添加了链接，现在让我们清理一下代码，把不再需要的注释删除掉吧！

# --hints--

你的代码中不应该有任何注释。

```js
assert.notMatch(code, /<!--/);
assert.notMatch(code, /-->/);
```

文本 `TODO: Add link to cat photos` 不应再出现。

```js
assert.notMatch(code, /TODO:?/i)
assert.notMatch(code, /Add\s*link\s*to\s*cat\s*photos/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
--fcc-editable-region--
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```