---
id: 671141f948cbab359e74cc93
title: æ­¥éª¤ 13
challengeType: 0
dashedName: step-13
---

# --description--

æ·»åŠ  `p` æ ‡ç­¾ï¼Œå°† `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` å˜æˆä¸€ä¸ªæ®µè½ã€?

# --hints--

ä½ åº”è¯¥åœ¨ `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` å‘¨å›´æ·»åŠ ä¸€ä¸?`p` å…ƒç´ ã€?

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P);
assert.strictEqual(P.innerText?.replace(/\s+/g, ' ').trim(), "See more cat photos in our gallery.");
```

åœ¨æ–°æ®µè½ä¸­ä»ç„¶åº”è¯¥æœ‰é”šå…ƒç´ ã€?

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P?.querySelector('a'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

