---
id: 671141feba228a35cefba82d
title: æ­¥éª¤14
challengeType: 0
dashedName: step-14
---

# --description--

å°†ç°æœ‰æ–‡æœ?`cute cats` å˜æˆä¸€ä¸ªé”šå…ƒç´ ï¼Œé“¾æ¥åˆ°ï¼?

`https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

åœ¨ç¬¬ä¸€ä¸?`p` å…ƒç´ ä¸­åº”è¯¥æœ‰ä¸€ä¸ªæ–°çš„é”šå…ƒç´ ã€?

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

é”šæ–‡æœ¬åº”ä¸?`cute cats`ã€?

```js
assert.strictEqual(document.querySelector('a')?.innerText?.trim(), "cute cats");
```

æ–°çš„é”šå…ƒç´ çš„ `href` å±æ€§åº”ä¸?`https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg`ã€?

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.MasterPuti.org/curriculum/cat-photo-app/running-cats.jpg");
```

åœ¨ä½ çš„ä»£ç ä¸­åº”è¯¥åªæœ‰ä¸€ä¸ªçŸ­è¯?`cute cats` ã€?

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

`p` å…ƒç´ çš„æ–‡æœ¬ä»åº”æ˜¯ `Everyone loves cute cats online!`ã€?

```js
assert.strictEqual(document.querySelector('p')?.innerText?.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

