---
id: 671141f948cbab359e74cc93
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

添加 `p` 标签，将 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` 变成一个段落。

继续保持 excellent 的学习状态！将文本包装在适当的标签中是网页设计的基础技能之一。

# --hints--

你应该在 `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` 周围添加一个 `p` 元素。

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P);
assert.strictEqual(P.innerText?.replace(/\s+/g, ' ').trim(), "See more cat photos in our gallery.");
```

在新段落中仍然应该有锚元素。

```js
const P = document.querySelectorAll('p')[1];
assert.exists(P?.querySelector('a'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```