---
id: 671b6e873249bb35c9debfcf
title: æ­¥éª¤ 12
challengeType: 0
dashedName: step-12
---

# --description--

åœ¨é”šå…ƒç´ ä¹‹å‰æ·»åŠ æ–‡å­— `See more`ï¼Œåœ¨é”šå…ƒç´ ä¹‹åæ·»åŠ æ–‡å­?`in our gallery`ã€?

# --hints--

ä½ åº”è¯¥ä½¿ç”¨ç°æœ‰çš„é”šå…ƒç´ ï¼Œè€Œä¸æ˜¯åˆ›å»ºä¸€ä¸ªæ–°ã€?

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

é”šå…ƒç´ å†…çš„æ–‡å­—ä¸åº”æ›´æ”¹ï¼Œå¿…é¡»æ˜?`cat photos`ã€?

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

åœ¨ä½ çš„ä»£ç ä¸­åº”è¯¥æœ?`See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery`ã€?

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

