---
id: 60a3e3396c7b40068ad6996b
title: æ­¥éª¤ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

å¼€å§‹é¡¹ç›®ï¼Œåœ?`body` å…ƒç´ ä¸­æ·»åŠ ä¸€ä¸?`img` å…ƒç´ ï¼Œå…¶ `src` ä¸?`https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png` ã€?

# --hints--

ä½ çš„ä»£ç åº”è¯¥æœ‰ä¸€ä¸?`img` å…ƒç´ ã€?

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` å…ƒç´ çš?`src` åº”è¯¥æ˜?`https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png`ã€?

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

