---
id: 68a47dcdbce12193caf77012
title: Implement a Sorted Index Finder
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

在这个实验中，你将创建一个函数，返回一个值在数组按升序排序后应该被插入的最低索引。

**目标**：完成下面的用户故事并让所有测试通过以完成实验。

**用户需求：**

1. 你应该有一个 `getIndexToIns` 函数，它接受两个参数：一个数组和一个数字。
2. 你应该使用 `sort` 方法按升序对数组进行排序。
3. 你的 `getIndexToIns` 函数应该使用 `findIndex` 方法返回数字应该被插入的最低索引。
4. 你的 `getIndexToIns` 函数应该始终返回一个数字。

**提示：**  
`findIndex` 方法是JavaScript中的内置数组方法。它接受一个回调函数并返回满足条件的第一个元素的索引。`findIndex` 和 `sort` 都是高阶函数。

**示例：**

- `getIndexToIns([1, 2, 3, 4], 1.5)` 应该返回 `1`，因为 `1.5` 大于 `1`（索引 `0`）且小于 `2`（索引 `1`）。
- `getIndexToIns([20, 3, 5], 19)` 应该返回 `2`，因为在排序为 `[3, 5, 20]` 后，`19` 小于 `20`（索引 `2`）且大于 `5`（索引 `1`）。


# --hints--

你应该有一个 `getIndexToIns` 函数。

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` 应该始终返回一个数字。

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

你的函数应该使用 `sort` 方法。

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


你的函数应该使用 `findIndex` 方法。

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` 应该返回 `3`。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` 应该返回 `2`。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` 应该返回 `1`。

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` 应该返回 `0`。

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` 应该返回 `2`。

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` 应该返回 `2`。

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` 应该返回 `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` 应该返回 `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--
```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
``` 

