---
id: 68e4276483b0c71658aab1cd
title: 步骤 5
challengeType: 1
dashedName: step-5
---

# --description--

现在是时候看看结果了。

首先添加一个 `console.log()`，输出字符串 `"Result of trimming whitespace from both ends:"`。

然后在那个 `console.log()` 下方，添加另一个 `console.log()` 输出 `cleanedInput` 变量。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该向控制台输出字符串 `"Result of trimming whitespace from both ends:"`。

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

你应该向控制台输出 `cleanedInput` 变量。

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
