---
id: 68e4276483b0c71658aab1d3
title: 步骤 14
challengeType: 1
dashedName: step-14
---

# --description--

正如你在之前的课程和工作坊中回忆的那样，你可以使用 `slice()` 方法提取字符串的一部分。

这里是一个提醒：

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

将 `camelCasedVersion` 的值从空字符串更改为使用 `slice()` 方法处理 `lowercaseWord` 变量的结果。传入数字 `0` 和 `5`，它们代表 `slice()` 方法的起始和结束索引。

现在你应该在控制台中看到单词 `"camel"`。

# --hints--

你的 `camelCasedVersion` 变量不应该是一个空字符串。

```js
assert.isNotEmpty(camelCasedVersion);
```

你应该将 `lowercaseWord.slice(0, 5)` 的结果赋值给 `camelCasedVersion` 变量。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

你不应该将字面量字符串 `"camel"` 赋值给 `camelCasedVersion`。请重新阅读说明以获得正确的方法。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method — trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
