---
id: 68e4276483b0c71658aab1ce
title: 步骤 6
challengeType: 1
dashedName: step-6
---

# --description--

有时候你只想去除字符串开头的空格。这时候 `trimStart()` 方法就派上用场了。

这里有一个例子：

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

记住字符串是不可变的。这个方法会创建一个带有结果的新字符串，而不是修改原始字符串。

创建一个变量 `trimmedStart`，并将使用 `trimStart()` 方法处理 `userInput` 变量的结果赋值给它。

# --hints--

你应该有一个 `trimmedStart` 变量。

```js
assert.exists(trimmedStart);
```

你的 `trimmedStart` 变量应该是一个字符串。

```js
assert.isString(trimmedStart);
```

你应该将 `userInput.trimStart()` 的结果赋值给 `trimmedStart` 变量。

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
