---
id: 68e4276483b0c71658aab1ca
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个工作坊中，你将练习使用各种字符串方法来格式化字符串。

首先创建一个名为 `userInput` 的变量，并赋值如下内容：`"   Hello World!   "`。

# --hints--

你应该有一个 `userInput` 变量。

```js
assert.exists(userInput);
```

你的 `userInput` 变量应该是一个字符串。

```js
assert.isString(userInput);
```

你应该给 `userInput` 变量赋值 `"   Hello World!   "`。确保在 `H` 前面有三个空格，在 `!` 后面也有三个空格。

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
