---
id: 68f3f5ca1d2a490c97f14f03
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

现在是时候看看去除字符串末尾空格的结果了。

在你的 `trimmedEnd` 变量下方，添加一个 `console.log()`，输出字符串 `"After using the trimEnd() method — trailing spaces removed:"`。

然后在那个 `console.log()` 下方，添加另一个 `console.log()` 输出 `trimmedEnd` 变量。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该向控制台输出字符串 `"After using the trimEnd() method — trailing spaces removed:"`。

```js
assert.equal(getLogs()[6], "After using the trimEnd() method — trailing spaces removed:");
```

你应该向控制台输出 `trimmedEnd` 变量。

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method — leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--
const trimmedEnd = userInput.trimEnd();

--fcc-editable-region--
```
