---
id: 68f3f3d65d51d80ac0b412fc
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

现在是时候看看去除字符串开头空格的结果了。

在你的 `trimmedStart` 变量下方，添加一个 `console.log()`，输出字符串 `"After using the trimStart() method — leading spaces removed:"`。

然后在那个 `console.log()` 下方，添加另一个 `console.log()` 输出 `trimmedStart` 变量。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该向控制台输出字符串 `"After using the trimStart() method — leading spaces removed:"`。

```js
assert.equal(getLogs()[4], "After using the trimStart() method — leading spaces removed:");
```

你应该向控制台输出 `trimmedStart` 变量。

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

--fcc-editable-region--
const trimmedStart = userInput.trimStart();

--fcc-editable-region--
```
