---
id: 6140cd32d018ed0f600eefce
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

将 `.wheel` 选择器的 `position` 设置为 `absolute`，并将 `height` 和 `width` 都设置为 `55vw`，以便根据视口宽度调整大小。


# --hints--

`.wheel` 的 `position` 应为 `absolute`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.position, 'absolute');
```

`.wheel` 的 `height` 应为 `55vw`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.height, '55vw');
```

`.wheel` 的 `width` 应为 `55vw`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.width, '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
