---
id: 6140cbeec34e970dfe75e710
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

为 `.wheel` 元素创建一个选择器。首先设置 `border: 2px solid black`、`border-radius: 50%` 和 `margin-left: 50px`，让它看起来像摩天轮的外框。


# --hints--

应包含 `.wheel` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

`.wheel` 选择器的 `border` 属性应为 `2px solid black`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.border, '2px solid black');
```

`.wheel` 选择器的 `border-radius` 应为 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius, '50%');
```

`.wheel` 选择器的 `margin-left` 应为 `50px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft, '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
