---
id: 6140cdebd39d6a101e747529
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

为 `.line` 元素创建一个选择器。首先设置 `background-color: black`、`width: 50%` 和 `height: 2px`，为线段设定颜色与尺寸。


# --hints--

应包含 `.line` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line'));
```

`.line` 的 `background-color` 应为 `black`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor, "black");
```

`.line` 的 `width` 应为 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.width, "50%");
```

`.line` 的 `height` 应为 `2px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.height, "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
