---
id: 6140cfc08ca9c5128c3e6478
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

将 `.line` 选择器的 `position` 设置为 `absolute`，并将 `left` 与 `top` 都设置为 `50%`，以便把线段定位到容器中心。


# --hints--

`.line` 选择器的 `position` 属性应为 `absolute`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.position, 'absolute');
```

`.line` 的 `left` 属性应为 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.left, '50%');
```

`.line` 的 `top` 属性应为 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.top, '50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
}
--fcc-editable-region--
```
