---
id: 617ace7d831f9c16a569b38a
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

为 `.wheel` 选择器添加 `max-height` 和 `max-width` 属性，两个值都设置为 `500px`，这样可以限制摩天轮的最大尺寸，让布局更稳定。


# --hints--

`.wheel` 选择器的 `max-height` 属性应设置为 `500px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxHeight, '500px');
```

`.wheel` 选择器的 `max-width` 属性应设置为 `500px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxWidth, '500px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
}
--fcc-editable-region--
```
