---
id: 674ef2d357676e50e469165b
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--
现在让 `App` 也有一个漂亮的布局容器吧！

请把空字符串换成一个 `div`，并赋予它 `className="flex-container"`。这样等会放多个卡片时就能自动排版，好看又整齐～✨

# --hints--
`App` 不应继续渲染空字符串。

```js
  const testElem = await __helpers.prepTestComponent(window.index.App);
  assert.notEqual(testElem.innerHTML, '');
```

需要返回一个单一的 `div` 元素。

```js
  const testElem = await __helpers.prepTestComponent(window.index.App);
  assert.equal(testElem.firstElementChild?.tagName, 'DIV');
  assert.lengthOf(testElem.children, 1);
```

这个 `div` 必须带 `className="flex-container"`。

```js
  const testElem = await __helpers.prepTestComponent(window.index.App);
  assert.include([...testElem.firstElementChild.classList], 'flex-container');
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  return (
  --fcc-editable-region--
    ""
  --fcc-editable-region--
  );
}
```
