---
id: 68c34f4aeb2b1a4170506751
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

接下来，向表格添加第五行。首先添加另一?`tr` 元素。在?`tr` 元素内，添加一?`th` 元素，将 `scope` 属性设置为 `"row"`，文本内容为 `12:30 PM`。然后在?`th` 元素下方，添加一?`td` 元素，文本内容为 `Lunch Break`。你?`td` 元素还应该有一?`colspan` 属性设置为 `3`，以便它跨越所有三个轨道?

# --hints--

你应该在 `tbody` 元素内有第五?`tr` 元素?

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

你应该在第五?`tr` 元素内有一?`th` 元素?

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

你的 `th` 元素应该有文本内?`12:30 PM`?

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

你的 `th` 元素应该有一?`scope` 属性设置为 `"row"`?

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

你应该在第五?`tr` 元素内有一?`td` 元素?

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

你的 `td` 元素应该有文本内?`Lunch Break`?

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
```

你的 `td` 元素应该有一?`colspan` 属性设置为 `3`?

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--

      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
