---
id: 684aaf9ec670c68d20efd0d0
title: 构建用户配置管理器
challengeType: 27
dashedName: build-a-user-configuration-manager
---

# --description--

在这个实验中，你将构建一个用户配置管理器，允许用户管理他们的设置，如主题、语言和通知。你将实现添加、更新、删除和查看用户设置的功能。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该定义一个名为`add_setting`的函数，它有两个参数，分别代表设置字典和包含键值对的元组

1. `add_setting`函数应该：

   - 将键和值转换为小写。
   - 如果键设置已存在，返回`Setting '[key]' already exists! Cannot add a new setting with this name.`
   - 如果键设置不存在，将键值对添加到给定的设置字典中并返回`Setting '[key]' added with value '[value]' successfully!`。
   - 返回的消息中键和值应该为小写。

1. 你应该定义一个名为`update_setting`的函数，它有两个参数，分别代表设置字典和包含键值对的元组。

1. `update_setting`函数应该：

   - 将键和值转换为小写。
   - 如果键设置存在，更新给定设置字典中的值并返回：`Setting '[key]' updated to '[value]' successfully!`
   - 如果键设置不存在，返回`Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - 返回的消息中键和值应该为小写。

1. 你应该定义一个名为`delete_setting`的函数，它有两个参数，分别代表设置字典和一个键。

1. `delete_setting`函数应该：

   - 将传入的键转换为小写。
   - 如果键设置存在，从给定的设置字典中移除键值对并返回`Setting '[key]' deleted successfully!`
   - 如果键设置不存在，返回`Setting not found!`
   - 返回的消息中键应该为小写。

1. 你应该定义一个名为`view_settings`的函数，它有一个参数，代表设置字典。

1. `view_settings`函数应该：

   - 如果给定的设置字典为空，返回`No settings available.`
   - 如果字典包含任何设置，返回一个显示设置的字符串。该字符串应该以`Current User Settings:`开头，后跟键值对，每个键值对占一行，键要大写。例如，`view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})`应该返回：

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high

   ```

1. 为了测试代码，你应该创建一个名为`test_settings`的字典来存储一些用户配置偏好。

# --hints--

你应该创建一个名为`test_settings`的字典并向其中添加一些值。

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

你应该定义一个名为`add_setting`的函数。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

`add_setting`函数应该有两个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting`应该将键转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting`应该将值转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))`应该返回错误消息`Setting 'theme' already exists! Cannot add a new setting with this name.`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))`应该添加一个新的键值对并返回成功消息`Setting 'volume' added with value 'high' successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting`应该正确地将给定的键值对添加到字典中。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

你应该定义一个名为`update_setting`的函数。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

`update_setting`函数应该有两个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`update_setting`函数应该将键转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting`函数应该将值转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # 使用现有的"theme"键但使用大写值
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))`应该更新现有键并返回成功消息`Setting 'theme' updated to 'dark' successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))`当键不存在时应该返回错误消息`Setting 'volume' does not exist! Cannot update a non-existing setting.`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting`应该正确地更新字典中的给定键值对。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

你应该定义一个名为`delete_setting`的函数。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

`delete_setting`函数应该有两个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting`应该将键转换为小写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')`应该移除现有键并返回成功消息`Setting 'theme' deleted successfully!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting`当键不存在时应该返回错误消息`Setting not found!`。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting`应该正确地从字典中移除给定的键。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

你应该定义一个名为`view_settings`的函数。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

`view_settings`函数应该有一个参数。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings`如果给定的字典为空，应该返回消息`No settings available.`。

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings`应该为非空字典返回格式化的设置。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings`应该将每个设置名称的第一个字母大写。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings`应该显示正确的结果并以换行符结尾。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: English\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "English",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```