---
id: 687d8e05a529a4838f1c961c
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

接下来，移除 `Node` 类中的 `pass` 关键字，然后给 `Node` 类添加一个带有两个参数的 `__init__` 方法，这两个参数是 `self` 和 `element`。

在方法内部，将 `element` 赋值给 `self.element`。这样就存储了创建节点时传入的值。做得很好！

# --hints--

你应该在 `Node` 类中定义一个 `__init__` 方法。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

你的 `__init__` 方法应该有两个参数，`self` 和 `element`。

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

你应该在 `__init__` 方法中将 `element` 赋值给 `self.element`。

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```