---
id: 687d7deea6a2a22fa8816a73
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在之前的课程中，你学习了一些常见的数据结构。在这个工作坊中，你将实现一个链表。太棒了！

首先，创建一个名为 `LinkedList` 的类来表示将保存列表节点的结构。在类中添加 `pass` 关键字，这样测试就能通过啦。

# --hints--

你应该创建一个名为 `LinkedList` 的类。记住要添加 `pass` 关键字哦。

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```