---
id: 687d8c6f066877768f7ee441
title: 步骤 2
challengeType: 20
dashedName: step-2
---

# --description--

在你的 `LinkedList` 类中，将 `pass` 关键字替换为一个名为 `Node` 的内部类。这将代表存储列表中数据的各个节点。继续加油！

# --hints--

你应该在 `LinkedList` 类中创建一个 `Node` 类。记住要在 `Node` 类中添加 `pass` 关键字哦。

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```