---
id: 687e366b3004a9831e4cbd41
title: 步骤 7
challengeType: 20
dashedName: step-7
---

# --description--

添加一个名为 `is_empty` 的方法，返回 `self.length == 0`。

你将使用这个方法来跟踪列表中有多少个节点。所以，当 `self.length == 0` 时，它返回 `True`，表示列表没有节点。

如果 `self.length` 大于 `0`，它返回 `False`，表示列表至少有一个节点。太棒了！

# --hints--

你应该创建一个 `is_empty` 方法，返回 `self.length == 0`。

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--

--fcc-editable-region--
```