---
id: 688649c8673f842f49228a5e
title: 步骤 11
challengeType: 20
dashedName: step-11
---

# --description--

如果列表不为空，方法应该找到最后一个节点并将新节点链接到它的 `next` 指针。

要做到这一点，在 `if` 语句块之后创建一个 `else` 语句块。在 `else` 语句块中，创建一个 `current_node` 变量并将 `self.head` 赋值给它。继续加油！

# --hints--

你应该在 `else` 语句块中将 `self.head` 赋值给 `current_node`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```