---
id: 6888fe6c745b6852ed240229
title: 步骤 18
challengeType: 20
dashedName: step-18
---

# --description--

接下来，你将循环遍历链表以找到包含要移除的指定元素的节点。

创建一个 `while` 循环，检查 `current_node` 是否不为 `None` 以及 `current_node.element` 是否不等于 `element`。在 `while` 循环中，将 `current_node` 赋值给 `previous_node`。太聪明了！

# --hints--

你应该在 `remove` 方法中有一个 `while` 循环。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
`)})
```

你的 `while` 循环应该检查 `current_node` 是否不为 `None` 以及 `current_node.element` 是否不等于 `element`。

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

你应该在 `while` 循环中将 `current_node` 赋值给 `previous_node`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```