---
id: 688ac806194cfc60ba1e1767
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

在 `while` 循环之后，创建一个检查 `current_node` 是否为 `None` 的 `if` 语句。

当这种情况发生时，链表已被遍历但未找到要移除的元素，函数应该停止运行。因此，在 `if` 语句体内添加 `return` 关键字。做得很好！

# --hints--

你应该创建一个检查 `current_node` 是否为 `None` 的 `if` 语句，并在其中添加 `return` 关键字。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--

--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```