---
id: a6b0bb188d873cb2c8729495
title: Implement an HTML Entity Converter
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

这个实验是关于将字符串中的特殊字符转换为相应的HTML实体。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个`convertHTML`函数，该函数接受一个字符串作为参数。
1. `convertHTML`函数应该通过将参数字符串中的特殊字符转换为相应的HTML实体来返回一个新字符串。

   - `&` 应该转换为 `&amp;`。
   - `<` 应该转换为 `&lt;`。
   - `>` 应该转换为 `&gt;`。
   - `"` 应该转换为 `&quot;`。
   - `'` 应该转换为 `&apos;`。

# --hints--

你应该有一个`convertHTML`函数。

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")`应该返回字符串`Dolce &amp; Gabbana`。

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")`应该返回字符串`Hamburgers &lt; Pizza &lt; Tacos`。

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")`应该返回字符串`Sixty &gt; twelve`。

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')`应该返回字符串`Stuff in &quot;quotation marks&quot;`。

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")`应该返回字符串`Schindler&apos;s List`。

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")`应该返回字符串`&lt;&gt;`。

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")`应该返回字符串`abc`。

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js
function convertHTML(str) {
  // 你的代码在这里
}
```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```