---
id: 6852ea0c3cadb54e897acdac
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在这个工作坊中，你将会构建一个*电子邮件模拟器*，它可以模拟不同用户之间发送、接收和管理电子邮件的过程。通过这个项目，你会学到类、对象以及如何以面向对象的方式组织代码。

让我们从创建一个名为 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 的类开始吧，要使用 [class](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d7ca10173919bc46251f.md#L15-L16) 关键字哦！

# --hints--

你应该创建一个名为 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 的类。

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
