---
id: 6852f0c3fcedb964ab959921
title: 第 6 步
challengeType: 20
dashedName: step-6
---

# --description--

现在让我们测试一下 [read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f0c3fcedb964ab959921.md#L21-L21) 属性是否已正确添加到你的 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中。由于你已经在前面的步骤中有了一个电子邮件对象，所以打印 [read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f0c3fcedb964ab959921.md#L21-L21) 属性，看看它现在默认为 [False](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L33-L33)。

# --hints--

你应该打印 [email_obj.read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f0c3fcedb964ab959921.md#L17-L17) 来测试 read 属性。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
