---
id: 6852f06fdde1b4634a263f03
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

现在让我们通过创建一个电子邮件对象来测试 `Email` 类，并检查它是否正确地存储了信息。你可以打印几个属性来验证一切是否正常工作。

创建一个名为 `email_obj` 的电子邮件对象，发件人设置为 alice@example.com，收件人设置为 bob@example.com，主题设置为 Hello，正文设置为 Hi Bob!。然后分别打印 `sender` 和 `subject` 属性，以验证它们被正确存储。

# --hints--

你应该创建一个名为 `email_obj` 的电子邮件对象，并指定相应参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

你应该打印 `email_obj.sender`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

你应该打印 `email_obj.subject`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
--fcc-editable-region--
```
