---
id: 6853f34ff93594934f6636d9
title: 第 27 步
challengeType: 20
dashedName: step-27
---

# --description--

让我们为 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类添加一个字符串表示形式，这样我们就可以显示简短的邮件摘要了。

在 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中添加一个接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 作为参数的 [__str__](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 方法。

# --hints--

你应该在 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中定义一个 [__str__](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__str__")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
