---
id: 6853f4584f508e95cc7a9f4d
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

现在你将创建一个方法来列出收件箱中的所有邮件。这个方法应该处理收件箱为空的情况，并在有邮件时显示编号列表。

在你的 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L37-L37) 类中添加一个名为 [list_emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L39-L47) 的方法，该方法接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 作为参数。首先，创建一个 [if](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L32-L33) 语句，通过测试 [emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L38-L38) 列表来检查收件箱是否为空。如果为空，打印消息 [Your inbox is empty.\n](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L44-L44)。添加一个 [return](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L39-L39) 语句退出该方法。

# --hints--

你应该在 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L37-L37) 类中定义一个名为 [list_emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L39-L47) 的方法，该方法接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

你应该使用 [if not self.emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L42-L42) 来检查 [self.emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L38-L38) 是否为空。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

如果没有邮件，你应该打印消息 [Your inbox is empty.\n](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L44-L44)。

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

你应该在 [if](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L32-L33) 语句的末尾有一个 [return](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L39-L39) 语句。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--

```
