---
id: 6853f2da9b8ed891e9fdc5fc
title: 第 23 步
challengeType: 20
dashedName: step-23
---

# --description--

现在让我们显示发件人和收件人信息。为了显示谁发送和接收了邮件，向 [display_full_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f2da9b8ed891e9fdc5fc.md#L25-L28) 方法中添加两个打印语句，格式如下：

- [From: sender](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f2da9b8ed891e9fdc5fc.md#L27-L27) 其中 [sender](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 被替换为发件人的名字
- [To: receiver](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f2da9b8ed891e9fdc5fc.md#L28-L28) 其中 [receiver](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 被替换为收件人的名字

# --hints--

你应该使用 f-string 来打印发件人姓名，使用 [From: {self.sender.name}](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f2da9b8ed891e9fdc5fc.md#L27-L27)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

你应该使用 f-string 来打印收件人姓名，使用 [To: {self.receiver.name}](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f2da9b8ed891e9fdc5fc.md#L28-L28)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
