---
id: 68540aa7dfd449ad074c285c
title: 第 51 步
challengeType: 20
dashedName: step-51
---

# --description--

现在你将演示收件箱功能。Ramy 将检查他的收件箱，阅读第一封邮件，删除第二封邮件，然后再次检查他的收件箱以查看变化。

在你的 [main](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68540aa7dfd449ad074c285c.md#L153-L169) 函数中，发送邮件后，添加代码来：

- 使用 [check_inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L51-L54) 方法让 Ramy 检查他的收件箱。
- 让 Ramy 阅读第一封邮件。
- 让 Ramy 删除第一封邮件。
- 让 Ramy 再次检查他的收件箱。

完成这些后，你就完成了电子邮件模拟器！

# --hints--

你应该调用 [ramy.check_inbox()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L51-L54) 来显示 Ramy 的邮件。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().has_stmt("ramy.check_inbox()")`))
    }
})
```

你应该调用 [ramy.read_email(1)](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f996671e63a0b4fbf143.md#L21-L21) 来阅读第一封邮件。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().has_stmt("ramy.read_email(1)")`))
    }
})

```

你应该调用 [ramy.delete_email(1)](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4ded9f84896ff182dd8.md#L90-L90) 来删除第一封邮件。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().has_stmt("ramy.delete_email(1)")`))
    }
})

```

你应该再次调用 [ramy.check_inbox()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L51-L54) 来显示更新后的收件箱。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().is_ordered("ramy.delete_email(1)", "ramy.check_inbox()")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"
class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

    def check_inbox(self):
        print(f"\n{self.name}'s Inbox:")
        self.inbox.list_emails()

    def read_email(self, index):
        self.inbox.read_email(index)

    def delete_email(self, index):
        self.inbox.delete_email(index)

--fcc-editable-region--
def main():
    tory = User('Tory')
    ramy = User('Ramy')        

    tory.send_email(ramy, 'Hello', 'Hi Ramy, just saying hello!')
    ramy.send_email(tory, 'Re: Hello', 'Hi Tory, hope you are fine.')

--fcc-editable-region--
if __name__ == '__main__':
    main()