---
id: 68594df7779a164e6a61e8ae
title: 第 32 步
challengeType: 20
dashedName: step-32
---

# --description--

收件箱需要一个方法来阅读特定的电子邮件。当用户想要阅读电子邮件时，他们会指定想要查看的电子邮件编号（从索引 0 开始），该方法将显示完整的电子邮件内容。

在你的 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 类中添加一个名为 [read_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f996671e63a0b4fbf143.md#L21-L21) 的方法，该方法接受一个 [index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4ded9f84896ff182dd8.md#L80-L80) 参数。首先，检查收件箱是否为空，如果是，则打印消息 [Inbox is empty.\n](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68594df7779a164e6a61e8ae.md#L25-L25)。在之后添加一个 [return](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L43-L43) 语句以退出该方法。

# --hints--

你应该在 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 类中定义一个名为 [read_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f996671e63a0b4fbf143.md#L21-L21) 的方法，该方法接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 和 [index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4ded9f84896ff182dd8.md#L80-L80) 作为参数。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

你应该使用 [if not self.emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68594df7779a164e6a61e8ae.md#L25-L25) 来检查 [self.emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L39-L39) 是否为空。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

如果没有邮件，你应该打印 [Inbox is empty.\n](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/68594df7779a164e6a61e8ae.md#L25-L25) 并返回。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--

--fcc-editable-region--
```
