---
id: 685c1e32254e5993a5c8823e
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

现在你已经完成了 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类的初始设置，是时候创建能够发送和接收电子邮件的用户了。每个用户都有一个名字，并且能够执行电子邮件操作。

创建一个 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类，其中包含一个接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 和 [name](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685c1e32254e5993a5c8823e.md#L21-L21) 作为参数的 [__init__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L22-L22) 方法。在类中，将 [name](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685c1e32254e5993a5c8823e.md#L21-L21) 参数分配给 [self.name](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685c1e32254e5993a5c8823e.md#L21-L21)。

# --hints--

你应该创建一个 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

[User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 类应该有一个 [__init__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L22-L22) 方法，该方法接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 和 [name](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685c1e32254e5993a5c8823e.md#L21-L21) 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

你应该将 [name](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685c1e32254e5993a5c8823e.md#L21-L21) 参数分配给 [self.name](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685c1e32254e5993a5c8823e.md#L21-L21)。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
