---
id: 685cdfd1932d8ac6ad986813
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

在 [__str__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L49-L49) 方法中，你将显示电子邮件是否已被阅读。

条件表达式允许你根据条件在一行中分配值。

下面是一个例子：

```py
x = 10
y = 'Even' if x % 2 == 0 else 'Odd' # y 将是偶数
```

在方法中，在之前使用条件表达式，如果电子邮件已读则将字符串 [Read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L50-L50) 分配给变量 [status](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685cdfd1932d8ac6ad986813.md#L25-L25)，如果未读则分配 [Unread](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L50-L50)。

# --hints--

[__str__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L49-L49) 方法应该创建一个使用条件表达式的 [status](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685cdfd1932d8ac6ad986813.md#L25-L25) 变量。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_stmt("status = 'Read' if self.read else 'Unread'")`));
    }
})
```

[__str__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L49-L49) 方法应该返回包含状态的电子邮件信息。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_variable("status")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        pass
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
