---
id: 685d53ec88a16a4643c2437b
title: 第 44 步
challengeType: 20
dashedName: step-44
---

# --description--

现在让我们也在电子邮件列表中显示时间戳。更新 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中的 [__str__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L49-L49) 方法，在主题后包含时间戳。

修改返回语句，在末尾包含格式化为 [%Y-%m-%d %H:%M](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f70470f45d9c280af659.md#L25-L25) 的时间戳，用 [| Time:](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685d53ec88a16a4643c2437b.md#L25-L25) 分隔。

完整格式应为：

```py
[status] From: sender | Subject: subject | Time: time 
```

其中 [status](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685cdfd1932d8ac6ad986813.md#L25-L25) 是电子邮件的状态，[sender](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 是发送者的名字，[subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 是电子邮件的主题，[time](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108) 的格式为 [%Y-%m-%d %H:%M](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f70470f45d9c280af659.md#L25-L25)。

# --hints--

你应该更新 [__str__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L49-L49) 方法以包含格式化的时间戳。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_stmt("return f\\"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}\\"")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')
```
