---
id: 685d2fd1f309b9191f257c55
title: 第 40 步
challengeType: 20
dashedName: step-40
---

# --description--

在将时间戳集成到我们的电子邮件系统之前，让我们先练习使用日期时间格式化。[datetime.datetime.now()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L107-L107) 函数给我们当前的日期和时间，我们可以使用 [strftime()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108) 方法以不同的方式格式化它。

以下是 [strftime()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108) 使用格式代码的工作方式：

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # 输出: 2024-03-15 (年-月-日 使用 - 分隔符)
```

像 [%Y](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108)（年）、[%m](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108)（月）、[%d](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108)（日）这样的格式代码告诉 [strftime()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108) 包含什么内容，你可以在它们之间添加分隔符如 [-](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108)。

在你的代码底部，创建一个名为 [current_time](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L107-L107) 的变量，并将其分配给 [datetime.datetime.now()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L107-L107)。然后使用 [strftime()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108) 以小时:分钟:秒的格式打印时间，使用 [:](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685d2fd1f309b9191f257c55.md#L23-L23) 作为分隔符。

使用这些格式代码：[%H](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685d2fd1f309b9191f257c55.md#L23-L23) 表示小时（24小时制），[%M](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685d2fd1f309b9191f257c55.md#L23-L23) 表示分钟，和 [%S](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685d2fd1f309b9191f257c55.md#L23-L23) 表示秒。

# --hints--

你应该创建一个名为 [current_time](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L107-L107) 的变量，并将其分配给 [datetime.datetime.now()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L107-L107)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

你应该使用 [strftime()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f68dd3affa9adb2e6d99.md#L108-L108) 以 [%H:%M:%S](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685d2fd1f309b9191f257c55.md#L23-L23) 格式打印时间。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
