---
id: 68875554c6c32de788afc64e
title: 第 21 步
challengeType: 20
dashedName: step-21
---

# --description--

用户需要一种方式来正确阅读他们的电子邮件。

在 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中添加一个 [display_full_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L31-L31) 方法，该方法只接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 作为参数。在此方法中，调用 [mark_as_read()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25) 方法，当有人查看邮件时将其标记为已读。

此方法将用于以美观的格式显示电子邮件的完整内容。

# --hints--

你应该在 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中定义一个 [display_full_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L31-L31) 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

[display_full_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L31-L31) 方法应该只接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 作为参数。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

该方法应该调用 [self.mark_as_read()](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25) 将电子邮件标记为已读。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
