---
id: 68b936bf4f47d6c29a669f99
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

修改 [send_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L46-L46) 方法，通过调用 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 类的 [receive_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L41-L41) 方法将电子邮件添加到接收者的收件箱中。

# --hints--

你应该在 [send_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f90fd03acc9fd75ff860.md#L46-L46) 方法中添加 [receiver.inbox.receive_email(email)](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L42-L42)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_stmt("receiver.inbox.receive_email(email)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
