---
id: 68be6af1f43cc44541debe41
title: 第 37 步
challengeType: 20
dashedName: step-37
---

# --description--

你应该处理用户尝试使用无效索引删除电子邮件的情况，就像你在 [read_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L59-L59) 方法中所做的那样。

在 [delete_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L70-L70) 方法中：

- 通过减去 1 将基于 1 的 [index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L59-L59) 参数转换为基于 0 的索引，并将其存储在名为 [actual_index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L64-L64) 的变量中。
- 创建一个 if 语句来检查 [actual_index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L64-L64) 是否小于 0 或大于等于 [self.emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L39-L39) 列表的长度。如果是，则打印消息 [Invalid email number.\n](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L67-L67) 并使用 [return](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-lists/6852e41901d911133812e044.md#L31-L31) 退出方法。

# --hints--

你应该从 [index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L59-L59) 参数中减去 1，并将其存储在名为 [actual_index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L64-L64) 的变量中。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_body().find_variable("actual_index").is_equivalent("actual_index = index - 1")`))
    }
})
```

你应该创建一个 if 语句来检查 [actual_index](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L64-L64) 是否小于 0 或大于等于 [self.emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L39-L39) 列表的长度。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_conditions()[0].is_equivalent("actual_index < 0 or actual_index >= len(self.emails)")`))
    }
})
```

在 [if](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-lists/6852e41901d911133812e044.md#L23-L23) 语句中，你应该打印消息 [Invalid email number.\n](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L67-L67) 并使用 [return](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-lists/6852e41901d911133812e044.md#L31-L31) 退出方法。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[0].is_equivalent("print('Invalid email number.\\\\n')")`))
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

--fcc-editable-region--
    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return

--fcc-editable-region--
```
