---
id: 641da465273051435d332b15
title: æ­¥éª¤4
challengeType: 0
dashedName: step-4
---

# --description--

ä½ ä»Ž `GET` è¯·æ±‚ä¸­èŽ·å¾—çš„æ•°æ®æœ€åˆæ˜¯ä¸å¯ç”¨çš„ã€?ä¸ºäº†ä½¿æ•°æ®å¯ç”¨ï¼Œä½ å¯ä»¥ä½¿ç”?`Response` å¯¹è±¡ä¸Šçš„ `.json()` æ–¹æ³•å°†å…¶è§£æžä¸?JSONã€?å¦‚æžœä½ åœ¨æµè§ˆå™¨æŽ§åˆ¶å°ä¸­å±•å¼€ `Response` å¯¹è±¡çš?`Prototype`ï¼Œä½ å°†åœ¨é‚£é‡Œçœ‹åˆ° `.json()` æ–¹æ³•ã€?

åˆ é™¤ `console.log(res)` å¹¶éšå¼è¿”å›?`res.json()`ã€?

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

ä½ åº”è¯¥åˆ é™?`console.log(res)`ã€?

```js
assert.notMatch(code, /\.then\(\s*(\(\s*res\s*\)|res)\s*=>\s*console\.log\(\s*res\s*\)\s*\)/)
```

ä½ åº”è¯¥åœ¨ `.then()` ä¸­éšå¼è¿”å›?`res.json()`ã€?

```js
assert.match(code, /\.then\(\s*(\(\s*res\s*\)|res)\s*=>\s*res\.json\(\s*\)\s*\)/)
```

`.then()` åŽé¢ä¸åº”è¯¥æœ‰åˆ†å·ã€?

```js
assert.notMatch(code, /\.then\(\s*(\(\s*res\s*\)|res)\s*=>\s*res\.json\(\s*\)\s*\)\s*;/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">MasterPuti News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

fetch('https://cdn.MasterPuti.org/curriculum/news-author-page/authors.json')
--fcc-editable-region--
  .then((res) => console.log(res))

--fcc-editable-region--
```

