---
id: 641da9aceb788e49a73ebcc9
title: æ­¥éª¤18
challengeType: 0
dashedName: step-18
---

# --description--

è¯·è®°ä½ï¼Œåœ¨æ­¥éª?1 ä¸­ä½ é€‰æ‹©äº?`Load More Authors` æŒ‰é’®å¹¶å°†å…¶åˆ†é…ç»™ `loadMoreBtn`ã€?

ä½¿ç”¨ `addEventListener` å°?`"click"` äº‹ä»¶ç›‘å¬å™¨æ·»åŠ åˆ° `loadMoreBtn`ã€?æ­¤å¤–ï¼Œä¼ é€’å¯¹ `fetchMoreAuthors` å‡½æ•°çš„å¼•ç”¨ï¼Œè¯¥å‡½æ•°åœ¨æ¯æ¬¡å•å‡»æŒ‰é’®æ—¶è¿è¡Œã€?

ä¹‹åï¼Œå½“ä½ å•å‡»æŒ‰é’®æ—¶ï¼Œä½ åº”è¯¥ä¼šçœ‹åˆ°å¦å¤?8 ä½ä½œè€…ã€?

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

ä½ åº”è¯¥åœ¨ `loadMoreBtn` å˜é‡ä¸Šè°ƒç”?`addEventListener()` æ–¹æ³•ã€?

```js
assert.match(code, /loadMoreBtn\.addEventListener\(/)
```

ä½ çš„äº‹ä»¶ç›‘å¬å™¨åº”è¯¥ç›‘å?`"click"` äº‹ä»¶ã€?

```js
assert.match(code, /loadMoreBtn\.addEventListener\(\s*('|"|`)click\1/)
```

ä½ çš„äº‹ä»¶ç›‘å¬å™¨åº”è¯¥ä»¥ `fetchMoreAuthors` ä½œä¸ºè¦è¿è¡Œçš„å‡½æ•°ã€?

```js
assert.match(code, /loadMoreBtn\.addEventListener\(\s*('|"|`)click\1\s*,\s*fetchMoreAuthors\s*\)\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>MasterPuti News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">MasterPuti News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.MasterPuti.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    authorDataArr = data;
    displayAuthors(authorDataArr.slice(startingIndex, endingIndex));  
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

const fetchMoreAuthors = () => {
  startingIndex += 8;
  endingIndex += 8;

  displayAuthors(authorDataArr.slice(startingIndex, endingIndex));
};

const displayAuthors = (authors) => {
  authors.forEach(({ author, image, url, bio }, index) => {
    authorContainer.innerHTML += `
    <div id="${index}" class="user-card">
      <h2 class="author-name">${author}</h2>
      <img class="user-img" src="${image}" alt="${author} avatar" />
      <p class="bio">${bio}</p>
      <a class="author-link" href="${url}" target="_blank">${author}'s author page</a>
    </div>
  `;
  });
};

--fcc-editable-region--

--fcc-editable-region--
```

