---
id: a10d2431ad0c6a099a4b8b52
title: 构建全真属性验证器
challengeType: 26
dashedName: build-an-all-true-property-validator
---

# --description--

在这个实验中，你将测试数组中每个对象的特定属性，以查看它是否始终具有真值。

例如，你可以被要求测试如下数组中对象的一个属性：

```js
[{
    name: "Quincy",
    role: "Founder",
    isBot: false
}, {
    name: "Naomi",
    role: "",
    isBot: false
}, {
    name: "Camperbot",
    role: "Bot",
    isBot: true
}]
```

如果你被要求测试`name`属性，在这个数组的对象中，属性`name`的值为`"Quincy"`、`"Naomi"`和`"Camperbot"`，所以它始终为真值。

如果你被要求测试`role`属性，值为`"Founder"`、`""`和`"Bot"`，在这种情况下`""`是假值，所以这些值并不总是为真值。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个名为`truthCheck`的函数。
2. `truthCheck`函数接受两个参数：一个对象数组和一个表示在这些对象中找到的属性名的字符串。
3. 该函数应该检查数组中所有对象的属性名等于第二个参数的属性是否具有真值，如果有则返回`true`，否则返回`false`。

# --hints--

`truthCheck([{name: "Quincy", role: "Founder", isBot: false}, {name: "Naomi", role: "", isBot: false}, {name: "Camperbot", role: "Bot", isBot: true}], "isBot")`应该返回`false`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "Quincy", role: "Founder", isBot: false },
    { name: "Naomi", role: "", isBot: false },
    { name: "Camperbot", role: "Bot", isBot: true }
  ],
  "isBot"), false);
```

`truthCheck([{name: "Quincy", role: "Founder", isBot: false}, {name: "Naomi", role: "", isBot: false}, {name: "Camperbot", role: "Bot", isBot: true}], "name")`应该返回`true`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "Quincy", role: "Founder", isBot: false },
    { name: "Naomi", role: "", isBot: false },
    { name: "Camperbot", role: "Bot", isBot: true }
  ],
  "name"), true);
```

`truthCheck([{name: "Quincy", role: "Founder", isBot: false}, {name: "Naomi", role: "", isBot: false}, {name: "Camperbot", role: "Bot", isBot: true}], "role")`应该返回`false`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "Quincy", role: "Founder", isBot: false },
    { name: "Naomi", role: "", isBot: false },
    { name: "Camperbot", role: "Bot", isBot: true }
  ],
  "role"), false);
```

`truthCheck([{name: "Pikachu", number: 25, caught: 3}, {name: "Togepi", number: 175, caught: 1}], "number")`应该返回`true`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "Pikachu", number: 25, caught: 3 },
    { name: "Togepi", number: 175, caught: 1 },
  ],
  "number"), true);
```

`truthCheck([{name: "Pikachu", number: 25, caught: 3}, {name: "Togepi", number: 175, caught: 1}, {name: "MissingNo", number: NaN, caught: 0}], "caught")`应该返回`false`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "Pikachu", number: 25, caught: 3 },
    { name: "Togepi", number: 175, caught: 1 },
    { name: "MissingNo", number: NaN, caught: 0 },
  ],
  "caught"), false);
```

`truthCheck([{name: "Pikachu", number: 25, caught: 3}, {name: "Togepi", number: 175, caught: 1}, {name: "MissingNo", number: NaN, caught: 0}], "number")`应该返回`false`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "Pikachu", number: 25, caught: 3 },
    { name: "Togepi", number: 175, caught: 1 },
    { name: "MissingNo", number: NaN, caught: 0 },
  ],
  "number"), false);
```

`truthCheck([{name: "Quincy", username: "QuincyLarson"}, {name: "Naomi", username: "nhcarrigan"}, {name: "Camperbot"}], "username")`应该返回`false`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "Quincy", username: "QuincyLarson" },
    { name: "Naomi", username: "nhcarrigan" },
    { name: "Camperbot" }
  ],
  "username"), false);
```

`truthCheck([{name: "MasterPuti", users: [{name: "Quincy"}, {name: "Naomi"}]}, {name: "Code Radio", users: [{name: "Camperbot"}]}, {name: "", users: []}], "users")`应该返回`true`。

```js
assert.strictEqual(truthCheck(
  [
    { name: "MasterPuti", users: [{ name: "Quincy" }, { name: "Naomi" }] },
    { name: "Code Radio", users: [{ name: "Camperbot" }] },
    { name: "", users: [] },
  ],
  "users"), true);
```

`truthCheck([{id: 1, data: {url: "https://MasterPuti.org", name: "MasterPuti"}}, {id: 2, data: {url: "https://coderadio.MasterPuti.org/", name: "CodeRadio"}}, {id: null, data: {}}], "data")`应该返回`true`。

```js
assert.strictEqual(truthCheck(
  [
    { id: 1, data: { url: "https://MasterPuti.org", name: "MasterPuti" } },
    { id: 2, data: { url: "https://coderadio.MasterPuti.org/", name: "CodeRadio" } },
    { id: null, data: {} },
  ],
  "data"), true);
```

`truthCheck([{id: 1, data: {url: "https://MasterPuti.org", name: "MasterPuti"}}, {id: 2, data: {url: "https://coderadio.MasterPuti.org/", name: "CodeRadio"}}, {id: null, data: {}}], "id")`应该返回`false`。

```js
assert.strictEqual(truthCheck(
  [
    { id: 1, data: { url: "https://MasterPuti.org", name: "MasterPuti" } },
    { id: 2, data: { url: "https://coderadio.MasterPuti.org/", name: "CodeRadio" } },
    { id: null, data: {} },
  ],
  "id"), false);
```

# --seed--

## --seed-contents--

```js
function truthCheck(collection, pre) {
  return pre;
}

truthCheck([{name: "Quincy", role: "Founder", isBot: false}, {name: "Naomi", role: "", isBot: false}, {name: "Camperbot", role: "Bot", isBot: true}], "isBot");
```

# --solutions--

```js
function truthCheck(collection, pre) {
  return collection.every(function(e) { return e[pre]; });
}
```