---
id: 68552464ef378e1051a55c28
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在第二个 `div` 中，创建两个锚点元素。第一个应该有文本 `了解更多`，第二个应该有文本 `开始聆听`。

对于两个锚点元素的 `href` 属性，都将它们设置为 `#`。

# --hints--

你应该在第二个 `div` 元素内部创建两个 `a` 元素

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

你的第一个 `a` 元素应该在第二个 `div` 元素内部有文本 `了解更多` 和 `href` 属性 `#`。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "了解更多")
assert.equal(links[0]?.getAttribute("href"), "#")
```

你的第二个 `a` 元素应该在第二个 `div` 元素内部有文本 `开始聆听` 和 `href` 属性 `#`。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "开始聆听")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
      <div>
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```