---
id: 685523c05fb41d0848c4b52f
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在第一个 `div` 元素内部，创建一个文本为 `Soundflow` 的 `span` 元素，一个文本为 `发现新音乐` 的 `h1` 元素，以及一个文本为 `播放你喜欢的歌曲并发现新艺术家。` 的 `p` 元素。

# --hints--

你应该在第一个 `div` 元素内部创建一个文本为 `Soundflow` 的 `span` 元素。

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

你应该在第一个 `div` 元素内部创建一个文本为 `发现新音乐` 的 `h1` 元素。

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "发现新音乐")
```

你应该在第一个 `div` 元素内部创建一个文本为 `播放你喜欢的歌曲并发现新艺术家。` 的 `p` 元素。

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "播放你喜欢的歌曲并发现新艺术家。")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--

      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```