---
id: 6858f90c54647a48ddb8822d
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

此时，内容，即外部 `div` 内的两个 `div` 元素，无论屏幕大小如何，都会堆叠在一起。

这是因为你在包裹的 `div` 元素中有 `flex-col` 类。要解决这个问题，使用 `lg:` 前缀将 `flex-row` 类添加到 `div`。

这意味着在具有较大屏幕的设备上，弹性方向将是 `row`。要再次看到效果，请将浏览器调整为更大的宽度。

# --hints--

你的外部 `div` 元素应该具有实用类 `lg:flex-row`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
       <div>
         <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```