---
id: 6858ef974c8ad5dcd7abb9f1
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

`div` 内的内容可以更好地间隔和排列。对于其余的间距和排列，你需要使用 flexbox。

将类 `flex` 添加到 `div`，使两个子 `div` 元素并排显示。

同时，添加 `justify-around` 使内容周围和之间的空间均匀分布，并添加 `items-center` 使整个内容最终在 `div` 中居中。

# --hints--

你的外部 `div` 元素应该具有实用类 `flex`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

你的外部 `div` 元素应该具有实用类 `justify-around`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

你的外部 `div` 元素应该具有实用类 `items-center`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
       <div>
         <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```