---
id: 6858e6408a0999603b2a4fee
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

对于 `h1`，添加类 `font-bold` 和 `text-4xl`。这将使文本更粗、更大，并且增加行间距。

同时，添加类 `my-4` 来增加 `h1` 的上下边距 16 像素，或 `1rem`。

# --hints--

你的 `h1` 元素应该具有实用类 `font-bold`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

你的 `h1` 元素应该具有实用类 `text-4xl`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

你的 `h1` 元素应该具有实用类 `my-4`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>发现新音乐</h1>
         --fcc-editable-region--
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
       <div>
        <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
  </body>
</html>
```