---
id: 6858e9f87e4faa903267d810
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

由于外部 `div` 中的文本默认为黑色，因此不太容易看清。因此，将类 `text-white` 添加到 `div`。

# --hints--

你的外部 `div` 元素应该具有实用类 `text-white`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600">
      <div>
        <span>Soundflow</span>
        <h1>发现新音乐</h1>
        <p>播放你喜欢的歌曲并发现新艺术家。</p>
      </div>
       <div>
        <a href="#">了解更多</a>
        <a href="#">开始聆听</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```