---
id: 68e4465fba0bd256868b7deb
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

继续探索新的字符串方法吧。

创建一个名为`learningSentence` 的变量，并赋值为 `"I love learning!"`。

然后依次输出 `"Original learning sentence:"` 和`learningSentence`。保持耐心，你的进步在一点点积累。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该创建一个`learningSentence` 变量。

```js
assert.exists(learningSentence);
```

你的 `learningSentence` 变量应该是一个字符串。

```js
assert.isString(learningSentence);
```

你需要将字符串`"I love learning!"` 赋值给 `learningSentence`。

```js
assert.equal(learningSentence, "I love learning!");
```

你需要输出`"Original learning sentence:"`。

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

你还需要输出`learningSentence` 的内容，让它被看到。继续保持笑容！

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--
const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--