---
id: 68e42b62d7093345c6975f22
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

继续保持势头！在你的 `originalString` 变量下面，添加一个 `console.log()`，输出字符串 `"Original string:"`，像是在为原始内容挂上一块可爱的标签。

然后再在它下面添加一个 `console.log()`，输出变量 `originalString`，这样我们就能在控制台里看到它的闪亮登场。一步一步来，你已经做得很棒。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你需要在控制台输出 `"Original string:"` 这一行，给原始字符串一个暖心的开场白。

```js
assert.equal(getLogs()[0], "Original string:");
```

你还需要在控制台输出 `originalString` 变量，让它正式亮相。继续加油！

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
