---
id: 68e43e884ef09d4d2ab05631
title: 步骤 4
challengeType: 1
dashedName: step-4
---

# --description--

是时候检验刚才的替换成果啦！

先添加一个 `console.log()`，内容为 `"After using the replace() method:"`，像是在宣布魔法完成。

再添加一条输出 `replacedString` 的 `console.log()`。

看看控制台，你应该看到那句话已经变成了 `"I love dogs."`，原来的小猫被小狗温柔替换了。干得漂亮！

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你需要在控制台输出 `"After using the replace() method:"` 这一行，让结果仪式感满满。

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

你还需要输出 `replacedString` 变量，以展示成功的变换。继续保持这份节操！

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--
const replacedString = originalString.replace("cats", "dogs");

--fcc-editable-region--
```