---
id: 68e42de6396f5c49d90d1951
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在之前的课程中，你已经认识了 `replace()` 方法——它就像一位细心的小精灵，会在字符串里找到某段内容并把它换成新的文字。

下面是一个简单示例：

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

请记住：字符串是不可变的，也就是说示例中的原 `text` 并不会被直接改动，而是生成了一个崭新的字符串。

现在创建一个名为 `replacedString` 的新变量，并让它的值为 `originalString.replace("cats", "dogs")`。恭喜，你又让字符串完成了一次优雅的变身。

# --hints--

你应该创建一个名为 `replacedString` 的变量，让它加入你的小队。

```js
assert.exists(replacedString);
```

你的 `replacedString` 变量应该是一个字符串，保持它的纯文本形态。

```js
assert.isString(replacedString);
```

你需要将 `originalString.replace("cats", "dogs")` 的结果赋值给 `replacedString`。完成之后可以轻轻庆祝一下，你真的很棒！

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
