---
id: 687ad7be1ff45032ccf1d92f
title: Build a Theme Switcher
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

在本练习中，你将构建一个可以切换不同主题的应用：切换主题后，背景色会改变，同时显示对应的提示信息。

**目标：**实现下面的用户故事并通过所有测试，完成本练习。

**用户故事：**

1. 页面应包含一个文本为 `Switch Theme` 的 `button` 元素。
2. 该 `button` 应包含以下属性：
  - `id="theme-switcher-button"`
  - `aria-haspopup="true"`
  - `aria-expanded="false"`
  - `aria-controls="theme-dropdown"`
3. 页面应有一个 `ul` 元素，且包含属性：
  - `id="theme-dropdown"`
  - `role="menu"`
  - `aria-labelledby="theme-switcher-button"`
  - `hidden`（初始隐藏）
4. `#theme-dropdown` 内至少应有两个 `li`，每个 `li` 应有 `role="menuitem"`，并包含表示不同主题的文本。
5. 每个 `li` 的 `id` 应以 `theme-` 开头，并以对应主题名（小写）结尾，例如主题为 `Light` 则 `id="theme-light"`。
6. 页面应包含一个 `aria-live="polite"` 的元素，用于展示主题切换时的消息。
7. 应定义一个 `themes` 数组，至少包含两个对象，每个对象有 `name` 与 `message` 属性；`name` 对应 `li` 的文本（小写），`message` 在切换后显示。
8. 点击 `#theme-switcher-button` 时，应显示 `#theme-dropdown` 中的主题选项。
9. 当下拉显示时，按钮的 `aria-expanded` 应为 `"true"`。
10. 再次点击按钮时应关闭下拉，添加 `hidden` 属性，并将 `aria-expanded` 设为 `"false"`。
11. 选择某个主题后，应向 `body` 添加对应的 `theme-<name>` 类，并在 `aria-live` 区域显示该主题对应的消息。

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

下面是一些测试提示，帮助你逐步实现交互：

应存在 `button` 元素，文本为 `Switch Theme`。

```js
assert.exists(document.querySelector("button"));
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

应存在 `ul#theme-dropdown[role="menu"][aria-labelledby="theme-switcher-button"][hidden]`，且其中至少包含两个 `li[role="menuitem"]`，每个 `li` 的 `id` 以 `theme-` 开头并与其文本（小写）相对应。

```js
assert.exists(document.querySelector("ul"));
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
assert.equal(ul?.getAttribute("role"), "menu");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
assert.isTrue(ul?.hasAttribute("hidden"));
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
listItems.forEach(li => { assert.equal(li?.getAttribute("role"), "menuitem"); assert.isTrue(li?.textContent.trim().length > 0); });
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());
listItems.forEach((li, index) => { assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`); });
```

页面应包含 `aria-live="polite"` 的元素用于显示消息；并定义 `themes` 数组，至少包含两个对象（带 `name` 与 `message`）。

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
assert.isArray(themes);
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => { assert.isObject(theme); assert.property(theme, "name"); assert.property(theme, "message"); assert.isString(theme.name); assert.isNotEmpty(theme.name); assert.isString(theme.message); assert.isNotEmpty(theme.message); });
const namesArr = themes?.map(({name}) => name);
assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listTitles.includes(name)));
```

点击按钮应显示/隐藏下拉并更新 `aria-expanded`；选择主题后应给 `body` 添加相应类并在 `aria-live` 区域展示消息。

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
resetThemeSwitch()
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
resetThemeSwitch()
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
resetThemeSwitch()
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");

themeDropdownBtn.dispatchEvent(new Event("click"));
const clickedItem = listItems[1]; 
clickedItem?.dispatchEvent(new Event("click"), {bubbles: true});
const expectedClass = clickedItem?.id; 
const actualClass = document.body.className;
assert.equal(actualClass, expectedClass);

const liveRegion = document.querySelector('[aria-live="polite"]');
themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click"), {bubbles: true});
const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);
assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```
```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>

  <div id="status" aria-live="polite"></div>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);

    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
