---
id: 68ba9c4f1688914d72876097
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `allow` 属性的现有值中添加 `encrypted-media`、`gyroscope` 和 `web-share`。

这三项将允许使用加密媒体扩展来保护视频，让 iframe 在需要时可以弹出到画中画模式，并允许通过设备的原生分享对话框分享 iframe 内容。

# --hints--

你的 `iframe` 元素的 `allow` 属性应该有 `encrypted-media` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const encryptedMedia = iframeElAllowAttr.trim().split(' ')[3]
assert.strictEqual(encryptedMedia, 'encrypted-media')
```

你的 `iframe` 元素的 `allow` 属性应该有 `gyroscope` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const gyroscope = iframeElAllowAttr.trim().split(' ')[4]
assert.strictEqual(gyroscope, 'gyroscope')
```

你的 `iframe` 元素的 `allow` 属性应该有 `web-share` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const webShare = iframeElAllowAttr.trim().split(' ')[5]
assert.strictEqual(webShare, 'web-share')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe 视频展示</h1>
--fcc-editable-region--
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```