---
id: 68ba9c4f1688914d72876099
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

最后但同样重要的是，你要添加的属性是 `allowfullscreen`。顾名思义，它允许视频以全屏模式观看。

这样，工作坊就完成了！

# --hints--

你的 `iframe` 元素应该有一个 `allowfullscreen` 属性。

```js
const iframeEl = document.querySelector('iframe')
assert.exists(iframeEl?.getAttribute('allowfullscreen'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe 视频展示</h1>
--fcc-editable-region--
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
      referrerpolicy="strict-origin-when-cross-origin"
      allowfullscreen
    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```