---
id: 68b94f23f7b3ee14078c89c5
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将使用 `iframe` 元素来显示视频。基本的 HTML 模板已经为你准备好了。

首先创建一个 `h1` 元素，文本为 `iframe 视频展示`。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'))
```

你的 `h1` 元素应该有文本 `iframe 视频展示`。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe 视频展示')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
<h1>iframe 视频展示</h1>
--fcc-editable-region--
  </body>
</html>
```