---
id: 68ba9c4f1688914d72876096
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

其中一个属性是 `allow`。它就像一个权限列表，告诉浏览器 `iframe` 被允许使用哪些功能。

下面是一个带有 `allow` 属性的 `iframe` 元素示例：

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

添加 `allow` 属性，值为 `accelerometer`、`autoplay` 和 `clipboard-write`。

`accelerometer` 让 `iframe` 能够使用运动传感器，这样它就可以检测设备倾斜和旋转等操作。`autoplay` 让视频能够自动播放，而 `clipboard-write` 让 iframe 能够向用户的剪贴板写入数据。

# --hints--

你的 `iframe` 元素应该有一个 `allow` 属性，值为 `accelerometer`。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const accelerometer = iframeElAllowAttr.trim().split(' ')[0]
assert.strictEqual(accelerometer, 'accelerometer')
```

你的 `iframe` 元素的 `allow` 属性应该有 `autoplay` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const autoplay = iframeElAllowAttr.trim().split(' ')[1]
assert.strictEqual(autoplay, 'autoplay')
```

你的 `iframe` 元素的 `allow` 属性应该有 `clipboard-write` 作为其值之一。

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const clipboardWrite = iframeElAllowAttr.trim().split(' ')[2]
assert.strictEqual(clipboardWrite, 'clipboard-write')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe 视频展示</h1>
--fcc-editable-region--
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write"
    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```