---
id: 68ba9c4f1688914d72876095
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

`iframe` 元素还可以使用 `src` 属性，其值表示要显示的资源的 URL 或路径。

给你的 `iframe` 元素添加一个 `src` 属性，值为 `https://www.youtube.com/embed/I0_951_MPE0`。

此时，你应该能看到页面上显示的视频，但还需要添加一些其他属性。

# --hints--

你的 `iframe` 元素应该有一个 `src` 属性，值为 `https://www.youtube.com/embed/I0_951_MPE0`。

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe 视频展示</h1>
--fcc-editable-region--
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```