---
id: 686b9720ee1d032bd77a480a
title: Debug an ISBN Validator
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN（国际标准书号）是分配给商业书籍的唯一标识符。它可以是10位或13位数字长，最后一位是根据其他数字计算得出的校验位。

Camperbot试图构建他们自己的ISBN验证器。然而，他们在过程中犯了一些错误。

在这个实验中，你将修复现有代码并使其正常运行。

预期行为：当用户运行程序时，它将显示提示`Enter ISBN and length:`。用户可以输入他们想要验证的ISBN代码，格式为`ISBN,length`。ISBN代码不应包含连字符，后跟其长度（`10`或`13`），用逗号分隔。

示例输入：`1530051126,10`用于ISBN-10代码。`9781530051120,13`用于ISBN-13代码。

如何找到校验位：

- 你不必了解此实验中的详细计算逻辑。函数`calculate_check_digit_10`和`calculate_check_digit_13`将处理计算并返回预期的校验位字符串。
- ISBN-10代码的校验位可以是`0`到`9`的数字或大写字母`X`。
- ISBN-13代码的校验位可以是`0`到`9`的数字。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该修复当前代码中的`IndentationError`。
1. 即使用户没有输入逗号分隔的值，程序也应该处理`IndexError`而不崩溃。
1. 当用户没有输入逗号分隔的值时，他们应该在控制台中看到消息`Enter comma-separated values.`，然后程序应该终止。
1. 即使用户为长度输入非数字值，程序也应该处理`ValueError`而不崩溃。
1. 当用户为长度输入非数字值时，他们应该在控制台中看到消息`Length must be a number.`，然后程序应该终止。
1. 你应该修复`validate_isbn`函数中的差一错误。
1. 你应该修复当前代码中用户输入有效ISBN代码时发生的`TypeError`。
1. 你应该修复当前代码中用户输入有效ISBN代码时发生的`IndexError`。
1. 即使用户输入包含非数字字符的不正确ISBN代码，程序也应该处理发生的`ValueError`而不崩溃。
1. 当用户输入包含非数字字符的不正确ISBN代码时，他们应该在控制台中看到消息`Invalid character was found.`。
1. 当用户输入`1530051126,10`时，他们应该看到消息`Valid ISBN Code.`
1. 当用户输入`9781530051120,13`时，他们应该看到消息`Valid ISBN Code.`

**重要**：你需要注释掉全局空间中的`main()`调用，以便测试能正常运行。

当你完成项目时，用户应该根据他们输入的值看到以下消息。

| ISBN代码                                                    | 长度                                        | 消息                                                                                                                     | 示例输入                         |
| ------------------------------------------------------------ | --------------------------------------------- | --------------------------------------------------------------------------------------------------------------------------- | ------------------------------------- |
| 有效                                                        | 有效                                         | `Valid ISBN Code.`                                                                                                          | `1530051126,10`                       |
| 无效数字                                               | 有效                                         | `Invalid ISBN Code.`                                                                                                        | `1530051125,10`                       |
| 与指定长度不匹配或留空                | 有效                                         | `ISBN-10 code should be 10 digits long.` 或 `ISBN-13 code should be 13 digits long.`，取决于他们输入的长度。 | `9781530051120,10` 或 `1530051126,13` |
| 包含非数字字符（校验位除外） | 有效                                         | `Invalid character was found.`                                                                                              | `15-0051126,10`                       |
| 任何                                                          | 无效数字                                | `Length should be 10 or 13.`                                                                                                | `1530051126,9`                        |
| 任何                                                          | 包含非数字字符或留空 | `Length must be a number.`                                                                                                  | `1530051125,A`                        |
| 非逗号分隔                                          | 非逗号分隔                           | `Enter comma-separated values.`                                                                                             | `1530051125`                          |

如果你想，可以使用以下值手动测试你的代码。

有效ISBN-10代码的示例输入：`1530051126,10` `9971502100,10` `080442957X,10`

有效ISBN-13代码的示例输入：`9781530051120,13` `9781947172104,13`

# --hints--

你应该注释掉对`main`函数的调用，以允许其余测试正常工作。

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

你应该有一个`validate_isbn`函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

你应该有一个`calculate_check_digit_10`函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

你应该有一个`calculate_check_digit_13`函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

你应该有一个`main`函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

当用户输入不是逗号分隔的值时，你应该在控制台中看到消息`Enter comma-separated values.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

当用户为长度输入非数字值时，你应该在控制台中看到消息`Length must be a number.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

当用户输入包含非数字字符的不正确ISBN代码时，你应该在控制台中看到消息`Invalid character was found.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

当用户输入`1530051126,10`时，你应该在控制台中看到消息`Valid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入`9781530051120,13`时，你应该看到消息`Valid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入`1530051125,10`时，你应该看到消息`Invalid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

当用户输入`9781530051120,10`时，你应该看到消息`ISBN-10 code should be 10 digits long.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

当用户输入`1530051126,13`时，你应该看到消息`ISBN-13 code should be 13 digits long.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

当用户输入`15-0051126,10`时，你应该看到消息`Invalid character was found.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

当用户输入`1530051126,9`时，你应该看到消息`Length should be 10 or 13.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

当用户输入`1530051125,A`时，你应该看到消息`Length must be a number.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

当用户输入`1530051125`时，你应该看到消息`Enter comma-separated values.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

当用户输入`9971502100,10`时，你应该看到消息`Valid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入`080442957X,10`时，你应该看到消息`Valid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入`9781947172104,13`时，你应该看到消息`Valid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)

print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    # 修复这里的错误
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    # 修复这里的差一错误
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # 从其他数字计算校验位
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # 检查给定的校验位是否与计算的校验位匹配
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')

def calculate_check_digit_10(main_digits_list):
    # 注意：你不必完全理解这个函数中的逻辑。
    digits_sum = 0
    # 将前9位数字分别乘以其对应的权重（10到2）并求和
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # 找到将总和除以11的余数，然后从11中减去它
    result = 11 - digits_sum % 11
    # 计算结果可以是1到11之间的数字。
    # 如果结果是11，使用0。
    # 如果结果是10，使用大写X。
    # 对于其他数字，使用原值。
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit

def calculate_check_digit_13(main_digits_list):
    # 注意：你不必完全理解这个函数中的逻辑。
    digits_sum = 0
    # 将前12位数字交替乘以1和3（从1开始），然后求和
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # 找到将总和除以10的余数，然后从10中减去它
    result = 10 - digits_sum % 10
    # 计算结果可以是1到10之间的数字。
    # 如果结果是10，使用0。
    # 对于其他数字，使用原值。
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit

def main():
    user_input = input('Enter ISBN and length: ')
    # 处理用户没有输入逗号分隔值的情况
    values = user_input.split(',')
    isbn = values[0]
    # 处理用户为长度输入非数字值的情况
    length = int(values[1])
    if length == 10 or length == 13:
        # 修复这里的缩进错误
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')

# main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # 从其他数字计算校验位
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # 检查给定的校验位是否与计算的校验位匹配
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # 注意：你不必完全理解这个函数中的逻辑。

    digits_sum = 0
    # 将前9位数字分别乘以其对应的权重（10到2）并求和
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # 找到将总和除以11的余数，然后从11中减去它
    result = 11 - digits_sum % 11

    # 计算结果可以是1到11之间的数字。
    # 如果结果是11，使用0。
    # 如果结果是10，使用大写X。
    # 对于其他数字，使用原值。
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # 注意：你不必完全理解这个函数中的逻辑。

    digits_sum = 0
    # 将前12位数字交替乘以1和3（从1开始），然后求和
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # 找到将总和除以10的余数，然后从10中减去它
    result = 10 - digits_sum % 10

    # 计算结果可以是1到10之间的数字。
    # 如果结果是10，使用0。
    # 对于其他数字，使用原值。
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```