---
id: 579e2a2c335b9d72dd32e05c
title: 实现切片和拼接算法
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

在这个实验中，你需要创建一个算法来合并两个数组。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 创建一个 [frankenSplice](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90) 函数，该函数接受两个数组和一个索引。
1. 按顺序将第一个数组的每个元素复制到第二个数组中，从给定的索引开始，并返回结果数组。
1. 函数运行后，输入数组应保持不变。

# --hints--

[frankenSplice([1, 2, 3], [4, 5], 1)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90) 应返回 [4, 1, 2, 3, 5](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90)。

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

[frankenSplice([1, 2], ["a", "b"], 1)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90) 应返回 ["a", 1, 2, "b"](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90)。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

[frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90) 应返回 ["head", "shoulders", "claw", "tentacle", "knees", "toes"](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90)。

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

第一个数组的所有元素都应按原始顺序添加到第二个数组中。[frankenSplice([1, 2, 3, 4], [], 0)](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90) 应返回 [1, 2, 3, 4](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-slice-and-splice/579e2a2c335b9d72dd32e05c.md#L84-L90)。

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

函数运行后，第一个数组应保持不变。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

函数运行后，第二个数组应保持不变。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
