---
id: a39963a4c10bc8b4d4f06d7e
title: 实现一个值移除函数
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

在这个实验中，你将创建一个函数，它以一个初始数组作为第一个参数，后面跟着一个或多个额外的参数，这些参数代表要移除的值。

**目标**：完成下面的用户故事并通过所有测试来完成这个实验。

**用户故事：**

1. 你应该有一个 `destroyer` 函数，它接受一个数组和一个或多个额外的参数。
1. `destroyer` 函数应该返回一个新数组，排除第一个参数中与后续任何参数匹配的所有元素。
1. 函数必须接受不定数量的参数。

# --hints--

你应该有一个 `destroyer` 函数。

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` 应该返回 `[1, 1]`。

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` 应该返回 `[1, 5, 1]`。

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` 应该返回 `[1]`。

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` 应该返回 `[]`。

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` 应该返回 `["hamburger"]`。

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` 应该返回 `[12, 92, 65]`。

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```