---
id: 5688e62ea601b2482ff8422b
title: 构建一个资料查找器
challengeType: 26
dashedName: lab-profile-lookup
---

# --description--

在这个实验中，您将构建一个资料查找器，用于查找联系人列表中人员的信息。

对于这个例子，假设有一个名为 Akira Laine 的联系人，`lookUpProfile("Akira", "lastName")` 应该返回 `Laine`。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 您应该创建一个名为 `lookUpProfile` 的函数，该函数接受姓名和属性作为参数。
2. 您应该从提供的 `contacts` 数组中检索联系人信息。
3. 如果函数接收到联系人姓名且相关联系人上存在该属性，则应返回该属性的值。
4. 如果传递给函数的姓名与联系人数组中的任何联系人都不匹配，则函数应返回 `"No such contact"`。
5. 如果在找到的联系人上不存在该属性，则函数应返回 `"No such property"`。

# --before-each--

```js
contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];
```

# --hints--

`contacts` 数组应该仍然存在于全局作用域中。如果您删除了它，请重置课程以恢复它。

```js
assert.isArray(contacts);
```

您应该有一个名为 `lookUpProfile` 的函数。

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")` 应该返回一个字符串。

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")` 应该返回字符串 `Vos`

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")` 应该返回 `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` 应该返回一个数组

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")` 应该返回字符串 `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")` 应该返回字符串 `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")` 应该返回字符串 `No such property`

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```