---
id: 685cdabd4cded3b340193f72
title: 第 24 步
challengeType: 20
dashedName: step-24
---

# --description--

现在为电子邮件显示添加主题行。在 [display_full_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L31-L31) 方法中，添加一个打印语句，以以下格式显示电子邮件的主题：

- [Subject: subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685cdabd4cded3b340193f72.md#L25-L25) 其中 [subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 被替换为电子邮件的主题

# --hints--

你应该使用 f-string 来打印主题，使用 [Subject: {self.subject}](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685cdabd4cded3b340193f72.md#L25-L25)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Subject: {self.subject}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')

--fcc-editable-region--

class User:
```
