---
id: 688754f3494b17e0b352d67d
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

现在是时候测试我们完整的电子邮件系统了！让我们创建一些用户，看看电子邮件功能的实际效果。

创建两个 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 对象：名为 ["Alice"](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L21-L21) 的 [alice](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L21-L21) 和名为 ["Bob"](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L21-L21) 的 [bob](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L22-L22)。这将演示如何创建用户以及用户之间如何交互。

# --hints--

你应该创建一个名为 [alice](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L21-L21) 的 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 对象，名称为 ["Alice"](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L21-L21)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

你应该创建一个名为 [bob](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L22-L22) 的 [User](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f11d412c29665985a915.md#L21-L21) 对象，名称为 ["Bob"](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/688754f3494b17e0b352d67d.md#L21-L21)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
