---
id: 68b9350183ff05bd106c79e4
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

现在，测试你在 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 类中创建的 [mark_as_read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25) 方法。在现有的 [Email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L17-L17) 对象 [email_obj](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24) 上使用该方法来更改状态。将电子邮件标记为已读后，打印 [email_obj](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24) 的 [read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L22-L22) 属性以确认更改。

# --hints--

你应该在 [email_obj](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24) 上调用 [mark_as_read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25) 方法将电子邮件标记为已读。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

你应该在调用 [mark_as_read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ef95ccc27660275692cd.md#L25-L25) 方法后打印 [email_obj](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L24-L24) 的 [read](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852ea0c3cadb54e897acdac.md#L22-L22) 属性。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
