---
id: 68b9366fae8b13c1278761b1
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

你的收件箱需要一种接收新邮件的方式。当有人向用户发送电子邮件时，它应该被添加到他们的收件箱中。

在你的 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 类中添加一个名为 [receive_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L41-L41) 的方法。该方法接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 和一个电子邮件对象 [email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L41-L41) 作为参数。在方法体中，使用 [append](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-lists/6852e41901d911133812e044.md#L19-L19) 方法将 [email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L41-L41) 添加到 [emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L39-L39) 列表中。

# --hints--

你应该在 [Inbox](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L37-L37) 类中定义一个名为 [receive_email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L41-L41) 的方法，该方法接受 [self](file:///F:/code/blocks_Chinese/blocks_english/lecture-working-with-classes-and-objects/6852d829161bd898603c0643.md#L35-L40) 和 [email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L41-L41) 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

你应该将 [email](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L41-L41) 添加到 [self.emails](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f34ff93594934f6636d9.md#L39-L39)。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []


--fcc-editable-region--
```
