---
id: 68bf0f09b7ec33be465c0c35
title: 第 29 步
challengeType: 20
dashedName: step-29
---

# --description--

在 [__str__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L49-L49) 方法的末尾，以以下格式返回电子邮件摘要：

```py
[status] From: sender | Subject: subject.
```

其中 [status](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/685cdfd1932d8ac6ad986813.md#L25-L25) 是电子邮件的状态，[sender](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 是发送者的名字，[subject](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6852f01116912861c7d1b0e6.md#L21-L21) 是电子邮件的主题。

# --hints--

[__str__](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L49-L49) 方法应该返回 [f"[{status}] From: {self.sender.name} | Subject: {self.subject}"](file:///F:/code/blocks_Chinese/blocks_english/workshop-email-simulator/6853f4584f508e95cc7a9f4d.md#L50-L50)。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_return("f\\"[{status}] From: {self.sender.name} | Subject: {self.subject}\\"")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        status = 'Read' if self.read else 'Unread'

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
