---
id: 686386e76197458b7dfd52ec
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--
在刚才的 `h2` 下面添加一个 `p`，文字内容为 `$0`。接着在这个 `p` 里、紧跟 `$0` 后面再放一个 `span`，写 `/month`。这样我们的免费方案价格标签就萌萌地亮相啦～

# --hints--
需要一个文本为 `$0` 的 `p` 元素。

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

在该 `p` 里要有一个文本为 `/month` 的 `span`。

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--

            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
