---
id: 686386e76197458b7dfd52ee
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在这个 `ul` 里添加两个列表项。第一个 `li` 里先放一个 `span`，内容是实体 `&#10003;`，紧跟在这个 `span` 后（仍在同一个 `li` 内）加上文字 `Ad-supported streaming`。

第二个 `li` 同样先放一个内容为实体 `&#10003;` 的 `span`，然后在它后面加入文字 `Curated playlists`。两个勾勾让功能显得很靠谱，是不是很有成就感～

# --hints--

`ul` 里面应该正好有两个 `li`。

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

第一个 `li` 里需要一个包含实体 `&#10003;` 的 `span`。

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

在第一个 `li` 中，该 `span` 后面紧跟文字 `Ad-supported streaming`。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

第二个 `li` 里也要有一个文本为实体 `&#10003;` 的 `span`。

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

在第二个 `li` 中，该 `span` 后面紧跟文字 `Curated playlists`。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--

            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
