---
id: 686383abac6c625a84458ea9
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--
在这个可爱的小工作坊里，你会通过使用 Tailwind CSS 提供的网格类（grid utilities）来动手构建一个“订阅价格组件”，更深入地掌握 Tailwind 的用法。冲鸭～

基础的 HTML 模板已经准备好啦。第一步：创建一个 `h1` 标题，内容写上 `Choose your listening plan`，让页面自带仪式感！

# --hints--
你需要创建一个 `h1` 元素。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素文本应该是 `Choose your listening plan`。

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
