---
id: 686386e76197458b7dfd52eb
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--
在刚才的 `h1` 下面创建一个外层 `div`。然后在里面再套一个 `div`，再往里继续嵌一个第三层 `div`（小盒子叠叠乐！）。

最里层的那个 `div` 里添加一个 `h2`，文本写 `Listener`，它就是第一个订阅方案的标题啦。继续加油，你做得很好～

# --hints--
需要创建一个 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

需要在外层 `div` 里再创建一个 `div`。

```js
assert.exists(document.querySelector("div > div"));
```

再在第二层 `div` 里嵌一个第三层 `div`。

```js
assert.exists(document.querySelector("div > div > div"));
```

在最内层 `div` 中放一个文本为 `Listener` 的 `h2`。

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
