---
id: 6842a6cd9836f0114a5b7a8a
title: Build a Number Pattern Generator
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

在这个实验中，您将通过构建一个创建数字模式的小应用程序来练习Python的基础知识。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 您应该定义一个名为[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)的函数，该函数接受一个参数n（代表一个正整数）。
2. [number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)应该使用[for](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L125-L125)循环。
3. [number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)(n)应该返回一个字符串，包含从1开始到n（包含）的所有整数，用空格分隔。例如，[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)(4)应该返回字符串[1 2 3 4](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L49-L49)。
4. 如果传递给函数的参数不是整数值，函数应该返回[Argument must be an integer value.](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L87-L87)。
5. 如果传递给函数的参数小于1，函数应该返回[Argument must be an integer greater than 0.](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L101-L101)。

# --hints--

您应该有一个[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)函数。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)函数应该有一个名为n的参数。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)(4)应返回[1 2 3 4](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L49-L49)。

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)(12)应返回[1 2 3 4 5 6 7 8 9 10 11 12](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L59-L59)。

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)应该为任何正整数返回一个用空格分隔的数字列表。

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2

    `)
})
```

当传递一个非整数值时，[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)应该返回[Argument must be an integer value.](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L87-L87)。

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

当传递一个非正整数时，[number_pattern](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L121-L129)应该返回[Argument must be an integer greater than 0.](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-number-pattern-generator/6842a6cd9836f0114a5b7a8a.md#L101-L101)。

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."
