---
id: 56533eb9ac21ba0edf2244cf
title: 构建唱片集
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

你正在创建一个有助于维护音乐专辑集的函数。该集合被组织为一个包含多个专辑的对象，这些专辑也是对象。每个专辑在集合中都以唯一的 [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.typedarray.d.ts#L46-L46) 作为属性名来表示。在每个专辑对象中，有各种属性描述有关专辑的信息。并非所有专辑都有完整的信息。

[updateRecords](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 函数接受 4 个参数，由以下函数参数表示：

-   [records](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) - 包含多个单独专辑的对象
-   [id](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/node_modules/.pnpm/@types+node@18.19.31/node_modules/@types/node/globals.typedarray.d.ts#L46-L46) - 代表 [records](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 对象中特定专辑的数字
-   [prop](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) - 表示要更新的专辑属性名称的字符串
-   [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) - 包含用于更新专辑属性的信息的字符串

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你的函数必须始终返回整个 [records](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 对象。
2. 如果 [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 是空字符串，则从专辑中删除给定的 [prop](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 属性。
3. 如果 [prop](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 不是 [tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 且 [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 不是空字符串，则将 [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 分配给该专辑的 [prop](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176)。
4. 如果 [prop](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 是 [tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 且 [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 不是空字符串，但专辑没有 [tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 属性，则创建一个空数组并将 [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 添加到其中。
5. 如果 [prop](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 是 [tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 且 [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 不是空字符串，则将 [value](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L176) 添加到专辑现有的 [tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 数组的末尾。

**注意：**测试中使用了 [recordCollection](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L129-L147) 对象的副本。你的函数不应直接引用 [recordCollection](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L129-L147) 对象，只能引用函数参数。

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

你应该有一个 [updateRecords](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 函数。

```js
assert.isFunction(updateRecords);
```

在 [updateRecords(recordCollection, 5439, "artist", "ABBA")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 之后，[artist](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L132-L132) 应该是字符串 [ABBA](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L132-L132)

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

在 [updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 之后，[tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 应该将字符串 [Take a Chance on Me](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 作为最后一个也是唯一的元素。

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

在 [updateRecords(recordCollection, 2548, "artist", "")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 之后，[artist](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L132-L132) 不应被设置

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

在 [updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 之后，[tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 应该将字符串 [Addicted to Love](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 作为最后一个元素。

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

在 [updateRecords(recordCollection, 2468, "tracks", "Free")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 之后，[tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 应该将字符串 [1999](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L138-L138) 作为第一个元素。

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

在 [updateRecords(recordCollection, 2548, "tracks", "")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 之后，[tracks](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L133-L133) 不应被设置

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

在 [updateRecords(recordCollection, 1245, "albumTitle", "Riptide")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L176-L186) 之后，[albumTitle](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L131-L131) 应该是字符串 [Riptide](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-record-collection/56533eb9ac21ba0edf2244cf.md#L131-L131)

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
