---
id: 673b567e3ba535dda140d278
title: 构建一个投票系统
challengeType: 26
dashedName: build-a-voting-system
---

# --description--

在这个实验中，你将构建一个使用 `Map` 创建投票和使用 `Set` 防止重复投票的投票系统。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该初始化一个 `poll` 变量为一个新的 `Map` 对象。

2. 你应该有一个函数 `addOption`，它接受一个参数 `option`。

3. 在 `addOption` 函数中：

   - 如果 `option` 在投票中尚不存在，则应将其添加到投票中，并使用一个空的 `Set` 作为其值来跟踪投票者。你还应该返回消息 `选项 "<option>" 已添加到投票中。`
   
   - 如果 `option` 已经存在，则应返回消息 `选项 "<option>" 已存在。`。
   
   - 如果你尝试添加一个空选项，函数应该返回消息 `选项不能为空。`。

4. 你应该有一个函数 `vote`，它接受两个参数，`option`（要投票的选项）和 `voterId`（投票者的唯一ID）。

5. 在 `vote` 函数中：

   - 如果 `option` 在投票中不存在，函数应该返回消息 `选项 "<option>" 不存在。`。
   
   - 如果 `option` 存在，函数应该检查 `voterId` 是否已经为此 `option` 投过票。
   
   - 如果投票者已经投过票，函数应该返回消息 `投票者 <voterId> 已经为 "<option>" 投过票了。`
   
   - 如果投票者尚未投票，`voterId` 应该被添加到此选项的 `voters` 的 `Set` 中。函数应该返回消息 `投票者 <voterId> 为 "<option>" 投票了。`

6. 你的 `poll` 中应该至少有三个选项。

7. 你的 `poll` 应该至少有三票。

8. 你应该有一个函数 `displayResults`，它以以下格式返回投票结果：

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.
.
/*
示例输出

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

你应该有一个 `poll` 变量初始化为一个新的 `Map` 对象。

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

你应该定义一个函数 `addOption`，它接受一个参数 `option`。

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

你应该定义一个函数 `vote`，它接受两个参数，`option` 和 `voterId`。

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

你应该定义一个函数 `displayResults` 来显示投票结果。

```js
assert.isFunction(displayResults);
```

你的 `poll` 中应该至少有三个选项。

```js
assert.isAtLeast(poll.size, 3);
```

你的 `poll` 应该至少有三票。

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

你应该确保每个投票选项映射到一个 `Set` 对象。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` 应该返回 `选项 "Egypt" 已添加到投票中。`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), '选项 "Egypt" 已添加到投票中。');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

添加一个空选项应该返回 `"选项不能为空。"`

```js
assert.equal(addOption(''), '选项不能为空。');
```

当 `Turkey` 已经被添加时，`addOption("Turkey")` 应该返回 `选项 "Turkey" 已存在。`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), '选项 "Turkey" 已存在。');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

当 `Malaysia` 存在于投票选项中时，`vote("Malaysia", "traveler1")` 应该返回 `投票者 traveler1 为 "Malaysia" 投票了。`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    '投票者 traveler1 为 "Malaysia" 投票了。'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` 应该更新选项的投票者 `Set`。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

当 `traveler1` 再次尝试为 `Algeria` 投票时，`vote("Algeria", "traveler1")` 应该返回 `投票者 traveler1 已经为 "Algeria" 投过票了。`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    '投票者 traveler1 已经为 "Algeria" 投过票了。'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

重复投票不应该增加 `Set` 的大小。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

当 `Nigeria` 不在投票选项中时，`vote("Nigeria", "traveler2")` 应该返回 `选项 "Nigeria" 不存在。`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    '选项 "Nigeria" 不存在。'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

一个独特的选项应该能够收到多个投票。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), '投票者 traveler1 为 "Bali" 投票了。');
  assert.equal(vote('Bali', 'traveler2'), '投票者 traveler2 为 "Bali" 投票了。');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` 应该以正确的格式返回结果。

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return '选项不能为空。';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `选项 "${option}" 已添加到投票中。`;
  } else {
    return `选项 "${option}" 已存在。`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `选项 "${option}" 不存在。`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `投票者 ${voterId} 已经为 "${option}" 投过票了。`;
  } else {
    voters.add(voterId);
    return `投票者 ${voterId} 为 "${option}" 投票了。`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```