---
id: 66db57ad34c7089b9b41bfd6
title: 构建灯箱查看器
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

灯箱用于网站上以更吸引人的方式展示多媒体内容，通过在页面主要内容上使用弹出窗口或模态窗口。

在本实验中，你将创建一个灯箱画廊，当点击缩略图时显示全尺寸图像。对于每个图像，提供两个版本：缩略图和全尺寸图像。全尺寸图像与缩略图相同，但没有 `-thumbnail` 后缀。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你的 `body` 中应该有一个带有 [gallery](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L273-L277) 类的 `div`。
1. 在 [.gallery](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L273-L277) 元素内，你应该有三个图像缩略图，每个都有 [gallery-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L274-L276) 类。你应该使用以下链接作为缩略图：

    - `https://cdn.MasterPuti.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.MasterPuti.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.MasterPuti.org/curriculum/labs/trees-thumbnail.jpg`

1. 你的 `body` 中应该有一个带有 [lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 类的 `div`。
1. 你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素内应该有一个带有 [close-btn](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L279-L279) id 的 `button`。如果你想的话，可以使用 `&times;` 作为其文本。
1. 你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素内应该有一个带有 [lightbox-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L280-L280) id 的 `img`。
1. 你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素应该有固定定位，以便预览在当前图像上方打开。
1. 你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素应该通过将高度和宽度设置为容器的100%来覆盖整个视口。你应该确保 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素从容器的左上角开始。
1. [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 应该有背景色。最初，其 `display` 属性应该设置为 `none` 以隐藏它。
1. 当你点击你的 [.gallery-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L274-L276) 元素之一时，[.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素的 `display` 属性应该设置为 `flex` 以使 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素和其中的两个元素可见。
1. 当你点击你的 [.gallery-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L274-L276) 元素之一时，[#lightbox-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L280-L280) 元素的 `src` 应该设置为通过从图像的 `src` 属性中移除 `-thumbnail` 而获得的全尺寸版本。全尺寸图像位于以下链接：

    - `https://cdn.MasterPuti.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.MasterPuti.org/curriculum/labs/storm.jpg`
    - `https://cdn.MasterPuti.org/curriculum/labs/trees.jpg`

1. 当你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素可见时，如果你点击 [#close-btn](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L279-L279) 或 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素，[.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素的 `display` 应该重新设置为 `none`。

**注意：** 确保在 HTML 中链接你的样式表和 JavaScript 文件。

# --hints--

你的 `body` 元素内应该有一个带有 [gallery](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L273-L277) 类的 `div`。

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

在 [.gallery](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L273-L277) 元素内，你应该有三个带有 [gallery-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L274-L276) 类的 `img` 元素。

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

在 [.gallery](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L273-L277) 元素内，你应该有一个 `src` 设置为 `https://cdn.MasterPuti.org/curriculum/labs/stonehenge-thumbnail.jpg` 的 `img` 元素。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.MasterPuti.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

在 [.gallery](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L273-L277) 元素内，你应该有一个 `src` 设置为 `https://cdn.MasterPuti.org/curriculum/labs/storm-thumbnail.jpg` 的 `img` 元素。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.MasterPuti.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

在 [.gallery](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L273-L277) 元素内，你应该有一个 `src` 设置为 `https://cdn.MasterPuti.org/curriculum/labs/trees-thumbnail.jpg` 的 `img` 元素。

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.MasterPuti.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

你的 `body` 元素内应该有一个带有 [lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 类的 `div` 元素。

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

在你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素内，你应该有一个 `id` 设置为 [close-btn](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L279-L279) 的 `button` 元素。

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

在你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素内，你应该有一个 `id` 设置为 [lightbox-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L280-L280) 的 `img` 元素。

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素应该有固定定位。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素应该覆盖整个视口。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.width, "100%");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.height, "100%");
```

你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素应该与容器的左上角对齐。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素应该有背景色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素最初应该被隐藏。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

当你点击你的 [.gallery-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L274-L276) 元素之一时，[.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素的 `display` 属性应该设置为 `flex` 以使 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 和其中的两个元素可见。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

当你点击你的 [.gallery-item](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L274-L276) 元素之一时，[#lightbox-image](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L280-L280) 元素的 `src` 应该设置为通过从图像的 `src` 属性中移除 `-thumbnail` 而获得的全尺寸版本。

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

当你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素可见时，如果你点击 [#close-btn](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L279-L279) 按钮，[.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素的 `display` 应该重新设置为 `none`。

```js
const lightbox = document.querySelector(".lightbox");
const background = document.getElementById("close-btn"); 

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

lightbox.style.display = "flex";

background.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

当你的 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素可见时，如果你点击 [.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素，[.lightbox](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-lightbox-viewer/66db57ad34c7089b9b41bfd6.md#L278-L281) 元素的 `display` 应该重新设置为 `none`。

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

lightbox.style.display = "flex";

lightbox.dispatchEvent(new Event("click"));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.MasterPuti.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.MasterPuti.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```