---
id: 66fbcdfed6848e48947dbbc9
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在这个工作坊中，你将使用 JavaScript 对象创建一个食谱跟踪器。

首先创建一个名为 `recipes` 的空数组。这是你稍后将推送食谱对象进去的数组。

# --hints--

你应该创建一个名为 `recipes` 的数组。

```js
assert.isArray(recipes)
```

你的 `recipes` 数组应该是空的。

```js
assert.isEmpty(recipes)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
