---
id: 66fbcf750a62784cf11f562d
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

创建一个名为 `recipe1` 的对象。在 `recipe1` 对象内部，创建一个值为 `"Spaghetti Carbonara"` 的 `name` 属性。

同样在 `recipe1` 对象内部，创建一个值为数组的 `ingredients` 属性。数组中应该包含 `spaghetti`、`Parmesan cheese`、`pancetta` 和 `black pepper`。

创建一个值为数组的 `ratings` 属性。数组中应该包含 `4`、`5`、`4` 和 `5`。

# --hints--

你应该创建一个名为 `recipe1` 的对象。

```js
assert.isObject(recipe1)
```

你的 `recipe1` 对象应该有一个 `name` 属性。

```js
assert.property(recipe1, 'name');
```

你的 `name` 属性应该设置为 `Spaghetti Carbonara`。

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

你的 `recipe1` 对象应该有一个 `ingredients` 属性。

```js
assert.property(recipe1, 'ingredients');
```

你的 `ingredients` 属性应该有一个数组值。

```js
assert.isArray(recipe1.ingredients)
```

你的 `ingredients` 属性的数组值应该包含 `spaghetti`、`Parmesan cheese`、`pancetta` 和 `black pepper`。

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

你的 `recipe1` 对象应该有一个 `ratings` 属性。

```js
assert.property(recipe1, "ratings");
```

你的 `ratings` 属性应该有一个数组值。

```js
assert.isArray(recipe1.ratings)
```

你的 `ratings` 属性的数组值应该包含 `4`、`5`、`4` 和 `5`。

```js
assert.deepEqual(recipe1.ratings, [4, 5, 4, 5]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
