---
id: 66fbcf750a62784cf11f5631
title: 步骤 7
challengeType: 1
dashedName: step-7
---

# --description--

在接下来的几个步骤中，你将计算 `recipes` 数组中每个食谱的平均评分、总配料数和难度等级。

首先创建一个 `getAverageRating` 函数，它接受一个参数，即一个包含评分的数组。在函数内部，计算从传递给函数的数组中得出的平均评分。

你的 `getAverageRating` 函数必须返回一个数字。

# --hints--

你应该创建一个 `getAverageRating` 函数。

```js
assert.isFunction(getAverageRating)
```

你的 `getAverageRating` 函数应该有一个参数。

```js
assert.lengthOf(getAverageRating, 1);
```

你的 `getAverageRating` 函数应该返回一个数字。

```js
assert.isNumber(getAverageRating(recipe1.ratings))
```

`getAverageRating(recipe1.ratings)` 应该返回 `4.5`。

```js
assert.strictEqual(getAverageRating(recipe1.ratings), 4.5)
```

`getAverageRating(recipe2.ratings)` 应该返回 `4.75`。

```js
assert.strictEqual(getAverageRating(recipe2.ratings), 4.75)
```

`getAverageRating(recipe3.ratings)` 应该返回 `4`。

```js
assert.strictEqual(getAverageRating(recipe3.ratings), 4)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
