---
id: 66fbcf750a62784cf11f5632
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

创建一个 `getTotalIngredients` 函数，它接受一个参数，表示一个包含配料的数组，并返回从传递给函数的数组中的配料数量。

# --hints--

你应该创建一个 `getTotalIngredients` 函数。

```js
assert.isFunction(getTotalIngredients)
```

你的 `getTotalIngredients` 函数应该有一个参数。

```js
assert.lengthOf(getTotalIngredients, 1);
```

你的 `getTotalIngredients` 函数应该返回一个数字。

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients))
```

`getTotalIngredients(recipe1.ingredients)` 应该返回 `4`。

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` 应该返回 `5`。

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` 应该返回 `3`。

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
