---
id: 66fbcf750a62784cf11f5630
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

第三个 `recipe3` 对象已经为你填写好了。它具有与 `recipe1` 和 `recipe2` 相同的属性。

你现在应该将这三个对象推入 `recipes` 数组中。为此，你可以使用 `push()` 方法。

使用 `push()` 方法将所有食谱对象推入 `recipes` 数组中。确保按 `recipe1`、`recipe2` 和 `recipe3` 的顺序推入。

同时删除 `recipe1Name`、`recipe2Name`、`recipe1CookingTime` 和 `recipe2CookingTime` 变量，以及记录这些变量的 `console.log` 语句。

# --hints--

你应该删除 `recipe1Name` 变量。

```js
assert.throws(() => recipe1Name, ReferenceError);
```

你应该删除 `console.log(recipe1Name);` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

你应该删除 `recipe2Name` 变量。

```js
assert.throws(() => recipe2Name, ReferenceError);
```

你应该删除 `console.log(recipe2Name);` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2Name\s*\)\s*;?/);
```

你应该删除 `recipe1CookingTime` 变量。

```js
assert.throws(() => recipe1CookingTime, ReferenceError);
```

你应该删除 `console.log(recipe1CookingTime);` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1CookingTime\s*\)\s*;?/);
```

你应该删除 `recipe2CookingTime` 变量。

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

你应该删除 `console.log(recipe2CookingTime);` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

你应该将所有食谱对象推入 `recipes` 数组中。确保按声明的顺序推入它们。

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2Name = recipe2.name;
console.log(recipe2Name);

const recipe1CookingTime = recipe1.cookingTime;
console.log(recipe1CookingTime);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);
--fcc-editable-region--
```
