---
id: 66fbcf750a62784cf11f562e
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

创建一个具有以下属性和值的 `recipe2` 对象：

| 键名                | 值                                                                   |
| ------------------ | ----------------------------------------------------------------------- |
| `name`             | `Chicken Curry`                                                         |
| `ingredients`      | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `ratings`          | `[4, 5, 5, 5]`                                                          |
| `cookingTime`      | `42`                                                                    |
| `totalIngredients` | `null`                                                                  |
| `difficultyLevel`  | `''`                                                                    |
| `averageRating`    | `null`                                                                  |

# --hints--

你应该创建一个名为 `recipe2` 的对象。

```js
assert.isObject(recipe2)
```

你的 `recipe2` 对象应该有一个 `name` 属性。

```js
assert.property(recipe2, "name");
```

你的 `name` 属性应该设置为 `Chicken Curry`。

```js
assert.equal(recipe2.name, "Chicken Curry");
```

你的 `recipe2` 对象应该有一个 `ingredients` 属性。

```js
assert.property(recipe2, "ingredients");
```

你的 `ingredients` 属性应该有一个数组值。

```js
assert.isArray(recipe2.ingredients)
```

你的 `ingredients` 属性的数组值应该包含 `chicken breast`、`coconut milk`、`curry powder`、`onion` 和 `garlic`。

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

你的 `recipe2` 对象应该有一个 `ratings` 属性。

```js
assert.property(recipe2, "ratings");
```

你的 `ratings` 属性应该有一个数组值。

```js
assert.isArray(recipe2.ratings)
```

你的 `ratings` 属性的数组值应该包含 `4`、`5`、`5` 和 `5`。

```js
assert.deepEqual(recipe2.ratings, [4, 5, 5, 5]);
```

你的 `recipe2` 对象应该有一个 `cookingTime` 属性。

```js
assert.property(recipe2, "cookingTime");
```

你的 `cookingTime` 属性值应该是一个数字。

```js
assert.isNumber(recipe2.cookingTime);
```

你的 `cookingTime` 属性应该设置为 `42`。

```js
assert.equal(recipe2.cookingTime, 42);
```

你的 `recipe2` 对象应该有一个 `totalIngredients` 属性。

```js
assert.property(recipe2, "totalIngredients");
```

你的 `totalIngredients` 属性应该设置为 `null`。

```js
assert.isNull(recipe2.totalIngredients);
```

你的 `recipe2` 对象应该有一个 `difficultyLevel` 属性。

```js
assert.property(recipe2, "difficultyLevel");
```

你的 `difficultyLevel` 属性应该设置为空字符串。

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

你的 `recipe2` 对象应该有一个 `averageRating` 属性。

```js
assert.property(recipe2, "averageRating");
```

你的 `averageRating` 属性应该设置为 `null`。

```js
assert.isNull(recipe2.averageRating);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

--fcc-editable-region--

--fcc-editable-region--
```
