---
id: 66fbcf750a62784cf11f5634
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

是时候测试每个函数了。你可以使用任何食谱进行测试，但本教程将从 `recipe1` 开始。

创建三个新变量：`recipe1AverageRating`、`recipe1TotalIngredients` 和 `recipe1DifficultyLevel`。通过调用每个变量对应的函数并将适当的 `recipe1` 属性传递给它们来为它们赋值。

最后，将每个变量记录到控制台以查看结果。

# --hints--

你应该创建一个 `recipe1AverageRating` 变量。

```js
assert.isNotNull(recipe1AverageRating)
```

你的 `recipe1AverageRating` 变量应该设置为 `getAverageRating(recipe1.ratings)`。

```js
assert.equal(recipe1AverageRating, getAverageRating(recipe1.ratings));
```

你应该将 `recipe1AverageRating` 记录到控制台。

```js
assert.match(code, /console\.log\(recipe1AverageRating\);?/)
```

你应该创建一个 `recipe1TotalIngredients` 变量。

```js
assert.isNotNull(recipe1TotalIngredients)
```

你的 `recipe1TotalIngredients` 变量应该设置为 `getTotalIngredients(recipe1.ingredients)`。

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

你应该将 `recipe1TotalIngredients` 记录到控制台。

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/)
```

你应该创建一个 `recipe1DifficultyLevel` 变量。

```js
assert.isNotNull(recipe1DifficultyLevel)
```

你的 `recipe1DifficultyLevel` 变量应该设置为 `getDifficultyLevel(recipe1.cookingTime)`。

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

你应该将 `recipe1DifficultyLevel` 记录到控制台。

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length;
}

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return 'easy';
  } else if (cookingTime <= 60) {
    return 'medium';
  } else {
    return 'hard';
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
