---
id: 66ff02792b3042a2b811f475
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

将以下属性添加到 `recipe1` 对象中：

| 键名                | 值  |
| ------------------ | ------ |
| `cookingTime`      | `22`   |
| `totalIngredients` | `null` |
| `difficultyLevel`  | `""`   |
| `averageRating`    | `null` |

值为 `null` 和空字符串的属性将在你计算它们之后进行更新。

# --hints--

你的 `recipe1` 对象应该有一个 `cookingTime` 属性。

```js
assert.property(recipe1, "cookingTime");
```

你的 `cookingTime` 属性应该设置为 `22`。

```js
assert.equal(recipe1.cookingTime, 22);
```

你的 `recipe1` 对象应该有一个 `totalIngredients` 属性。

```js
assert.property(recipe1, "totalIngredients");
```

你的 `totalIngredients` 属性应该设置为 `null`。

```js
assert.isNull(recipe1.totalIngredients);
```

你的 `recipe1` 对象应该有一个 `difficultyLevel` 属性。

```js
assert.property(recipe1, "difficultyLevel");
```

你的 `difficultyLevel` 属性应该设置为空字符串。

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

你的 `recipe1` 对象应该有一个 `averageRating` 属性。

```js
assert.property(recipe1, "averageRating");
```

你的 `averageRating` 属性应该设置为 `null`。

```js
assert.isNull(recipe1.averageRating);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  ratings: [4, 5, 4, 5],
--fcc-editable-region--

--fcc-editable-region--
};
```
