---
id: 670e4f45f7116c0f216a5177
title: 步骤 5
challengeType: 1
dashedName: step-5
---

# --description--

在继续之前，你应该练习如何从对象中访问属性。

你可以使用点符号（`.`）或括号符号（`[]`）来实现这一点。这是一个示例：

```js
const person = {
  name: 'John',
  age: 30,
  job: 'Software Engineer'
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

访问 `recipe1` 和 `recipe2` 的 `name` 属性，并分别将它们赋值给变量 `recipe1Name` 和 `recipe2Name`。

接下来，访问两个食谱的 `cookingTime` 属性，并分别将它们赋值给变量 `recipe1CookingTime` 和 `recipe2CookingTime`。

确保你创建的所有变量都记录在控制台中。

# --hints--

你应该创建一个 `recipe1Name` 变量。

```js
assert.isDefined(recipe1Name);
```

你应该将 `recipe1` 的 `name` 属性的值赋给你的 `recipe1name` 变量。

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

你应该将 `recipe1Name` 记录到控制台。

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

你应该创建一个 `recipe2Name` 变量。

```js
assert.isDefined(recipe2Name);
```

你应该将 `recipe2` 的 `name` 属性的值赋给你的 `recipe2name` 变量。

```js
assert.strictEqual(recipe2Name, recipe2.name);
```

你应该将 `recipe2Name` 记录到控制台。

```js
assert.match(code, /console\.log\(recipe2Name\)/);
```

你应该创建一个 `recipe1CookingTime` 变量。

```js
assert.isDefined(recipe1CookingTime);
```

你应该将 `recipe1` 的 `cookingTime` 属性的值赋给你的 `recipe1CookingTime` 变量。

```js
assert.strictEqual(recipe1CookingTime, recipe1.cookingTime);
```

你应该将 `recipe1CookingTime` 记录到控制台。

```js
assert.match(code, /console\.log\(recipe1CookingTime\)/);
```

你应该创建一个 `recipe2CookingTime` 变量。

```js
assert.isDefined(recipe2CookingTime);
```

你应该将 `recipe2` 的 `cookingTime` 属性的值赋给你的 `recipe2CookingTime` 变量。

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

你应该将 `recipe2CookingTime` 记录到控制台。

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

--fcc-editable-region--

--fcc-editable-region--
```
