---
id: 685db0bd1280154ae85bcf9f
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊中，你将使用 HTML、CSS 和 JavaScript 构建一个可访问的行星选项卡列表。HTML 模板和 CSS 已经为你提供。

在前几个步骤中，你将构建选项卡列表的其余 HTML。

首先创建一个 `class` 设置为 `"tabs"` 的 `div` 元素。在该 `div` 元素内部，添加一个 `id` 属性设置为 `"tabs-title"` 且文本内容为 `Planets` 的 `h2` 元素。

# --hints--

你应该有一个 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你的 `div` 元素应该有一个设置为 `"tabs"` 的 class。

```js
assert.exists(document.querySelector("div.tabs"));
```

你的 `div` 元素内部应该有一个 `h2` 元素。

```js
assert.exists(document.querySelector(".tabs h2"));
```

你的 `h2` 元素应该有一个设置为 `"tabs-title"` 的 `id` 属性。

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

你的 `h2` 元素应该有文本内容 `Planets`。

```js
assert.equal(document.querySelector(".tabs h2").textContent, "行星");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```