---
id: 641cdefa704f232675ed98aa
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

创建一个名为 `isSpam` 的函数，它接受一个参数。现在让它返回 `false` 吧！

# --hints--

你应该有一个名为 `isSpam` 的函数。

```js
assert.isFunction(isSpam);
```

你的 `isSpam` 函数应该有一个参数。

```js
assert.lengthOf(isSpam, 1);
```

你的 `isSpam` 函数应该返回 `false`。

```js
assert.isFalse(isSpam(""));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
      <form id="message-form">
        <label for="message-input">Input message:</label>
        <input type="text" id="message-input" />
        <button id="check-message-btn" type="button">
          Check message
        </button>
      </form>
      <h2 id="result-message"></h2>
    </header>
    <main>
      <div class="message-container">
        <p>Here are some similar messages:</p>
        <div class="message-list">
          <p>earn money fast</p>
          <p>free money</p>
          <p>click here</p>
          <p>meet singles</p>
        </div>
      </div>
    </main>
    <script src="script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #f5f6f7;
  --dark-blue: #1b1b32;
  --fcc-blue: #0a0a23;
  --white: #ffffff;
  --grey: #ccc;
}

body {
  font-family: "Lato", Arial, sans-serif;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

.main-text {
  text-align: center;
  max-width: 600px;
  margin: 0 auto;
  padding: 10px;
}

.title {
  font-size: 2.5rem;
  margin-bottom: 10px;
}

.description {
  font-size: 1.2rem;
  margin-bottom: 20px;
}

#message-form {
  display: flex;
  flex-direction: column;
  align-items: center;
  gap: 10px;
}

#message-input {
  width: 100%;
  padding: 10px;
  font-size: 1rem;
  border: 1px solid var(--grey);
  border-radius: 5px;
  background-color: var(--dark-blue);
  color: var(--white);
}

#check-message-btn {
  padding: 10px 20px;
  font-size: 1rem;
  background-color: var(--fcc-blue);
  color: var(--white);
  border: none;
  border-radius: 5px;
  cursor: pointer;
}

#check-message-btn:hover {
  background-color: #0056b3;
}

#result-message {
  margin-top: 20px;
  font-size: 1.5rem;
  text-align: center;
}

.message-container {
  background-color: var(--dark-blue);
  border-radius: 5px;
  padding: 20px;
  margin: 20px auto;
  max-width: 600px;
}

.message-list {
  text-align: left;
  margin-top: 10px;
}

.message-list p {
  margin: 5px 0;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");
--fcc-editable-region--

--fcc-editable-region--
```