---
id: 60a3e3396c7b40068ad6997e
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

现在 `.one` 在水平方向已居中，但它的上外边距会越过画布并作用到框架的边框上，使整个画布下移 20px。

为 `.canvas` 添加 `padding: 1px`，这样 `.one` 就有了一个“可推靠”的内侧边界，布局更稳妥。继续，很棒的细节处理！

# --hints--

请把 `padding` 设置为 `1px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

`.canvas` 的 `padding` 应为 `1px`。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
