---
id: 60a3e3396c7b40068ad6996b
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

开始项目：在 `body` 元素中添加一个 `img` 元素，并将它的 `src` 属性设为 `https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png`。小小一步，太棒了！

# --hints--

你的代码应该包含一个 `img` 元素。

```js
assert(document.querySelectorAll('img').length === 1);
```

该 `img` 的 `src` 应设置为 `https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

