---
id: 60a3e3396c7b40068ad6998d
title: 步骤 34
challengeType: 0
dashedName: step-34
---

# --description--

为 `.one` 和 `.two` 创建一个共同的规则，并将它们的模糊（`blur`）增加 1 像素，使画面更柔和。

你做得很棒，继续保持细心！

# --hints--









请创建一个 `.one, .two` 的选择器列表。

```js
const oneTwo = new __helpers.CSSHelp(document).getStyle('.one, .two');
assert(oneTwo);
```

将 `filter` 属性设置为 `blur(1px)`。

```js
const hasFilter = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.filter === 'blur(1px)');
assert(hasFilter)
```

`.one` 的 `filter` 应为 `blur(1px)`。

```js
const one = document.querySelector('.one');
const oneFilter = getComputedStyle(one).filter;
assert(oneFilter === 'blur(1px)');
```

`.two` 的 `filter` 也应为 `blur(1px)`。

```js
const two = document.querySelector('.two');
const twoFilter = getComputedStyle(two).filter;
assert(twoFilter === 'blur(1px)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
  filter: blur(2px);
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: 0 auto 20px;
}

--fcc-editable-region--

--fcc-editable-region--

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
