---
id: 60a3e3396c7b40068ad69974
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

每幅画都需要一个框架（frame）。

把 `.canvas` 元素放入另一个 `div` 中，并给这个 `div` 添加 `class="frame"`，这样画面会更完整。做得好，继续！
# --hints--

请添加一个新的 `div` 元素作为框架。

```js
assert(document.querySelectorAll('div').length === 2)
```

你的 `.canvas` 应嵌套在这个新 `div` 中。

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

新 `div` 的 `class` 应包含 `frame`。

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

新 `div` 应位于 `body` 内。

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
