---
id: 60a3e3396c7b40068ad69970
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在开始为 `div` 添加样式之前，需要在 HTML 中引入 CSS。

添加一个 `link` 元素来链接 `styles.css` 文件：将 `href` 设为 `styles.css`，并把 `rel` 设为 `stylesheet`。小步前进，棒棒的！

# --hints--

你的代码应包含一个 `link` 元素。

```js
assert(/<link/.test(code))
```

该 `link` 的 `rel` 属性应为 `stylesheet`。

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

该 `link` 的 `href` 应为 `styles.css`。

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
