---
id: 60a3e3396c7b40068ad6996c
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在之前关于 CSS 盒模型的课程中，你学习到每个 HTML 元素都被视为带有四个区域的“盒子”。

想象你收到一个包裹：内容物是盒子里的商品，在这里就是你的 header、段落或图片元素。

把页面中 `<img>` 的 `src` 从 `https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png` 改为 `https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-2.png`。做得很棒，继续前进！

# --hints--

`img` 元素的 `src` 应为 `https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-2.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```


