---
id: 60a3e3396c7b40068ad69978
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

使用 margin 调整元素外部的空间。

为 `.frame` 设置 `margin: 20px auto`，这会在垂直方向增加 20px 的外边距，并在水平方向居中框架。做得很好，继续！

# --hints--

请将 `margin` 属性设为 `20px auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

`.frame` 的 `margin` 应为 `20px auto`。

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
