---
id: 60a3e3396c7b40068ad69971
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

开始写 CSS 啦。

尽管你的 `<div>` 内没有文本，它仍然被视为一个有内容的盒子。请为 `.canvas` 类编写规则，并将其 `width` 设置为 `500px`。你很棒，动手试试看！

# --hints--

你的样式表中应包含 `.canvas` 选择器。

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

应该将 `width` 属性设置为 `500px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

`.canvas` 元素的 `width` 应为 `500px`。

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
