---
id: 60a3e3396c7b40068ad69973
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

将画布的 `background-color` 设置为 `#4d0f00`，这会给你的画面一个温暖的底色。继续，很棒的选择！

# --hints--

请将 `background-color` 属性设为 `#4d0f00`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

`.canvas` 的 `background-color` 应为 `#4d0f00`（即 rgb(77,15,0)）。

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
