---
id: 60a3e3396c7b40068ad6996d
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

内容周围有一层称为 padding 的空间，类似于气泡膜把物品和外盒隔开。

把 border 想象成包裹物的纸箱边框。

将 `<img>` 的 `src` 改为 `https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-3.png`，你做得很棒！

# --hints--

`img` 元素的 `src` 应为 `https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-3.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```

