---
id: 60a3e3396c7b40068ad6996f
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `body` 中添加一个 `div` 元素，并将它的 `class` 属性设为 `canvas`。

这个 `div` 就是你的画布（canvas），太棒了，继续构建画面！

# --hints--

你的代码应包含一个 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 1)
```

该 `div` 的 `class` 应包含 `canvas`。

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
