---
id: 60a3e3396c7b40068ad69977
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

使用 padding 来调整元素内部的空白。

在 `.frame` 中使用 `padding` 简写将 `.frame` 与 `.canvas` 之间的间距都增大 `50px`（上下左右都会生效）。太棒了，继续优化布局！

# --hints--

请将 `padding` 属性设置为 `50px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

`.frame` 的 `padding` 应为 `50px`。

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
