---
id: 60a3e3396c7b40068ad69980
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

在 `.one` 元素下面再添加一个 `div`，并把它的 `class` 设置为 `two`，这将成为你的第二个矩形。

# --hints--

不要修改已有的 `.one` 元素。

```js
assert(document.querySelectorAll('.one').length === 1);
```

在 `.canvas` 中应有第二个 `div` 元素。

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

第二个 `div` 的 `class` 应为 `two`。

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

`.two` 元素应位于 `.one` 元素之后。

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
