---
id: 60a3e3396c7b40068ad69976
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

现在画框太宽了。

请在 `.frame` 规则中将 `width` 设置为 `500px`，这样尺寸就合适啦。继续，很棒！

# --hints--

请将 `width` 属性设为 `500px`。

```js
const widthFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.width === '500px');
assert(widthFilter.length === 2);
```

`.frame` 的 `width` 应为 `500px`。

```js
const frameWidth = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('width');
assert(frameWidth === '500px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
