---
id: 60a3e3396c7b40068ad69979
title: 步骤 15
challengeType: 0
dashedName: step-15
---

# --description--

在 `.canvas` 内添加一个新的 `div` 元素，作为第一个矩形区域。

给这个新 `div` 添加 `class="one"`，加油，你正在搭建画面的第一块色块！

# --hints--

你应该创建一个新的 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 3);
```

把这个新 `div` 嵌套在 `.canvas` 中。

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

该新 `div` 的 `class` 应包含 `one`。

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
