---
id: 60a3e3396c7b40068ad69975
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

为 `.frame` 类写一条规则。

使用 `border` 简写将 `.frame` 设置为 `50px solid black`，让画框更有份量。继续，做得很好！

# --hints--

你的样式表应包含 `.frame` 选择器。

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

应将 `border` 设置为 `50px solid black`。

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

`.frame` 的 `border` 应为 `50px solid black`。

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
