---
id: 60a3e3396c7b40068ad6996e
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

外边距（margin）是盒子外部的空间，用来控制元素之间的间距。

在示例中，下方元素的上外边距较大，会把它往页面下方推移。

现在我们快速复习完盒模型，准备开始绘制 Rothko 风格的画作吧。

请从页面中移除 `img` 元素。做得很好，我们继续下一步！

# --hints--

你的代码中不应存在 `img` 元素。

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.MasterPuti.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```

