---
id: 68a23004e96a105df3a21fc4
title: 构建纸牌计数助手
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

在赌场的二十一点游戏中，玩家可以通过跟踪牌组中剩余的高低牌的相对数量来确定他们在下一局中是否比庄家有优势。这被称为纸牌计数。

牌组中剩余更多的高牌对玩家有利。当计数为正数时，玩家应该下大注。当计数为零或负数时，玩家应该下小注。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该使用 `let` 声明一个名为 `count` 的全局变量并将其设置为 `0`。
1. 你应该有一个名为 `cc` 的函数。
1. `cc` 函数应该接收一个 `card` 参数，该参数可以是数字或字符串。
    - 对于 `2` 到 `10` 之间的值，`card` 参数将是一个数字。
    - 对于所有其他值，`card` 参数将是一个字符串。
1. `cc` 函数应该根据特定标准修改全局 `count` 变量。
1. 对于牌 `2`、`3`、`4`、`5` 或 `6`，全局 `count` 变量应增加 `1`。
1. 对于牌 `7`、`8`、`9`，全局 `count` 变量应保持不变。
1. 对于牌 `10`、`"J"`、`"Q"`、`"K"`、`"A"`，全局 `count` 变量应减少 `1`。
1. 当计数为正数时，`cc` 函数应返回一个包含当前计数和字符串 `Bet` 的字符串。
1. 当计数小于或等于 `0` 时，`cc` 函数应返回一个包含当前计数和字符串 `Hold` 的字符串。
1. 在函数输出中，当前计数和玩家决策（`Bet` 或 `Hold`）之间应该用一个空格分隔。例如，`-3 Hold`。

# --hints--

你应该使用 `let` 声明一个名为 `count` 的全局变量并将其设置为 `0`。

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

你应该有一个名为 `cc` 的函数。

```js
assert.isFunction(cc);
```

你的函数应该返回 `count` 的值和文本（`Bet` 或 `Hold`），它们之间用一个空格字符分隔。

```js
count = 0;
let out = cc(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

在打出牌 `2`、`3`、`4`、`5` 后，调用 `cc(6)` 应该返回字符串 `5 Bet`。

```js
count = 0;
cc(2);
cc(3);
cc(4);
cc(5);
const out = cc(6);
assert.strictEqual(out, '5 Bet')
```

在打出牌 `7`、`8` 后，调用 `cc(9)` 应该返回字符串 `0 Hold`。

```js
count = 0;
cc(7);
cc(8);
const out = cc(9);
assert.strictEqual(out, '0 Hold');
```

在打出牌 `10`、`"J"`、`"Q"`、`"K"` 后，调用 `cc("A")` 应该返回字符串 `-5 Hold`。

```js
count = 0;
cc(10);
cc('J');
cc('Q');
cc('K');
const out = cc('A');
assert.strictEqual(out, '-5 Hold');
```

在打出牌 `3`、`7`、`"Q"`、`8` 后，调用 `cc("A")` 应该返回字符串 `-1 Hold`。

```js
count = 0;
cc(3);
cc(7);
cc('Q');
cc(8);
const out = cc('A');
assert.strictEqual(out, '-1 Hold');
```

在打出牌 `2`、`"J"`、`9`、`2` 后，调用 `cc(7)` 应该返回字符串 `1 Bet`。

```js
count = 0;
cc(2);
cc('J');
cc(9);
cc(2);
const out = cc(7);
assert.strictEqual(out, '1 Bet');
```

在打出牌 `2`、`2` 后，调用 `cc(10)` 应该返回字符串 `1 Bet`。

```js
count = 0;
cc(2);
cc(2);
const out = cc(10);
assert.strictEqual(out, '1 Bet')
```

在打出牌 `3`、`2`、`"A"`、`10` 后，调用 `cc("K")` 应该返回字符串 `-1 Hold`。

```js
count = 0;
cc(3);
cc(2);
cc('A');
cc(10);
const out = cc('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```