---
id: 612ea97df5742154772f312e
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

给 `.key.black--key::after` 添加设置为 `absolute` 的 `position` 属性和设置为 `-18px` 的 `left` 属性。

# --hints--

你的 `.key.black--key::after` 选择器应该将 `position` 属性设置为 `absolute`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.position, 'absolute');
```

你的 `.key.black--key::after` 选择器应该将 `left` 属性设置为 `-18px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.left, '-18px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--
.key.black--key::after {
  background-color: #1d1e22;
  content: "";
}
--fcc-editable-region--
```
