---
id: 612e89562043183c86df287c
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

浏览器经常对特定元素应用默认的边距和内边距值。为了确保你的钢琴布局正确显示，你需要重置盒模型。

在你的 CSS 文件中添加一个 `html` 规则选择器，并将 `box-sizing` 属性设置为 `border-box`。

# --hints--

你应该有一个 `html` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

你的 `html` 选择器应该将 `box-sizing` 属性设置为 `border-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
