---
id: 612ebcba99bfa46a15370b11
title: 步骤 25
challengeType: 0
dashedName: step-25
---

# --description--

`@media` at-rule，也称为媒体查询，用于有条件地应用 CSS。媒体查询通常用于根据视口宽度使用 `max-width` 和 `min-width` 属性来应用 CSS。

在下面的示例中，当视口宽度为 `960px` 及以下时，padding 会被应用到 `.card` 类。

```css
@media (max-width: 960px) {
  .card {
    padding: 2rem;
  }
}
```

添加一个媒体查询，当视口宽度为 `768px` 及以下时应用。

# --hints--

你应该添加一个新的 `@media` 查询。

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('media'), 1);
```

你的 `@media` 查询应该有一个 `768px` 的 `max-width`。

```js
assert.equal(new __helpers.CSSHelp(document).getCSSRules('media')[0]?.media?.mediaText, '(max-width: 768px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.MasterPuti.org/platform/universal/fcc_primary.svg" alt="MasterPuti Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
  position: relative;
  border-radius: 10px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
  border-radius: 0 0 3px 3px;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
  border-radius: 0 0 3px 3px;
}

.logo {
  width: 200px;
  position: absolute;
  top: 23px;
}

--fcc-editable-region--

--fcc-editable-region--
```

