---
id: 612ea4c4993aba52ab4aa32e
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

现在是时候使用你之前设置的伪选择器了。要创建黑键，添加一个新的 `.key.black--key::after` 选择器。这会定位同时具有 `key` 和 `black--key` 类的元素，并选择在 HTML 中这些元素后面创建的伪元素。

在新选择器中，将 `background-color` 设置为 `#1d1e22`。同时将 `content` 属性设置为 `""`。这将使伪元素为空。

`content` 属性用于设置或覆盖元素的内容。默认情况下，由 `::before` 和 `::after` 选择器创建的伪元素是空的，这意味着它们不会在页面上渲染。通过将 `content` 属性设置为空字符串 `""`，你可以确保伪元素被渲染，同时仍然显示为空。

如果你想进行实验，可以尝试移除 `background-color` 属性并为 `content` 属性设置不同的值，比如 `"♥"`。记住在完成时撤销这些更改，以便测试通过。

# --hints--

你应该有一个 `.key.black--key::after` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

你的 `.key.black--key::after` 选择器应该将 `background-color` 属性设置为 `#1d1e22`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

你的 `.key.black--key::after` 选择器应该将 `content` 属性设置为 `""`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
