---
id: 612e7d1c29fb872d6384379c
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你现有的 `div` 内嵌套第二个 `div`，并将 `class` 设置为 `keys`。

# --hints--

你应该创建第二个 `div` 元素。

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

你的新 `div` 元素应该在你现有的 `div` 元素内。

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

你的新 `div` 元素应该将 `class` 设置为 `keys`。

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
