---
id: 612e813b3ba67633222cbe54
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

记住，`class` 属性可以包含多个值。为了区分白色键和黑色键，给你的第二个、第三个、第五个、第六个和第七个 `.key` 元素添加第二个 `class` 值 `black--key`。

# --hints--

你的第二个 `.key` 元素也应该有 `black--key` 的 `class`。

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

你的第三个 `.key` 应该在 `class` 中有 `black--key`。

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

你的第五个 `.key` 应该在 `class` 中有 `black--key`。

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

你的第六个 `.key` 应该在 `class` 中有 `black--key`。

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

你的第七个 `.key` 应该在 `class` 中有 `black--key`。

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

你应该有五个 `.black--key` 元素哦～加油！

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

你应该有七个 `.key` 元素呢～太棒了！

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
