---
id: af7588ade1100bde429baf20
title: 构建缺失字母检测器
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个名为 [fearNotLetter](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 的函数。
1. [fearNotLetter](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 函数应该接受一个参数：一个表示按字母顺序排列的字母范围的字符串，该范围可能缺少一个字母。
1. 该函数应该找到传递的字母范围中缺失的字母并返回它。
1. 如果范围中所有字母都存在，该函数应该返回 `undefined`。

# --hints--

你应该有一个 [fearNotLetter](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 函数。

```js
assert.isFunction(fearNotLetter);
```

[fearNotLetter("abce")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 应该返回字符串 `d`。

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

[fearNotLetter("abcdefghjklmno")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 应该返回字符串 `i`。

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

[fearNotLetter("stvwx")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 应该返回字符串 `u`。

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

[fearNotLetter("bcdf")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 应该返回字符串 `e`。

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

[fearNotLetter("abcdefghijklmnopqrstuvwxyz")](file:///e:/1.work/byy/%E5%9C%A8%E7%BA%BF%E8%AF%BE%E7%A8%8B/blocks/blocks_english/lab-missing-letter-detector/af7588ade1100bde429baf20.md#L68-L77) 应该返回 `undefined`。

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```