---
id: 66e7ee20b79186306fc12da5
title: Build a Set of Football Team Cards
challengeType: 25
dashedName: lab-football-team-cards
demoType: onClick
---

# --description--

在这个实验中，你将构建一组足球队卡片。用户应该能够使用下拉菜单并根据球员的位置在不同球员之间进行筛选。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该创建一个具有以下属性的`footballTeam`对象：`team`、`year`、`headCoach`、`players`。
1. `team`属性应该包含一个字符串形式的球队名称。
1. `year`属性应该包含一个数字形式的年份。
1. `headCoach`属性应该包含一个字符串形式的主教练姓名。
1. `players`属性应该是一个包含至少四个元素的数组。
1. `players`数组中的每个元素都应该是一个具有`name`、`position`、`isCaptain`属性的对象。
1. `name`属性应该包含一个字符串形式的球员姓名。
1. `position`属性应该具有以下值之一：`"forward"`、`"midfielder"`、`"defender"`或`"goalkeeper"`。
1. `isCaptain`属性应该具有布尔值。其中一个球员的`isCaptain`属性应该设置为`true`。
1. 你应该在页面上显示`headCoach`、`team`和`year`的值。这些值应该显示在具有`id`值为`head-coach`、`team`和`year`的HTML元素中。
1. 你应该在`#player-cards`元素内显示球员数据，每个球员应该显示在一个具有`player-card`类的`div`中，并在其内部嵌套一个包含球员姓名的`h2`，如果是队长则包含`(Captain)`，以及一个包含`Position:`和球员位置的`p`。

    以下是一个HTML应该如何显示的示例：

    ```html
    <div class="player-card">
      <h2>Sergio Batista</h2>
      <p>Position: midfielder</p>
    </div>
    <div class="player-card">
      <h2>(Captain) Diego Maradona</h2>
      <p>Position: midfielder</p>
    </div>
    ```

1. 当使用下拉菜单选择其中一个位置时，应该只显示该位置球员的`.player-card`元素。如果选择了`"All Players"`选项，则所有球员都应该显示在页面上。

# --hints--

你应该有一个`footballTeam`变量。

```js
assert.exists(footballTeam);
```

`footballTeam`变量应该是一个具有四个属性的对象：`team`、`year`、`headCoach`和`players`。

```js
assert.isObject(footballTeam);
assert.containsAllKeys(footballTeam, ['team', 'year', 'headCoach', 'players']);
```

`team`属性应该是一个字符串。

```js
assert.isString(footballTeam.team);
```

`year`属性应该是一个数字。

```js
assert.isNumber(footballTeam.year);
```

`headCoach`属性应该是一个字符串。

```js
assert.isString(footballTeam.headCoach);
```

`players`属性应该是一个至少包含四个对象的数组，每个对象都应该具有`name`、`position`、`isCaptain`键。

```js
assert.isArray(footballTeam.players);
assert.isAtLeast(footballTeam.players.length, 4);
footballTeam.players.forEach(player => {
  assert.isObject(player);
  assert.containsAllKeys(player, ['name', 'position', 'isCaptain']);
})
```

`name`属性应该具有字符串值。

```js
footballTeam.players.forEach(({name}) => assert.isString(name));
```

`position`属性应该具有`"forward"`、`"midfielder"`、`"defender"`或`"goalkeeper"`值之一。

```js
footballTeam.players.forEach(({position}) => {
  assert.isString(position);
  assert.oneOf(position, ["forward", "midfielder", "defender", "goalkeeper"]);
});
```

`isCaptain`属性应该具有布尔值，并且应该只有一个队长。

```js
footballTeam.players.forEach(({isCaptain}) => assert.isBoolean(isCaptain));
const listOfCaptains = footballTeam.players.filter(({isCaptain}) => isCaptain);
assert.lengthOf(listOfCaptains, 1);
```

你应该在具有`id`值为`head-coach`、`team`和`year`的HTML元素中显示来自`footballTeam`对象的`headCoach`、`team`和`year`值。

```js
const teamElement = document.querySelector('.team-stats #team');
const yearElement = document.querySelector('.team-stats #year');
const headCoachElement = document.querySelector('.team-stats #head-coach');
assert.equal(teamElement?.innerText.trim(), footballTeam.team);
assert.equal(yearElement?.innerText.trim(), footballTeam.year);
assert.equal(headCoachElement?.innerText.trim(), footballTeam.headCoach);
```

当选中`All Players`选项时，所有球员都应该出现在`#player-cards`内。

```js
const select = document.querySelector('#players')
select.value = 'all';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, footballTeam.players);
```

当选中`Position Forward`选项时，只有前锋球员应该出现在`#player-cards`内。

```js
const forwards = footballTeam.players.filter(({position}) => position === 'forward')
const select = document.querySelector('#players')
select.value = 'forward';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, forwards);
```

当选中`Position Midfielder`选项时，只有中场球员应该出现在`#player-cards`内。

```js
const midfielders = footballTeam.players.filter(({position}) => position === 'midfielder')
const select = document.querySelector('#players')
select.value = 'midfielder';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, midfielders);
```

当选中`Position Defender`选项时，只有后卫球员应该出现在`#player-cards`内。

```js
const defenders = footballTeam.players.filter(({position}) => position === 'defender')
const select = document.querySelector('#players')
select.value = 'defender';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, defenders);
```

当选中`Position Goalkeeper`选项时，只有守门员球员应该出现。

```js
const goalkeepers = footballTeam.players.filter(({position}) => position === 'goalkeeper')
const select = document.querySelector('#players')
select.value = 'goalkeeper';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, goalkeepers);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>
      构建一组足球队卡片
    </title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1 class="title">球队统计</h1>
    <main>
      <div class="team-stats">
        <p>球队: <span id="team"></span></p>
        <p>年份: <span id="year"></span></p>
        <p>主教练: <span id="head-coach"></span></p>
      </div>
      <label class="options-label" for="players">筛选队友:</label>
      <select name="players" id="players">
        <option value="all">所有球员</option>
        <option value="forward">位置 前锋</option>
        <option value="midfielder">位置 中场</option>
        <option value="defender">位置 后卫</option>
        <option value="goalkeeper">位置 守门员</option>
      </select>
      <div class="cards" id="player-cards"></div>
    </main>
    <footer>&copy; MasterPuti</footer>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
    flex-direction: column;
  }
}

```

```js
// 你的代码在这里
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>
      构建一组足球队卡片
    </title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1 class="title">球队统计</h1>
    <main>
      <div class="team-stats">
        <p>球队: <span id="team"></span></p>
        <p>年份: <span id="year"></span></p>
        <p>主教练: <span id="head-coach"></span></p>
      </div>
      <label class="options-label" for="players">筛选队友:</label>
      <select name="players" id="players">
        <option value="all">所有球员</option>
        <option value="forward">位置 前锋</option>
        <option value="midfielder">位置 中场</option>
        <option value="defender">位置 后卫</option>
        <option value="goalkeeper">位置 守门员</option>
      </select>
      <div class="cards" id="player-cards"></div>
    </main>
    <footer>&copy; MasterPuti</footer>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
    flex-direction: column;
  }
}

```

```js
const teamName = document.getElementById("team");
const worldCupYear = document.getElementById("year");
const headCoach = document.getElementById("head-coach");
const playerCards = document.getElementById("player-cards");
const playersDropdownList = document.getElementById("players");
const footballTeam = {
  team: "阿根廷",
  year: 1986,
  headCoach: "卡洛斯·比拉尔多",
  players: [
    {
      name: "塞尔吉奥·阿尔米隆",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "塞尔吉奥·巴蒂斯塔",
      position: "midfielder",
      isCaptain: false,
    },
     {
      name: "里卡多·博奇尼",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "克劳迪奥·博尔吉",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "何塞·路易斯·布朗",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "丹尼尔·帕萨雷拉",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "豪尔赫·布尔查加",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "内斯托尔·克劳森",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "何塞·路易斯·库丘夫",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "迭戈·马拉多纳",
      position: "midfielder",
      isCaptain: true,
    },
    {
      name: "豪尔赫·瓦尔多诺",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "埃克托·恩里克",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "奥斯卡·加雷",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "里卡多·吉乌斯蒂",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "路易斯·伊尔萨斯",
      position: "goalkeeper",
      isCaptain: false,
    },
    {
      name: "胡利奥·奥拉尔蒂科切亚",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "佩德罗·帕斯库利",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "内里·蓬皮多",
      position: "goalkeeper",
      isCaptain: false,
    },
    {
      name: "奥斯卡·鲁杰里",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "卡洛斯·塔皮亚",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "马塞洛·特罗比亚尼",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "埃克托·泽拉达",
      position: "goalkeeper",
      isCaptain: false,
    },
  ],
};

const { sport, team, year, players } = footballTeam;
const coachName = footballTeam.headCoach;

teamName.textContent = team;
worldCupYear.textContent = year;
headCoach.textContent = coachName;

const setPlayerCards = (arr = players) => {
  playerCards.innerHTML += arr
    .map(
      ({ name, position, isCaptain }) =>
        `
        <div class="player-card">
          <h2>${isCaptain ? "(队长)" : ""} ${name}</h2>
          <p>位置: ${position}</p>
        </div>
      `
    )
    .join("");
};

playersDropdownList.addEventListener("change", (e) => {
  playerCards.innerHTML = "";

  switch (e.target.value) {
    case "forward":
      setPlayerCards(players.filter((player) => player.position === "forward"));
      break;
    case "midfielder":
      setPlayerCards(
        players.filter((player) => player.position === "midfielder")
      );
      break;
    case "defender":
      setPlayerCards(
        players.filter((player) => player.position === "defender")
      );
      break;
    case "goalkeeper":
      setPlayerCards(
        players.filter((player) => player.position === "goalkeeper")
      );
      break;
    default:
      setPlayerCards();
  }
});

setPlayerCards();
```