---
id: 63c9c09a7daa4f0ff92c4023
title: 第71步
challengeType: 0
dashedName: step-71
---

# --description--

使用相同的语法，查询[#lunch](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9c09a7daa4f0ff92c4023.md#L62-L62)元素中的[number](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c2187f55eb0f400269568f.md#L214-L214)输入框，并将它们分配给[lunchNumberInputs](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9c09a7daa4f0ff92c4023.md#L236-L236)。

# --hints--

你应该声明一个[lunchNumberInputs](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9c09a7daa4f0ff92c4023.md#L236-L236)变量。

```js
assert.match(calculateCalories.toString(), /lunchNumberInputs\s*=/);
```

你应该使用[#lunch input[type='number']](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9c09a7daa4f0ff92c4023.md#L62-L62)调用[document.querySelectorAll()](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L208-L208)。

```js
assert.match(calculateCalories.toString(), /document\.querySelectorAll\(\s*('|"|`)#lunch input\[\s*type\s*=\s*(('|")?)number\2\s*\]\s*\1\s*\)/);
```

你应该将[document.querySelectorAll()](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c21839f56eaf3ef4e027c4.md#L208-L208)调用的结果分配给[lunchNumberInputs](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9c09a7daa4f0ff92c4023.md#L236-L236)。

```js
assert.match(calculateCalories.toString(), /lunchNumberInputs\s*=\s*document\.querySelectorAll\(\s*('|"|`)#lunch input\[\s*type\s*=\s*(('|")?)number\2\s*\]\s*\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>卡路里计算器</title>
  </head>
  <body>
    <main>
      <h1>卡路里计算器</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">预算</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="每日卡路里预算"
            required
          />
          <fieldset id="breakfast">
            <legend>早餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>午餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>晚餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>零食</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>运动</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">添加食物或运动：</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>早餐</option>
                <option value="lunch">午餐</option>
                <option value="dinner">晚餐</option>
                <option value="snacks">零食</option>
                <option value="exercise">运动</option>
              </select>
              <button type="button" id="add-entry">添加条目</button>
            </span>
          </div>
          <div>
            <button type="submit">
              计算剩余卡路里
            </button>
            <button type="button" id="clear">清除</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

function cleanInputString(str) {
  const regex = /[+-\s]/g;
  return str.replace(regex, '');
}

function isInvalidInput(str) {
  const regex = /\d+e\d+/i;
  return str.match(regex);
}

function addEntry() {
  const targetInputContainer = document.querySelector(`#${entryDropdown.value} .input-container`);
  const entryNumber = targetInputContainer.querySelectorAll('input[type="text"]').length + 1;
  const HTMLString = `
  <label for="${entryDropdown.value}-${entryNumber}-name">条目 ${entryNumber} 名称</label>
  <input type="text" id="${entryDropdown.value}-${entryNumber}-name" placeholder="名称" />
  <label for="${entryDropdown.value}-${entryNumber}-calories">条目 ${entryNumber} 卡路里</label>
  <input
    type="number"
    min="0"
    id="${entryDropdown.value}-${entryNumber}-calories"
    placeholder="卡路里"
  />`;
  targetInputContainer.insertAdjacentHTML('beforeend', HTMLString);
}

function calculateCalories(e) {
  e.preventDefault();
  isError = false;

  const breakfastNumberInputs = document.querySelectorAll("#breakfast input[type='number']");
  const lunchNumberInputs = document.querySelectorAll("#lunch input[type='number']");

}

function getCaloriesFromInputs(list) {
  let calories = 0;

  for (const item of list) {
    const currVal = cleanInputString(item.value);
    const invalidInputMatch = isInvalidInput(currVal);

    if (invalidInputMatch) {
      alert(`无效输入: ${invalidInputMatch[0]}`);
      isError = true;
      return null;
    }
    calories += Number(currVal);
  }
  return calories;
}

addEntryButton.addEventListener("click", addEntry);