---
id: 63c9e5eea8261d22856ead1c
title: 第81步
challengeType: 0
dashedName: step-81
---

# --description--

你需要构建将在[output](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e5eea8261d22856ead1c.md#L253-L253)元素中显示的HTML字符串。首先在函数末尾的新行上将空模板字面量分配给[output](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e5eea8261d22856ead1c.md#L253-L253)元素的[innerHTML](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e5eea8261d22856ead1c.md#L253-L253)属性。

# --hints--

你应该访问[output](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e5eea8261d22856ead1c.md#L253-L253)元素的[innerHTML](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e5eea8261d22856ead1c.md#L253-L253)属性。

```js
assert.match(calculateCalories.toString(), /output\s*\.\s*innerHTML/);
```

你应该将空模板字面量分配给[output](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e5eea8261d22856ead1c.md#L253-L253)元素的[innerHTML](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e5eea8261d22856ead1c.md#L253-L253)属性。

```js
// again, template literals don't play well with the parser so we have to look at the raw code.
const htmlString = code.split(/output\s*\.\s*innerHTML\s*=\s*/)[1].split(/\s*;?\s*\}/)[0];
assert.equal(htmlString, '``');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="zh-CN">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>卡路里计算器</title>
  </head>
  <body>
    <main>
      <h1>卡路里计算器</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">预算</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="每日卡路里预算"
            required
          />
          <fieldset id="breakfast">
            <legend>早餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>午餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>晚餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>零食</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>运动</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">添加食物或运动：</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>早餐</option>
                <option value="lunch">午餐</option>
                <option value="dinner">晚餐</option>
                <option value="snacks">零食</option>
                <option value="exercise">运动</option>
              </select>
              <button type="button" id="add-entry">添加条目</button>
            </span>
          </div>
          <div>
            <button type="submit">
              计算剩余卡路里
            </button>
            <button type="button" id="clear">清除</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

function cleanInputString(str) {
  const regex = /[+-\s]/g;
  return str.replace(regex, '');
}

function isInvalidInput(str) {
  const regex = /\d+e\d+/i;
  return str.match(regex);
}

function addEntry() {
  const targetInputContainer = document.querySelector(`#${entryDropdown.value} .input-container`);
  const entryNumber = targetInputContainer.querySelectorAll('input[type="text"]').length + 1;
  const HTMLString = `
  <label for="${entryDropdown.value}-${entryNumber}-name">条目 ${entryNumber} 名称</label>
  <input type="text" id="${entryDropdown.value}-${entryNumber}-name" placeholder="名称" />
  <label for="${entryDropdown.value}-${entryNumber}-calories">条目 ${entryNumber} 卡路里</label>
  <input
    type="number"
    min="0"
    id="${entryDropdown.value}-${entryNumber}-calories"
    placeholder="卡路里"
  />`;
  targetInputContainer.insertAdjacentHTML('beforeend', HTMLString);
}

function calculateCalories(e) {
  e.preventDefault();
  isError = false;

  const breakfastNumberInputs = document.querySelectorAll("#breakfast input[type='number']");
  const lunchNumberInputs = document.querySelectorAll("#lunch input[type='number']");
  const dinnerNumberInputs = document.querySelectorAll("#dinner input[type='number']");
  const snacksNumberInputs = document.querySelectorAll("#snacks input[type='number']");
  const exerciseNumberInputs = document.querySelectorAll("#exercise input[type='number']");

  const breakfastCalories = getCaloriesFromInputs(breakfastNumberInputs);
  const lunchCalories = getCaloriesFromInputs(lunchNumberInputs);
  const dinnerCalories = getCaloriesFromInputs(dinnerNumberInputs);
  const snacksCalories = getCaloriesFromInputs(snacksNumberInputs);
  const exerciseCalories = getCaloriesFromInputs(exerciseNumberInputs);
  const budgetCalories = getCaloriesFromInputs([budgetNumberInput]);

  if (isError) {
    return;
  }

  const consumedCalories = breakfastCalories + lunchCalories + dinnerCalories + snacksCalories;
  const remainingCalories = budgetCalories - consumedCalories + exerciseCalories;
  const surplusOrDeficit = remainingCalories < 0 ? '盈余' : '赤字';

  output.innerHTML = ``;
}

function getCaloriesFromInputs(list) {
  let calories = 0;

  for (const item of list) {
    const currVal = cleanInputString(item.value);
    const invalidInputMatch = isInvalidInput(currVal);

    if (invalidInputMatch) {
      alert(`无效输入: ${invalidInputMatch[0]}`);
      isError = true;
      return null;
    }
    calories += Number(currVal);
  }
  return calories;
}

addEntryButton.addEventListener("click", addEntry);