---
id: 6604080b66ff6e942d8225b1
title: 步骤 36
challengeType: 0
dashedName: step-36
---

# --description--

当你打开控制台时，你应该看到这个结果：

```js
[ '1e3', index: 0, input: '1e3', groups: undefined ]
```

[match](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/6604080b66ff6e942d8225b1.md#L224-L224)方法返回一个包含在字符串中找到的所有匹配项的数组。

以下是该信息的完整分解：

- `"1e3"` 是针对[/\d+e\d+/i](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/6604080b66ff6e942d8225b1.md#L223-L223)正则表达式的匹配值。
- [index: 0](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225) 是字符串中匹配值的索引。
- [input: '1e3'](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/6604080b66ff6e942d8225b1.md#L222-L222) 是被匹配的原始字符串。
- [groups: undefined](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225) 是匹配的组，在这种情况下未使用。你将在以后的项目中了解更多关于组的信息。

现在是时候测试有效的输入了。将你的控制台语句更新为以下内容：[console.log(isInvalidInput("10"))](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225)。

打开控制台查看结果。你将在下一步中了解更多关于这个结果的含义。

# --hints--

在你的[isInvalidInput](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/6604080b66ff6e942d8225b1.md#L222-L224)函数下面应该有一个[console.log](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/63c9e7d5b21eee2776ecc226.md#L225-L225)。

```js
assert.match(code, /console\.log\(/);
```

你的控制台语句应该有[isInvalidInput("10")](file:///f:/code/blocks_Chinese/blocks_english/workshop-calorie-counter/6604080b66ff6e942d8225b1.md#L222-L222)的值。

```js
assert.match(code, /console\.log\(isInvalidInput\(('|")10\s*\1\);?\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>卡路里计算器</title>
  </head>
  <body>
    <main>
      <h1>卡路里计算器</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">预算</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="每日卡路里预算"
            required
          />
          <fieldset id="breakfast">
            <legend>早餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>午餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>晚餐</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>零食</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>运动</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">添加食物或运动：</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>早餐</option>
                <option value="lunch">午餐</option>
                <option value="dinner">晚餐</option>
                <option value="snacks">零食</option>
                <option value="exercise">运动</option>
              </select>
              <button type="button" id="add-entry">添加条目</button>
            </span>
          </div>
          <div>
            <button type="submit">
              计算剩余卡路里
            </button>
            <button type="button" id="clear">清除</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');
const addEntryButton = document.getElementById('add-entry');
const clearButton = document.getElementById('clear');
const output = document.getElementById('output');
let isError = false;

function cleanInputString(str) {
  const regex = /[+-\s]/g;
  return str.replace(regex, '');
}

function isInvalidInput(str) {
  const regex = /\d+e\d+/i;
  return str.match(regex);
}
--fcc-editable-region--
console.log(isInvalidInput("10"));
--fcc-editable-region--
```