---
id: 68507cadbf36aa089a84ce2d
title: 实现Luhn算法
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Luhn算法，也称为"模10"或"模10"算法，是一种简单的校验和公式，用于验证各种识别号码，如信用卡号码。以下是使用Luhn算法验证数字的步骤：

- 从右侧开始，排除最右边的数字（校验位），将每隔一个数字的值加倍。
- 如果将数字加倍的结果大于 `9`，则将数字相加得到个位数。或者，你可以从结果中减去 `9`。
- 计算包括校验位在内的所有数字的总和。
- 如果所有数字的总和是 `10` 的倍数，则该数字有效；否则无效。

假设我们有数字 `453914881`。使用Luhn算法验证它的步骤如下：

 ```md 
 账号                4   5   3   9   1   4   8   8   1 
 隔位加倍            4  10   3  18   1   8   8  16   1 
 相加两位数          4   1   3   9   1   8   8   7   1 
 ```

- 然后将所有数字相加，`4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`。
- 由于 `42` 不是 `10` 的倍数，该数字无效。

在本实验中，你将使用Luhn算法构建一个信用卡验证器。

**目标：** 实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该定义一个名为 `verify_card_number` 的函数，该函数接受一串数字（代表卡号）并根据Luhn算法验证其是否有效。

2. 在 `verify_card_number` 函数内：

   - 你应该处理卡号中可能存在的任何破折号或空格。
   - 如果卡号有效，返回 `VALID!`；否则，返回 `INVALID!`。

当你完成项目时，根据输入你应该看到以下消息：

| 卡号                | 消息     |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |

# --hints--

你应该有一个名为 `verify_card_number` 的函数。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` 应该返回 `VALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` 应该返回 `VALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` 应该返回 `INVALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` 应该返回 `INVALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

当使用有效的信用卡号调用时，`verify_card_number` 应该返回 `VALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

当使用无效的信用卡号调用时，`verify_card_number` 应该返回 `INVALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```