---
id: 615380dff67172357fcf0425
title: æ­¥éª¤ 9
challengeType: 0
dashedName: step-9
---

# --description--

ç°åœ¨ä½ çš„å›¾ç‰‡å¤ªå¤§äº†ã€?

åˆ›å»ºä¸€ä¸?`.gallery img` é€‰æ‹©å™¨æ¥é€‰å–ä½ çš„å›¾ç‰‡ã€?å°†å®ƒä»¬çš„ `width` éƒ½è®¾ç½®ä¸º `100%`ï¼?`max-width` éƒ½è®¾ç½®ä¸º `350px`ã€?

å¹¶å°† `height` å±æ€§è®¾ç½®ä¸º `300px`ï¼Œä»¥ä¿æŒå›¾ç‰‡å¤§å°ä¸€è‡´ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`.gallery img` é€‰æ‹©å™¨ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

ä½ çš„ `.gallery img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`width` å±æ€§ï¼Œå…¶å€¼è®¾ç½®ä¸º `100%`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width, '100%');
```

ä½ çš„ `.gallery img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`max-width` å±æ€§ï¼Œå…¶å€¼è®¾ç½®ä¸º `350px`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

ä½ çš„ `.gallery img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`height` å±æ€§ï¼Œå…¶å€¼è®¾ç½®ä¸º `300px`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```

