---
id: 61537d86bdc3dd343688fceb
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

接下来，根据文档中图片的顺序为每个[img]元素添加[src]属性。第一个[img]的[src]应该是[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg]。其余的图片应该相同，只是将[src]中的[1]替换为该图片在文档中的序号。

# --hints--

你所有的九个[img]元素都应该有[src]属性。

```js
const images = [...document.querySelectorAll('img')];
assert.isTrue(images.every(image => image.getAttribute('src')));
```

你的第一个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[0]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg');
```

你的第二个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[1]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg');
```

你的第三个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[2]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg');
```

你的第四个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[3]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg');
```

你的第五个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[4]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg');
```

你的第六个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[5]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg');
```

你的第七个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[6]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg');
```

你的第八个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[7]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg');
```

你的第九个[img]元素应该将[https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg]设置为[src]属性值。

```js
assert.equal(document.querySelectorAll('img')?.[8]?.getAttribute('src'), 'https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```