---
id: 6153893900438b4643590367
title: æ­¥éª¤11
challengeType: 0
dashedName: step-11
---

# --description--

Align your `.header` text in the center. Make the text uppercase using the `text-transform` property with `uppercase` as the value.

å°†å…¶æ‰€æœ‰æ–¹å‘çš„å†…è¾¹è·éƒ½è®¾ç½®ä¸?`32px`ã€?å°†èƒŒæ™¯é¢œè‰²è®¾ç½®ä¸º `#0a0a23`ï¼Œæ–‡æœ¬é¢œè‰²è®¾ç½®ä¸º `#fff` ã€?

æ·»åŠ  `border-bottom`ï¼Œå€¼ä¸º `4px solid #fdb347`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`.header` é€‰æ‹©å™¨ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

`.header` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `center` çš?`text-align` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

`.header` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `uppercase` çš?`text-transform` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

`.header` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `32px` çš?`padding` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

`.header` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `#0a0a23` çš?`background-color` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

`.header` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `#fff` çš?`color` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

`.header` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `4px solid #fdb347` çš?`border-bottom` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```

