---
id: 61537bb9b1a29430ac15ad38
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

通过将你的[styles.css]文件链接到页面来开始这个项目。

# --hints--

你的[link]元素应该有一个[href]属性，其值为[styles.css]。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```