---
id: 6153908a366afb4d57185c8d
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

Flexbox有主轴和交叉轴。主轴由[flex-direction]属性定义，它有四个可能的值：

- [row]（默认值）：水平轴，flex项目从左到右
- [row-reverse]：水平轴，flex项目从右到左
- [column]：垂直轴，flex项目从上到下
- [column-reverse]：垂直轴，flex项目从下到上

**注意**：轴线和方向会根据文本方向的不同而不同。所示的值适用于从左到右的文本方向。

尝试不同的值来查看它们如何影响布局。

完成后，在[.gallery]元素上设置一个明确的[flex-direction]值为[row]。

# --hints--

你的[.gallery]选择器应该有一个[flex-direction]属性，其值设置为[row]。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```