---
id: 6153a04847abee57a3a406ac
title: æ­¥éª¤ 19
challengeType: 0
dashedName: step-19
---

# --description--

Your images need some space between them.

`gap` CSS ç¼©å†™å±æ€§è®¾ç½®è¡Œå’Œåˆ—ä¹‹é—´çš„é—´éš™ï¼Œä¹Ÿè¢«ç§°ä¸ºç½‘æ ¼é—´è·ï¼ˆgutterï¼‰ã€?`gap` å±æ€§ä»¥å?`row-gap` å’?`column-gap` å­å±æ€§ç”¨æ¥è®¾ç½?flexã€grid å’Œå¤šåˆ—å¸ƒå±€çš„é—´éš™ã€?å¯ä»¥å°†æ­¤å±æ€§åº”ç”¨åˆ°å®¹å™¨å…ƒç´ ã€?

ä¸?`.gallery` flex å®¹å™¨æ·»åŠ ä¸€ä¸?`gap` å±æ€§ï¼Œå€¼è®¾ç½®ä¸º `16px`ã€?

# --hints--

`.gallery` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`gap` å±æ€§ï¼Œå…¶å€¼ä¸º `16px`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.gap, '16px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
}
```

