---
id: 615f171d05def3218035dc85
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

移除[body]元素的margin，将[font-family]设置为[sans-serif]，并将其[background-color]的值设为[#f5f6f7]。

# --hints--

你应该有一个[body]选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你的[body]选择器应该有一个[margin]属性，其值设置为0。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginTop, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginBottom, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginLeft, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginRight, '0px');
```

你的[body]选择器应该有一个[font-family]属性，其值设置为[sans-serif]。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily, 'sans-serif');
```

你的[body]选择器应该有一个[background-color]属性，其值设置为[#f5f6f7]。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}

```