---
id: 6493bc0d99879635209565aa
title: æ­¥éª¤7
challengeType: 0
dashedName: step-7
---

# --description--

`border-box` æ¨¡å‹ä¸?`content-box` çš„ä½œç”¨ç›¸åã€?å…ƒç´ çš„æ€»å®½åº¦ï¼ˆåŒ…æ‹¬è¾¹è·å’Œè¾¹æ¡†ï¼‰å°†æ˜¯æ˜ç¡®è®¾ç½®çš„å®½åº¦ã€?å…ƒç´ çš„å†…å®¹ä¼šç¼©å°ï¼Œä»¥ä¾¿ä¸ºå†…è¾¹è·å’Œè¾¹æ¡†ç•™ä¸‹ç©ºé—´ã€?

å°?`box-sizing` å±æ€§æ›´æ”¹ä¸º `border-box`ã€?æ³¨æ„ä½ çš„è“è‰²å›¾ç‰‡è¾¹æ¡†ç°åœ¨å¦‚ä½•é€‚åº”åœ¨ä½ çš„çº¢è‰²ç”»å»Šè¾¹æ¡†å†…çš„ã€?

# --hints--


ä½ çš„ * é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`box-sizing` å±æ€§ï¼Œå…¶å€¼è®¾ç½®ä¸º `border-box`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>


```

```css
--fcc-editable-region--
 * {
  box-sizing: content-box;
}
--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```

