---
id: 6494da0daf5df5197963671d
title: æ­¥éª¤6
challengeType: 0
dashedName: step-6
---

# --description--

æ³¨æ„åˆ°è“è‰²çš„å›¾ç‰‡è¾¹æ¡†ä¼šè¶…å‡ºçº¢è‰²çš„ç”»å»Šè¾¹æ¡†ã€?è¿™æ˜¯å› ä¸ºæµè§ˆå™¨è®¡ç®—å®¹å™¨å…ƒç´ å¤§å°çš„æ–¹å¼æ‰€è‡´ã€?

`box-sizing` å±æ€§ç”¨äºè®¾ç½®è¿™ç§è¡Œä¸ºã€?é»˜è®¤æƒ…å†µä¸‹ï¼Œä½¿ç”¨çš„æ˜¯ `content-box` æ¨¡å‹ã€?åœ¨è¿™ä¸ªé»˜è®¤æ¨¡å‹ä¸­ï¼Œå½“æŸä¸ªå…ƒç´ æœ‰ä¸€ä¸ªç‰¹å®šçš„å®½åº¦æ—¶ï¼Œè¯¥å®½åº¦ä»…åŸºäºå…ƒç´ çš„å†…å®¹æ¥è®¡ç®—ã€?å†…è¾¹è·å’Œè¾¹æ¡†å€¼ä¼šè¢«åŠ åˆ°æ€»å®½åº¦ä¸­ï¼Œå› æ­¤å…ƒç´ å®½åº¦ä¼šå¢é•¿ä»¥é€‚åº”è¿™äº›å€¼ã€?

å°è¯•ä½¿ç”¨å…¨å±€ `*` é€‰æ‹©å™¨ï¼Œå°?`box-sizing` è®¾ç½®ä¸?`content-box`ã€?æ­¤æ—¶ï¼Œä½ ä¸ä¼šçœ‹åˆ°ä»»ä½•å˜åŒ–ï¼Œå› ä¸ºä½ æ­£åœ¨ä½¿ç”¨é»˜è®¤å€¼ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`*` é€‰æ‹©å™¨ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

ä½ çš„ `*` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`box-sizing` å±æ€§ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

ä½ çš„ `*` é€‰æ‹©å™¨çš„ `box-sizing` å±æ€§åº”è¯¥è®¾ç½®ä¸º `content-box`ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```

