---
id: 64dd076a1b33c86d84b4232f
title: æ­¥éª¤5
challengeType: 0
dashedName: step-5
---

# --description--

ä¸ºäº†æ›´å¥½åœ°å¯è§†åŒ–ä½ å…ƒç´ çš„å°ºå¯¸ï¼Œæ·»åŠ è¾¹æ¡†æœƒå¾ˆæœ‰å¸®åŠ©ã€?

å°†ä½ çš?`.gallery` å…ƒç´ çš?`width` è®¾ç½®ä¸?`50%`ï¼Œå°† `border` è®¾ç½®ä¸?`5px solid red`ã€?

ç„¶åï¼Œå°† `img` å…ƒç´ çš?`width` è®¾ç½®ä¸?`100%`ï¼Œå°† `padding` è®¾ç½®ä¸?`5px`ï¼Œå°† `border` è®¾ç½®ä¸?`5px solid blue`ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`.gallery` é€‰æ‹©å™¨ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

ä½ çš„ `.gallery` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`width` å±æ€§ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

ä½ çš„ `.gallery` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `50%` çš?`width` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

ä½ çš„ `.gallery` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`border` å±æ€§ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

ä½ çš„ `.gallery` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `5px solid red` çš?`border` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

ä½ åº”è¯¥æœ‰ä¸€ä¸?`img` é€‰æ‹©å™¨ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

ä½ çš„ `img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`width` å±æ€§ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

ä½ çš„ `img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `100%` çš?`width` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

ä½ çš„ `img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`padding` å±æ€§ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

ä½ çš„ `img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `5px` çš?`padding` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

ä½ çš„ `img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸?`border` å±æ€§ã€?

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

ä½ çš„ `img` é€‰æ‹©å™¨åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `5px solid blue` çš?`border` å±æ€§ã€?

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

