---
id: 6153a3ebb4f7f05b8401b716
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

[alt]图片属性应该描述图片内容。屏幕阅读器会代替图片宣布替代文本。如果图片无法加载，此文本将代替图片显示。

为了完成项目，给你的所有九张猫图片添加[alt]属性来描述它们。每个描述使用至少五个字符的值。

# --hints--

你所有的九个[img]元素都应该有一个[alt]属性，其中包含描述每张图片的文本，每个描述使用至少[5]个字母。

```js
const images = [...document.querySelectorAll('img')];
images.forEach((image) => {
  const alt = image.getAttribute('alt');
  assert.isAtLeast(alt?.length, 5);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      --fcc-editable-region--
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  gap: 16px;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
  border-radius: 10px;
}

.gallery::after {
  content: "";
  width: 350px;
}

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg" alt="Gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg" alt="Gallery">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  gap: 16px;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
  border-radius: 10px;
}

.gallery::after {
  content: "";
  width: 350px;
}
```