---
id: 61539e07e7430b528fbffe21
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

给你的[.gallery]选择器添加一个[padding]属性，其值设置为[20px 10px]，在容器周围创建一些空间。

然后，给它一个[1400px]的[max-width]，并添加一个值为[0 auto]的[margin]使其居中。

# --hints--

你的[.gallery]选择器应该有一个[padding]属性，其值设置为[20px 10px]。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingTop, '20px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingBottom, '20px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingLeft, '10px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingRight, '10px');
```

你的[.gallery]选择器应该有一个[max-width]属性，其值设置为[1400px]。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.maxWidth, '1400px');
```

你的[.gallery]元素应该使用值为[0 auto]的[margin]来居中。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('.gallery').cssText, 'margin: 0px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.MasterPuti.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```