---
id: 61537c5f81f0cf325b4a854c
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在[body]元素内添加一个[header]元素，并为其分配一个值为[header]的[class]。

在[header]内创建一个[h1]元素，文本内容为[css flexbox photo gallery]。

# --hints--

你的[body]元素内应该有一个[header]元素。

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

你的[header]元素应该有一个值为[header]的[class]。

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

你的[header]元素内应该有一个[h1]元素。

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

你的[h1]元素内应该有文本[css flexbox photo gallery]。

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```