---
id: a5deed1811a43193f9f1c841
title: 实现元素跳过器
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

在这个实验中，你将创建一个函数，该函数跳过数组中的元素，直到根据特定的测试函数找到一个可接受的元素。

例如，对于像`[1, 1, 1, 2, 1, 1, 1]`这样的数组和测试函数`function(n) {return n === 2}`，第一个可接受的元素是索引3处的元素，所以在此之前的所有元素都需要被丢弃，输出应该是剩余的元素`[2, 1, 1, 1]`。

**目标**：实现以下用户需求并通过所有测试以完成实验。

**用户需求**

1. 你应该有一个`dropElements`函数，它接受一个数组(`arr`)和一个函数(`func`)作为参数。
1. `dropElements`函数应该遍历数组，从第一个元素开始删除元素，直到`func`为某个元素返回`true`。
1. `dropElements`函数应该返回数组中剩余的元素（如果满足条件）。
1. 如果条件从未满足，它应该返回一个空数组。

# --hints--

你应该有一个`dropElements`函数。

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})`应该返回`[3, 4]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})`应该返回`[1, 0, 1]`。

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})`应该返回`[1, 2, 3]`。

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})`应该返回`[]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})`应该返回`[7, 4]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})`应该返回`[3, 9, 2]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```