---
id: 66da326c02141df538f29ba5
title: åˆ›å»ºä¸€ä¸ªç»“è´¦é¡µé?
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**ç›®æ ‡ï¼?*å®ç°ä»¥ä¸‹ç”¨æˆ·éœ€æ±‚å¹¶é€šè¿‡æ‰€æœ‰æµ‹è¯•ä»¥å®Œæˆå®éªŒã€?

**ç”¨æˆ·éœ€æ±‚ï¼š**

1. ä½ åº”è¯¥æœ‰ä¸€ä¸?`h1` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Checkout` ã€?
1. `h1` å…ƒç´ ä¹‹ååº”ç´§è·Ÿä¸¤ä¸?`section` å…ƒç´ ã€?
1. åœ¨ç¬¬ä¸€éƒ¨åˆ†ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸?`h2` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Your Cart`ã€?
1. åº”åœ¨ç¬¬ä¸€éƒ¨åˆ†æä¾›é¡¹ç›®å›¾ç‰‡ï¼Œå¹¶é…ä»¥é€‚å½“çš„å¤‡ç”¨æ–‡å­—ã€?å¦‚æœä½ æ„¿æ„ï¼Œå¯ä»¥ä½¿ç”¨è¿™å¼ å›¾ç‰‡ï¼š`https://cdn.MasterPuti.org/curriculum/labs/cube.jpg`
1. åœ¨ç¬¬äºŒéƒ¨åˆ†ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸?`h2` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Payment Information`ã€?
1. åœ¨ç¬¬äºŒéƒ¨åˆ†ä¸­ï¼Œåº”æœ‰ä¸€ä¸?`form` å…ƒç´ ã€?
1. åœ¨è¡¨å•ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸ªè¾“å…¥å…ƒç´ ï¼Œå…?`id` å’?`name` å±æ€§å‡ä¸?`card-name`ï¼Œ`type` ä¸?`text`ï¼Œå¹¶ä¸”æœ‰ä¸ä¹‹å…³è”çš?`label`ã€?
1. åœ¨è¡¨å•ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸ªè¾“å…¥å…ƒç´ ï¼Œå…?`id` å’?`name` å±æ€§å‡ä¸?`card-number`ï¼Œ`type` ä¸?`text`ï¼Œå¹¶ä¸”æœ‰ä¸ä¹‹å…³è”çš?`label`ã€?
1. ä½ è‡³å°‘åº”åŒ…å«ä¸¤ä¸ªå¸¦æœ‰  `required` å±æ€§çš„ `input`å…ƒç´ ã€?
1. ä½ åœ¨æ¯ä¸ªå¿…å¡«è¾“å…¥çš?`label` å…ƒç´ å†…ï¼Œéƒ½åº”è¯¥æœ‰ä¸€ä¸?`span` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º  `*`ï¼Œ`aria-hidden` è®¾ç½®ä¸?`true`ï¼Œä»¥ä¾¿å¿…å¡«å­—æ®µèƒ½é€šè¿‡è§†è§‰æ–¹å¼æ ‡æ˜ã€?
1. ä½ åº”åœ¨å¡å·è¾“å…¥æ¡†çš„æ­£åæ–¹æ”¾ç½®ä¸€ä¸?`p` å…ƒç´ ï¼Œå…¶ä¸­åŒ…å«è¯´æ˜æ‰€éœ€å¡å·æ ¼å¼çš„å¸®åŠ©æ–‡æœ¬ã€?`p` åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `card-number-help` çš?`id`ï¼Œå¹¶è¢«å¡å·è¾“å…¥æ¡†ä½¿ç”¨ `aria-describedby` è¿›è¡Œå¼•ç”¨ã€?

# --hints--

ä½ åº”è¯¥æœ‰ä¸€ä¸?`h1` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Checkout` ã€?

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

ä½ çš„é¡µé¢ä¸Šåº”è¯¥åªæœ‰ä¸€ä¸?`h1` å…ƒç´ ã€?

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

åœ?`h1` å…ƒç´ ä¹‹åï¼Œä½ åº”è‡³å°‘åŒ…å«ä¸¤ä¸?`section` å…ƒç´ ã€?

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

åœ¨ç¬¬ä¸€éƒ¨åˆ†ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸?`h2` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Your Cart`ã€?

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

åœ¨ç¬¬ä¸€éƒ¨åˆ†ä¸­ï¼Œåº”è¯¥è‡³å°‘æœ‰ä¸€å¼ å¸¦æœ‰æ›¿ä»£æ–‡æœ¬çš„å›¾ç‰‡ã€?

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

åœ¨ç¬¬äºŒéƒ¨åˆ†ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸?`h2` å…ƒç´ ï¼Œæ–‡æœ¬ä¸º `Payment Information`ã€?

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

åœ¨ç¬¬äºŒä¸ª `section` ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸?`form` å…ƒç´ ã€?

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

åœ¨è¡¨å•ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸?`input` å…ƒç´ ï¼Œå…¶ `id` å’?`name` ä¸?`card-name`ã€?

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

ä½ çš„å¡ç‰‡åç§°è¾“å…¥æ¡†åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `text` çš?`type` å±æ€§ã€?

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

åœ¨è¡¨å•ä¸­ï¼Œåº”åŒ…å«ä¸€ä¸?`input` å…ƒç´ ï¼Œå…¶ `id` å’?`name` ä¸?`card-number`ã€?

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

ä½ çš„å¡å·è¾“å…¥æ¡†åº”è¯¥æœ‰ä¸€ä¸ªå€¼ä¸º `text` çš?`type` å±æ€§ã€?

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

æ‰€æœ‰é `submit` `type` çš?`input` å…ƒç´ ï¼Œéƒ½åº”æœ‰ä¸€ä¸ªä¸ä¹‹å…³è”çš„ `label` å…ƒç´ ã€?

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

ä½ è‡³å°‘åº”åŒ…å«ä¸¤ä¸ªå¸¦æœ‰ `required` å±æ€§çš„ `input` å…ƒç´ ã€?

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

å¯¹äºæ¯ä¸ªå¿…å¡«è¾“å…¥é¡¹ï¼Œåº”åœ¨ `label` å…ƒç´ å†…åŒ…å«ä¸€ä¸ªå¸¦æœ‰æ–‡æœ?`*` çš?`span` å…ƒç´ ã€?

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`);
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

ä½ çš„ `span` å…ƒç´ åº”å°† `aria-hidden` å±æ€§è®¾ç½®ä¸º `true`ã€?

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`);
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

åœ¨å¡å·è¾“å…¥æ¡†çš„æ­£ä¸‹æ–¹ï¼Œåº”æ”¾ç½®ä¸€ä¸ªå…·æœ?`id="card-number-help"` å±æ€§çš„ `p` å…ƒç´ ã€?

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
assert.equal(cardNumberInput.nextElementSibling, cardNumberHelp);
```

ä½ çš„å¡å·å¸®åŠ©æ–‡æœ¬ä¸åº”ä¸ºç©ºã€?

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

ä½ çš„å¡å·è¾“å…¥æ¡†åº”è®¾ç½® `aria-describedby` å±æ€§ä¸º `card-number-help`ã€?

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.MasterPuti.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```

